Elaborating (= (distinct (* (* 2 x) y) (* x (* 2 y))) (not (= (* (* 2 x) y) (* x (* 2 y)))))
Elaboration succeeded
Elaborating (= (* (* 2 x) y) (* 2 (* x y)))
Elaboration failed: At 8:5 of carcara/src/rare/meta.rs
Check failed: 
(= (goal_lhs ) (goal_rhs ))
Elaborating (= (* x (* 2 y)) (* 2 (* x y)))
Elaboration succeeded
Elaborating (= (= (* 2 (* x y)) (* 2 (* x y))) true)
Elaboration succeeded
Elaborating (= (not true) false)
Elaboration succeeded
holey
(step t0 (cl (not (= (distinct (* (* 2 x) y) (* x (* 2 y))) false)) (not (distinct (* (* 2 x) y) (* x (* 2 y)))) false) :rule equiv_pos2)
(step t1.t1 (cl (= (distinct (* (* 2 x) y) (* x (* 2 y))) (not (= (* (* 2 x) y) (* x (* 2 y)))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (distinct (* (* 2 x) y) (* x (* 2 y))) (not (= (* (* 2 x) y) (* x (* 2 y))))) 0 6))
(step t2.t1 (cl (= (* (* 2 x) y) (* 2 (* x y)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (* (* 2 x) y) (* 2 (* x y))) 3 7))
(step t3.t1 (cl (= (* x (* 2 y)) (* 2 (* x y)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (* x (* 2 y)) (* 2 (* x y))) 3 7))
(step t4 (cl (= (= (* (* 2 x) y) (* x (* 2 y))) (= (* 2 (* x y)) (* 2 (* x y))))) :rule cong :premises (t2.t1 t3.t1))
(step t5.t1 (cl (= (= (* 2 (* x y)) (* 2 (* x y))) true)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= (* 2 (* x y)) (* 2 (* x y))) true) 3 7))
(step t6 (cl (= (= (* (* 2 x) y) (* x (* 2 y))) true)) :rule trans :premises (t4 t5.t1))
(step t7 (cl (= (not (= (* (* 2 x) y) (* x (* 2 y)))) (not true))) :rule cong :premises (t6))
(step t8.t1 (cl (= (not true) false)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (not true) false) 1 7))
(step t9 (cl (= (not (= (* (* 2 x) y) (* x (* 2 y)))) false)) :rule trans :premises (t7 t8.t1))
(step t10 (cl (= (distinct (* (* 2 x) y) (* x (* 2 y))) false)) :rule trans :premises (t1.t1 t9))
(assume a1 (distinct (* (* 2 x) y) (* x (* 2 y))))
(step t11 (cl false) :rule resolution :premises (t0 t10 a1) :args ((= (distinct (* (* 2 x) y) (* x (* 2 y))) false) false (distinct (* (* 2 x) y) (* x (* 2 y))) false))
(step t12 (cl (not false)) :rule false)
(step t13 (cl) :rule resolution :premises (t11 t12) :args (false true))
