
(define-fun @quantifiers_skolemize_1 () S2 (choice ((?v0 S2)) (not (forall ((?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))))
(define-fun @quantifiers_skolemize_2 () S3 (choice ((?v1 S3)) (not (forall ((BOUND_VARIABLE_8839 S3)) (or (not (= (f3 (f4 f5 @quantifiers_skolemize_1) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 @quantifiers_skolemize_1) BOUND_VARIABLE_8839)))))))
(define-fun @quantifiers_skolemize_3 () S3 (choice ((BOUND_VARIABLE_8839 S3)) (not (or (not (= (f3 (f4 f5 @quantifiers_skolemize_1) @quantifiers_skolemize_2) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) @quantifiers_skolemize_2) f13) BOUND_VARIABLE_8839))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 @quantifiers_skolemize_1) BOUND_VARIABLE_8839))))))
(assume a0 (not (= f1 f2)))
(assume a1 (! (not (! (=> (forall ((?v0 S2) (?v1 S3)) (=> (= (f3 (f4 f5 ?v0) ?v1) f1) (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))))) (forall ((?v0 S2) (?v1 S3)) (=> (= (f3 (f4 f5 ?v0) ?v1) f1) (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1)))))) :named @p_1)) :named @p_2))
(assume a2 (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (=> (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) f1) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) f1))))
(assume a3 (forall ((?v0 S7)) (not (= ?v0 (f11 f12 ?v0)))))
(assume a4 (forall ((?v0 S7)) (not (= (f11 f12 ?v0) ?v0))))
(assume a5 (forall ((?v0 S7) (?v1 S7)) (= (= (f11 f12 ?v0) (f11 f12 ?v1)) (= ?v0 ?v1))))
(assume a6 (forall ((?v0 S7) (?v1 S7)) (=> (= (f11 f12 ?v0) (f11 f12 ?v1)) (= ?v0 ?v1))))
(assume a7 (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3) (?v4 S10) (?v5 S3) (?v6 S7) (?v7 S3)) (=> (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) f1) (=> (= (f3 (f6 (f7 (f8 f9 ?v4) ?v5) ?v6) ?v7) f1) (exists ((?v8 S7)) (and (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v8) ?v3) f1) (= (f3 (f6 (f7 (f8 f9 ?v4) ?v5) ?v8) ?v7) f1)))))))
(assume a8 (forall ((?v0 S3) (?v1 S7)) (= (f3 (f6 (f7 (f8 f9 f15) ?v0) ?v1) ?v0) f1)))
(assume a9 (forall ((?v0 S3) (?v1 S7) (?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f15) ?v0) ?v1) ?v2) f1) (=> (=> (= ?v2 ?v0) false) false))))
(assume a10 (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3) (?v4 S10) (?v5 S3)) (=> (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) f1) (=> (= (f3 (f6 (f7 (f8 f9 ?v4) ?v3) ?v2) ?v5) f1) (= (f3 (f6 (f7 (f8 f9 (f16 (f17 f18 ?v0) ?v4)) ?v1) ?v2) ?v5) f1)))))
(assume a11 (forall ((?v0 S10) (?v1 S10) (?v2 S3) (?v3 S7) (?v4 S3)) (=> (= (f3 (f6 (f7 (f8 f9 (f16 (f17 f18 ?v0) ?v1)) ?v2) ?v3) ?v4) f1) (=> (forall ((?v5 S3)) (=> (= (f3 (f6 (f7 (f8 f9 ?v0) ?v2) ?v3) ?v5) f1) (=> (= (f3 (f6 (f7 (f8 f9 ?v1) ?v5) ?v3) ?v4) f1) false))) false))))
(assume a12 (forall ((?v0 S10) (?v1 S10)) (not (= f15 (f16 (f17 f18 ?v0) ?v1)))))
(assume a13 (forall ((?v0 S10) (?v1 S10)) (not (= (f16 (f17 f18 ?v0) ?v1) f15))))
(assume a14 (forall ((?v0 S10) (?v1 S10) (?v2 S10) (?v3 S10)) (= (= (f16 (f17 f18 ?v0) ?v1) (f16 (f17 f18 ?v2) ?v3)) (and (= ?v0 ?v2) (= ?v1 ?v3)))))
(assume a15 (forall ((?v0 S4) (?v1 S10) (?v2 S10) (?v3 S3) (?v4 S7) (?v5 S3)) (=> (= (f3 (f6 (f7 (f8 f9 (f16 (f17 (f19 f20 ?v0) ?v1) ?v2)) ?v3) ?v4) ?v5) f1) (=> (=> (= (f3 ?v0 ?v3) f1) (=> (= (f3 (f6 (f7 (f8 f9 ?v1) ?v3) ?v4) ?v5) f1) false)) (=> (=> (not (= (f3 ?v0 ?v3) f1)) (=> (= (f3 (f6 (f7 (f8 f9 ?v2) ?v3) ?v4) ?v5) f1) false)) false)))))
(assume a16 (forall ((?v0 S4) (?v1 S3) (?v2 S10) (?v3 S7) (?v4 S3) (?v5 S10)) (=> (= (f3 ?v0 ?v1) f1) (=> (= (f3 (f6 (f7 (f8 f9 ?v2) ?v1) ?v3) ?v4) f1) (= (f3 (f6 (f7 (f8 f9 (f16 (f17 (f19 f20 ?v0) ?v2) ?v5)) ?v1) ?v3) ?v4) f1)))))
(assume a17 (forall ((?v0 S4) (?v1 S3) (?v2 S10) (?v3 S7) (?v4 S3) (?v5 S10)) (=> (not (= (f3 ?v0 ?v1) f1)) (=> (= (f3 (f6 (f7 (f8 f9 ?v2) ?v1) ?v3) ?v4) f1) (= (f3 (f6 (f7 (f8 f9 (f16 (f17 (f19 f20 ?v0) ?v5) ?v2)) ?v1) ?v3) ?v4) f1)))))
(assume a18 (forall ((?v0 S4) (?v1 S3) (?v2 S10) (?v3 S7) (?v4 S3) (?v5 S3)) (=> (= (f3 ?v0 ?v1) f1) (=> (= (f3 (f6 (f7 (f8 f9 ?v2) ?v1) ?v3) ?v4) f1) (=> (= (f3 (f6 (f7 (f8 f9 (f16 (f21 f22 ?v0) ?v2)) ?v4) ?v3) ?v5) f1) (= (f3 (f6 (f7 (f8 f9 (f16 (f21 f22 ?v0) ?v2)) ?v1) ?v3) ?v5) f1))))))
(assume a19 (forall ((?v0 S4) (?v1 S3) (?v2 S10) (?v3 S7)) (=> (not (= (f3 ?v0 ?v1) f1)) (= (f3 (f6 (f7 (f8 f9 (f16 (f21 f22 ?v0) ?v2)) ?v1) ?v3) ?v1) f1))))
(assume a20 (forall ((?v0 S4) (?v1 S10) (?v2 S10) (?v3 S4) (?v4 S10) (?v5 S10)) (= (= (f16 (f17 (f19 f20 ?v0) ?v1) ?v2) (f16 (f17 (f19 f20 ?v3) ?v4) ?v5)) (and (= ?v0 ?v3) (and (= ?v1 ?v4) (= ?v2 ?v5))))))
(assume a21 (forall ((?v0 S4) (?v1 S10) (?v2 S4) (?v3 S10)) (= (= (f16 (f21 f22 ?v0) ?v1) (f16 (f21 f22 ?v2) ?v3)) (and (= ?v0 ?v2) (= ?v1 ?v3)))))
(assume a22 (forall ((?v0 S4) (?v1 S10) (?v2 S10) (?v3 S4) (?v4 S10)) (not (= (f16 (f17 (f19 f20 ?v0) ?v1) ?v2) (f16 (f21 f22 ?v3) ?v4)))))
(assume a23 (forall ((?v0 S4) (?v1 S10) (?v2 S4) (?v3 S10) (?v4 S10)) (not (= (f16 (f21 f22 ?v0) ?v1) (f16 (f17 (f19 f20 ?v2) ?v3) ?v4)))))
(assume a24 (forall ((?v0 S4) (?v1 S10) (?v2 S10) (?v3 S10)) (not (= (f16 (f21 f22 ?v0) ?v1) (f16 (f17 f18 ?v2) ?v3)))))
(assume a25 (forall ((?v0 S10) (?v1 S10) (?v2 S4) (?v3 S10)) (not (= (f16 (f17 f18 ?v0) ?v1) (f16 (f21 f22 ?v2) ?v3)))))
(assume a26 (forall ((?v0 S4) (?v1 S10)) (not (= (f16 (f21 f22 ?v0) ?v1) f15))))
(assume a27 (forall ((?v0 S4) (?v1 S10)) (not (= f15 (f16 (f21 f22 ?v0) ?v1)))))
(assume a28 (forall ((?v0 S4) (?v1 S10) (?v2 S10) (?v3 S10) (?v4 S10)) (not (= (f16 (f17 (f19 f20 ?v0) ?v1) ?v2) (f16 (f17 f18 ?v3) ?v4)))))
(assume a29 (forall ((?v0 S10) (?v1 S10) (?v2 S4) (?v3 S10) (?v4 S10)) (not (= (f16 (f17 f18 ?v0) ?v1) (f16 (f17 (f19 f20 ?v2) ?v3) ?v4)))))
(assume a30 (forall ((?v0 S4) (?v1 S10) (?v2 S10)) (not (= (f16 (f17 (f19 f20 ?v0) ?v1) ?v2) f15))))
(assume a31 (forall ((?v0 S4) (?v1 S10) (?v2 S10)) (not (= f15 (f16 (f17 (f19 f20 ?v0) ?v1) ?v2)))))
(assume a32 (forall ((?v0 S7)) (not (= f23 (f11 f12 ?v0)))))
(assume a33 (forall ((?v0 S7)) (not (= (f11 f12 ?v0) f23))))
(assume a34 (forall ((?v0 S7)) (=> (= f23 (f11 f12 ?v0)) false)))
(assume a35 (forall ((?v0 S7)) (=> (= (f11 f12 ?v0) f23) false)))
(assume a36 (forall ((?v0 S4) (?v1 S10) (?v2 S3) (?v3 S7) (?v4 S3)) (=> (= (f3 (f6 (f7 (f8 f9 (f16 (f21 f22 ?v0) ?v1)) ?v2) ?v3) ?v4) f1) (=> (=> (= ?v4 ?v2) (=> (not (= (f3 ?v0 ?v2) f1)) false)) (=> (forall ((?v5 S3)) (=> (= (f3 ?v0 ?v2) f1) (=> (= (f3 (f6 (f7 (f8 f9 ?v1) ?v2) ?v3) ?v5) f1) (=> (= (f3 (f6 (f7 (f8 f9 (f16 (f21 f22 ?v0) ?v1)) ?v5) ?v3) ?v4) f1) false)))) false)))))
(assume a37 (forall ((?v0 S7)) (=> (not (= ?v0 f23)) (exists ((?v1 S7)) (= ?v0 (f11 f12 ?v1))))))
(assume a38 (forall ((?v0 S16) (?v1 S7)) (=> (= (f24 ?v0 f23) f1) (=> (forall ((?v2 S7)) (=> (= (f24 ?v0 ?v2) f1) (= (f24 ?v0 (f11 f12 ?v2)) f1))) (= (f24 ?v0 ?v1) f1)))))
(assume a39 (forall ((?v0 S16) (?v1 S7)) (=> (= (f24 ?v0 ?v1) f1) (=> (forall ((?v2 S7)) (=> (= (f24 ?v0 (f11 f12 ?v2)) f1) (= (f24 ?v0 ?v2) f1))) (= (f24 ?v0 f23) f1)))))
(assume a40 (forall ((?v0 S7)) (=> (=> (= ?v0 f23) false) (=> (forall ((?v1 S7)) (=> (= ?v0 (f11 f12 ?v1)) false)) false))))
(assume a41 (= (f25 f26 f15) f23))
(assume a42 (= (f25 f27 f15) f23))
(assume a43 (forall ((?v0 S18)) (= (= f28 ?v0) (= ?v0 f28))))
(assume a44 (forall ((?v0 S7)) (= (= f23 ?v0) (= ?v0 f23))))
(assume a45 (forall ((?v0 S10) (?v1 S10)) (= (f25 f27 (f16 (f17 f18 ?v0) ?v1)) (f11 (f29 f30 (f11 (f29 f30 (f25 f27 ?v0)) (f25 f27 ?v1))) (f11 f12 f23)))))
(assume a46 (forall ((?v0 S10) (?v1 S10)) (= (f25 f26 (f16 (f17 f18 ?v0) ?v1)) (f11 (f29 f30 (f11 (f29 f30 (f25 f26 ?v0)) (f25 f26 ?v1))) (f11 f12 f23)))))
(assume a47 (forall ((?v0 S4) (?v1 S10) (?v2 S10)) (= (f25 f27 (f16 (f17 (f19 f20 ?v0) ?v1) ?v2)) (f11 (f29 f30 (f11 (f29 f30 (f25 f27 ?v1)) (f25 f27 ?v2))) (f11 f12 f23)))))
(assume a48 (forall ((?v0 S4) (?v1 S10) (?v2 S10)) (= (f25 f26 (f16 (f17 (f19 f20 ?v0) ?v1) ?v2)) (f11 (f29 f30 (f11 (f29 f30 (f25 f26 ?v1)) (f25 f26 ?v2))) (f11 f12 f23)))))
(assume a49 (forall ((?v0 S4) (?v1 S10)) (= (f25 f27 (f16 (f21 f22 ?v0) ?v1)) (f11 (f29 f30 (f25 f27 ?v1)) (f11 f12 f23)))))
(assume a50 (forall ((?v0 S4) (?v1 S10)) (= (f25 f26 (f16 (f21 f22 ?v0) ?v1)) (f11 (f29 f30 (f25 f26 ?v1)) (f11 f12 f23)))))
(assume a51 (forall ((?v0 S20) (?v1 S21)) (= (f25 f27 (f31 (f32 f33 ?v0) ?v1)) f23)))
(assume a52 (forall ((?v0 S18) (?v1 S18) (?v2 S18)) (= (f34 (f35 f36 (f34 (f35 f36 ?v0) ?v1)) ?v2) (f34 (f35 f36 ?v0) (f34 (f35 f36 ?v1) ?v2)))))
(assume a53 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (= (f11 (f29 f30 (f11 (f29 f30 ?v0) ?v1)) ?v2) (f11 (f29 f30 ?v0) (f11 (f29 f30 ?v1) ?v2)))))
(assume a54 (forall ((?v0 S18) (?v1 S18) (?v2 S18)) (= (= (f34 (f35 f36 ?v0) ?v1) (f34 (f35 f36 ?v0) ?v2)) (= ?v1 ?v2))))
(assume a55 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (= (= (f11 (f29 f30 ?v0) ?v1) (f11 (f29 f30 ?v0) ?v2)) (= ?v1 ?v2))))
(assume a56 (forall ((?v0 S18) (?v1 S18) (?v2 S18)) (= (= (f34 (f35 f36 ?v0) ?v1) (f34 (f35 f36 ?v2) ?v1)) (= ?v0 ?v2))))
(assume a57 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (= (= (f11 (f29 f30 ?v0) ?v1) (f11 (f29 f30 ?v2) ?v1)) (= ?v0 ?v2))))
(assume a58 (forall ((?v0 S18) (?v1 S18) (?v2 S18)) (=> (= (f34 (f35 f36 ?v0) ?v1) (f34 (f35 f36 ?v0) ?v2)) (= ?v1 ?v2))))
(assume a59 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (=> (= (f11 (f29 f30 ?v0) ?v1) (f11 (f29 f30 ?v0) ?v2)) (= ?v1 ?v2))))
(assume a60 (forall ((?v0 S18) (?v1 S18) (?v2 S18)) (=> (= (f34 (f35 f36 ?v0) ?v1) (f34 (f35 f36 ?v2) ?v1)) (= ?v0 ?v2))))
(assume a61 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (=> (= (f11 (f29 f30 ?v0) ?v1) (f11 (f29 f30 ?v2) ?v1)) (= ?v0 ?v2))))
(assume a62 (forall ((?v0 S7) (?v1 S7)) (= (f11 (f29 f30 ?v0) ?v1) (f11 (f29 f30 ?v1) ?v0))))
(assume a63 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (= (f11 (f29 f30 ?v0) (f11 (f29 f30 ?v1) ?v2)) (f11 (f29 f30 ?v1) (f11 (f29 f30 ?v0) ?v2)))))
(assume a64 (forall ((?v0 S20) (?v1 S21) (?v2 S20) (?v3 S21)) (= (= (f31 (f32 f33 ?v0) ?v1) (f31 (f32 f33 ?v2) ?v3)) (and (= ?v0 ?v2) (= ?v1 ?v3)))))
(assume a65 (forall ((?v0 S18)) (= (f34 (f35 f36 ?v0) f28) ?v0)))
(assume a66 (forall ((?v0 S7)) (= (f11 (f29 f30 ?v0) f23) ?v0)))
(assume a67 (forall ((?v0 S18) (?v1 S26)) (= (= (f37 (f38 f39 ?v0) ?v1) f1) (= (f40 ?v1 ?v0) f1))))
(assume a68 (forall ((?v0 S18)) (= (= f28 (f34 (f35 f36 ?v0) ?v0)) (= ?v0 f28))))
(assume a69 (forall ((?v0 S18)) (= (f34 (f35 f36 f28) ?v0) ?v0)))
(assume a70 (forall ((?v0 S7)) (= (f11 (f29 f30 f23) ?v0) ?v0)))
(assume a71 (forall ((?v0 S7) (?v1 S7)) (= (= (f11 (f29 f30 ?v0) ?v1) f23) (and (= ?v0 f23) (= ?v1 f23)))))
(assume a72 (forall ((?v0 S7) (?v1 S7)) (=> (= (f11 (f29 f30 ?v0) ?v1) ?v0) (= ?v1 f23))))
(assume a73 (forall ((?v0 S7) (?v1 S7)) (= (f11 (f29 f30 ?v0) (f11 f12 ?v1)) (f11 f12 (f11 (f29 f30 ?v0) ?v1)))))
(assume a74 (forall ((?v0 S7) (?v1 S7)) (= (f11 (f29 f30 (f11 f12 ?v0)) ?v1) (f11 f12 (f11 (f29 f30 ?v0) ?v1)))))
(assume a75 (forall ((?v0 S7) (?v1 S7)) (= (f11 (f29 f30 (f11 f12 ?v0)) ?v1) (f11 (f29 f30 ?v0) (f11 f12 ?v1)))))
(assume a76 (forall ((?v0 S4) (?v1 S10) (?v2 S20) (?v3 S21)) (not (= (f16 (f21 f22 ?v0) ?v1) (f31 (f32 f33 ?v2) ?v3)))))
(assume a77 (forall ((?v0 S20) (?v1 S21) (?v2 S4) (?v3 S10)) (not (= (f31 (f32 f33 ?v0) ?v1) (f16 (f21 f22 ?v2) ?v3)))))
(assume a78 (forall ((?v0 S4) (?v1 S10) (?v2 S10) (?v3 S20) (?v4 S21)) (not (= (f16 (f17 (f19 f20 ?v0) ?v1) ?v2) (f31 (f32 f33 ?v3) ?v4)))))
(assume a79 (forall ((?v0 S20) (?v1 S21) (?v2 S4) (?v3 S10) (?v4 S10)) (not (= (f31 (f32 f33 ?v0) ?v1) (f16 (f17 (f19 f20 ?v2) ?v3) ?v4)))))
(assume a80 (forall ((?v0 S10) (?v1 S10) (?v2 S20) (?v3 S21)) (not (= (f16 (f17 f18 ?v0) ?v1) (f31 (f32 f33 ?v2) ?v3)))))
(assume a81 (forall ((?v0 S20) (?v1 S21) (?v2 S10) (?v3 S10)) (not (= (f31 (f32 f33 ?v0) ?v1) (f16 (f17 f18 ?v2) ?v3)))))
(assume a82 (forall ((?v0 S20) (?v1 S21)) (not (= (f31 (f32 f33 ?v0) ?v1) f15))))
(assume a83 (forall ((?v0 S20) (?v1 S21)) (not (= f15 (f31 (f32 f33 ?v0) ?v1)))))
(assume a84 (forall ((?v0 S7) (?v1 S7)) (= (= (f11 (f29 f30 ?v0) ?v1) (f11 f12 f23)) (or (and (= ?v0 (f11 f12 f23)) (= ?v1 f23)) (and (= ?v0 f23) (= ?v1 (f11 f12 f23)))))))
(assume a85 (forall ((?v0 S7) (?v1 S7)) (= (= (f11 f12 f23) (f11 (f29 f30 ?v0) ?v1)) (or (and (= ?v0 (f11 f12 f23)) (= ?v1 f23)) (and (= ?v0 f23) (= ?v1 (f11 f12 f23)))))))
(assume a86 (forall ((?v0 S20) (?v1 S21)) (= (f25 f26 (f31 (f32 f33 ?v0) ?v1)) f23)))
(assume a87 (forall ((?v0 S18) (?v1 S18)) (= (= ?v0 (f34 (f35 f36 ?v0) ?v1)) (= ?v1 f28))))
(assume a88 (forall ((?v0 S7) (?v1 S7)) (= (= ?v0 (f11 (f29 f30 ?v0) ?v1)) (= ?v1 f23))))
(assume a89 (forall ((?v0 S18)) (= (= (f34 (f35 f36 ?v0) ?v0) f28) (= ?v0 f28))))
(assume a90 (forall ((?v0 S29) (?v1 S21) (?v2 S10)) (= (f25 f26 (f16 (f41 (f42 f43 ?v0) ?v1) ?v2)) (f11 (f29 f30 (f25 f26 ?v2)) (f11 f12 f23)))))
(assume a91 (forall ((?v0 S29) (?v1 S21) (?v2 S10)) (= (f25 f27 (f16 (f41 (f42 f43 ?v0) ?v1) ?v2)) (f11 (f29 f30 (f25 f27 ?v2)) (f11 f12 f23)))))
(assume a92 (forall ((?v0 S7)) (= (f11 f44 (f11 f12 ?v0)) (f11 (f29 f30 (f11 f44 ?v0)) (f11 f12 f23)))))
(assume a93 (forall ((?v0 S29) (?v1 S21) (?v2 S10) (?v3 S29) (?v4 S21) (?v5 S10)) (= (= (f16 (f41 (f42 f43 ?v0) ?v1) ?v2) (f16 (f41 (f42 f43 ?v3) ?v4) ?v5)) (and (= ?v0 ?v3) (and (= ?v1 ?v4) (= ?v2 ?v5))))))
(assume a94 (= (f11 f44 f23) f23))
(assume a95 (forall ((?v0 S4) (?v1 S10) (?v2 S29) (?v3 S21) (?v4 S10)) (not (= (f16 (f21 f22 ?v0) ?v1) (f16 (f41 (f42 f43 ?v2) ?v3) ?v4)))))
(assume a96 (forall ((?v0 S29) (?v1 S21) (?v2 S10) (?v3 S4) (?v4 S10)) (not (= (f16 (f41 (f42 f43 ?v0) ?v1) ?v2) (f16 (f21 f22 ?v3) ?v4)))))
(assume a97 (forall ((?v0 S4) (?v1 S10) (?v2 S10) (?v3 S29) (?v4 S21) (?v5 S10)) (not (= (f16 (f17 (f19 f20 ?v0) ?v1) ?v2) (f16 (f41 (f42 f43 ?v3) ?v4) ?v5)))))
(assume a98 (forall ((?v0 S29) (?v1 S21) (?v2 S10) (?v3 S4) (?v4 S10) (?v5 S10)) (not (= (f16 (f41 (f42 f43 ?v0) ?v1) ?v2) (f16 (f17 (f19 f20 ?v3) ?v4) ?v5)))))
(assume a99 (forall ((?v0 S10) (?v1 S10) (?v2 S29) (?v3 S21) (?v4 S10)) (not (= (f16 (f17 f18 ?v0) ?v1) (f16 (f41 (f42 f43 ?v2) ?v3) ?v4)))))
(assume a100 (forall ((?v0 S29) (?v1 S21) (?v2 S10) (?v3 S10) (?v4 S10)) (not (= (f16 (f41 (f42 f43 ?v0) ?v1) ?v2) (f16 (f17 f18 ?v3) ?v4)))))
(assume a101 (forall ((?v0 S29) (?v1 S21) (?v2 S10) (?v3 S20) (?v4 S21)) (not (= (f16 (f41 (f42 f43 ?v0) ?v1) ?v2) (f31 (f32 f33 ?v3) ?v4)))))
(assume a102 (forall ((?v0 S20) (?v1 S21) (?v2 S29) (?v3 S21) (?v4 S10)) (not (= (f31 (f32 f33 ?v0) ?v1) (f16 (f41 (f42 f43 ?v2) ?v3) ?v4)))))
(assume a103 (forall ((?v0 S29) (?v1 S21) (?v2 S10)) (not (= f15 (f16 (f41 (f42 f43 ?v0) ?v1) ?v2)))))
(assume a104 (forall ((?v0 S29) (?v1 S21) (?v2 S10)) (not (= (f16 (f41 (f42 f43 ?v0) ?v1) ?v2) f15))))
(assume a105 (forall ((?v0 S18) (?v1 S18)) (= (f34 (f35 f36 ?v0) ?v1) (f34 (f35 f36 ?v1) ?v0))))
(assume a106 (forall ((?v0 S18) (?v1 S18) (?v2 S18)) (= (f34 (f35 f36 ?v0) (f34 (f35 f36 ?v1) ?v2)) (f34 (f35 f36 ?v1) (f34 (f35 f36 ?v0) ?v2)))))
(assume a107 (forall ((?v0 S18) (?v1 S18) (?v2 S18)) (= (f34 (f35 f36 ?v0) (f34 (f35 f36 ?v1) ?v2)) (f34 (f35 f36 (f34 (f35 f36 ?v0) ?v1)) ?v2))))
(assume a108 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (= (f11 (f29 f30 ?v0) (f11 (f29 f30 ?v1) ?v2)) (f11 (f29 f30 (f11 (f29 f30 ?v0) ?v1)) ?v2))))
(assume a109 (forall ((?v0 S18) (?v1 S18) (?v2 S18)) (= (f34 (f35 f36 (f34 (f35 f36 ?v0) ?v1)) ?v2) (f34 (f35 f36 (f34 (f35 f36 ?v0) ?v2)) ?v1))))
(assume a110 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (= (f11 (f29 f30 (f11 (f29 f30 ?v0) ?v1)) ?v2) (f11 (f29 f30 (f11 (f29 f30 ?v0) ?v2)) ?v1))))
(assume a111 (forall ((?v0 S18) (?v1 S18) (?v2 S18) (?v3 S18)) (= (f34 (f35 f36 (f34 (f35 f36 ?v0) ?v1)) (f34 (f35 f36 ?v2) ?v3)) (f34 (f35 f36 (f34 (f35 f36 ?v0) ?v2)) (f34 (f35 f36 ?v1) ?v3)))))
(assume a112 (forall ((?v0 S7) (?v1 S7) (?v2 S7) (?v3 S7)) (= (f11 (f29 f30 (f11 (f29 f30 ?v0) ?v1)) (f11 (f29 f30 ?v2) ?v3)) (f11 (f29 f30 (f11 (f29 f30 ?v0) ?v2)) (f11 (f29 f30 ?v1) ?v3)))))
(assume a113 (forall ((?v0 S7)) (= (f11 f45 (f11 f12 ?v0)) (f11 (f29 f30 (f11 f45 ?v0)) (f11 f12 f23)))))
(assume a114 (= (f46 f47 f1) f23))
(assume a115 (= (f46 f47 f2) f23))
(assume a116 (forall ((?v0 S20) (?v1 S21) (?v2 S3) (?v3 S7) (?v4 S3)) (=> (= (f3 (f6 (f7 (f8 f9 (f31 (f32 f33 ?v0) ?v1)) ?v2) ?v3) ?v4) f1) (=> (=> (= ?v4 (f48 (f49 (f50 f51 ?v2) ?v0) (f52 ?v1 ?v2))) false) false))))
(assume a117 (forall ((?v0 S20) (?v1 S21) (?v2 S3) (?v3 S7)) (= (f3 (f6 (f7 (f8 f9 (f31 (f32 f33 ?v0) ?v1)) ?v2) ?v3) (f48 (f49 (f50 f51 ?v2) ?v0) (f52 ?v1 ?v2))) f1)))
(assume a118 (forall ((?v0 S7)) (= (f11 f45 ?v0) ?v0)))
(assume a119 (= (f11 f45 f23) f23))
(assume a120 (forall ((?v0 S20) (?v1 S21) (?v2 S3) (?v3 S3)) (=> (= (f3 (f53 (f54 f55 (f31 (f32 f33 ?v0) ?v1)) ?v2) ?v3) f1) (=> (=> (= ?v3 (f48 (f49 (f50 f51 ?v2) ?v0) (f52 ?v1 ?v2))) false) false))))
(assume a121 (forall ((?v0 S20) (?v1 S21) (?v2 S3)) (= (f3 (f53 (f54 f55 (f31 (f32 f33 ?v0) ?v1)) ?v2) (f48 (f49 (f50 f51 ?v2) ?v0) (f52 ?v1 ?v2))) f1)))
(assume a122 (forall ((?v0 S1)) (= (f46 f56 ?v0) f23)))
(assume a123 (= (f46 f56 f2) f23))
(assume a124 (forall ((?v0 S4) (?v1 S3) (?v2 S10) (?v3 S3) (?v4 S3)) (=> (= (f3 ?v0 ?v1) f1) (=> (= (f3 (f53 (f54 f55 ?v2) ?v1) ?v3) f1) (=> (= (f3 (f53 (f54 f55 (f16 (f21 f22 ?v0) ?v2)) ?v3) ?v4) f1) (= (f3 (f53 (f54 f55 (f16 (f21 f22 ?v0) ?v2)) ?v1) ?v4) f1))))))
(assume a125 (forall ((?v0 S4) (?v1 S3) (?v2 S10)) (=> (not (= (f3 ?v0 ?v1) f1)) (= (f3 (f53 (f54 f55 (f16 (f21 f22 ?v0) ?v2)) ?v1) ?v1) f1))))
(assume a126 (forall ((?v0 S4) (?v1 S10) (?v2 S10) (?v3 S3) (?v4 S3)) (=> (= (f3 (f53 (f54 f55 (f16 (f17 (f19 f20 ?v0) ?v1) ?v2)) ?v3) ?v4) f1) (=> (=> (= (f3 ?v0 ?v3) f1) (=> (= (f3 (f53 (f54 f55 ?v1) ?v3) ?v4) f1) false)) (=> (=> (not (= (f3 ?v0 ?v3) f1)) (=> (= (f3 (f53 (f54 f55 ?v2) ?v3) ?v4) f1) false)) false)))))
(assume a127 (forall ((?v0 S4) (?v1 S3) (?v2 S10) (?v3 S3) (?v4 S10)) (=> (= (f3 ?v0 ?v1) f1) (=> (= (f3 (f53 (f54 f55 ?v2) ?v1) ?v3) f1) (= (f3 (f53 (f54 f55 (f16 (f17 (f19 f20 ?v0) ?v2) ?v4)) ?v1) ?v3) f1)))))
(assume a128 (forall ((?v0 S4) (?v1 S3) (?v2 S10) (?v3 S3) (?v4 S10)) (=> (not (= (f3 ?v0 ?v1) f1)) (=> (= (f3 (f53 (f54 f55 ?v2) ?v1) ?v3) f1) (= (f3 (f53 (f54 f55 (f16 (f17 (f19 f20 ?v0) ?v4) ?v2)) ?v1) ?v3) f1)))))
(assume a129 (forall ((?v0 S10) (?v1 S3) (?v2 S3) (?v3 S10) (?v4 S3)) (=> (= (f3 (f53 (f54 f55 ?v0) ?v1) ?v2) f1) (=> (= (f3 (f53 (f54 f55 ?v3) ?v2) ?v4) f1) (= (f3 (f53 (f54 f55 (f16 (f17 f18 ?v0) ?v3)) ?v1) ?v4) f1)))))
(assume a130 (forall ((?v0 S3)) (= (f3 (f53 (f54 f55 f15) ?v0) ?v0) f1)))
(assume a131 (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f53 (f54 f55 f15) ?v0) ?v1) f1) (=> (=> (= ?v1 ?v0) false) false))))
(assume a132 (forall ((?v0 S10) (?v1 S3) (?v2 S3)) (= (= (f3 (f53 (f54 f55 ?v0) ?v1) ?v2) f1) (exists ((?v3 S7)) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v3) ?v2) f1)))))
(assume a133 (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (=> (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) f1) (= (f3 (f53 (f54 f55 ?v0) ?v1) ?v3) f1))))
(assume a134 (forall ((?v0 S10) (?v1 S3) (?v2 S3) (?v3 S3)) (=> (= (f3 (f53 (f54 f55 ?v0) ?v1) ?v2) f1) (=> (= (f3 (f53 (f54 f55 ?v0) ?v1) ?v3) f1) (= ?v3 ?v2)))))
(assume a135 (= (f46 f56 f1) f23))
(assume a136 (forall ((?v0 S10) (?v1 S10) (?v2 S3) (?v3 S3)) (=> (= (f3 (f53 (f54 f55 (f16 (f17 f18 ?v0) ?v1)) ?v2) ?v3) f1) (=> (forall ((?v4 S3)) (=> (= (f3 (f53 (f54 f55 ?v0) ?v2) ?v4) f1) (=> (= (f3 (f53 (f54 f55 ?v1) ?v4) ?v3) f1) false))) false))))
(assume a137 (forall ((?v0 S4) (?v1 S10) (?v2 S3) (?v3 S3)) (=> (= (f3 (f53 (f54 f55 (f16 (f21 f22 ?v0) ?v1)) ?v2) ?v3) f1) (=> (=> (= ?v3 ?v2) (=> (not (= (f3 ?v0 ?v2) f1)) false)) (=> (forall ((?v4 S3)) (=> (= (f3 ?v0 ?v2) f1) (=> (= (f3 (f53 (f54 f55 ?v1) ?v2) ?v4) f1) (=> (= (f3 (f53 (f54 f55 (f16 (f21 f22 ?v0) ?v1)) ?v4) ?v3) f1) false)))) false)))))
(assume a138 (forall ((?v0 S10) (?v1 S3) (?v2 S3)) (=> (= (f3 (f53 (f54 f55 ?v0) ?v1) ?v2) f1) (exists ((?v3 S7)) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v3) ?v2) f1)))))
(assume a139 (forall ((?v0 S18)) (=> (= (f37 (f38 f39 ?v0) f57) f1) (= (= (f34 (f35 f36 ?v0) ?v0) f28) (= ?v0 f28)))))
(assume a140 (! (= (! (f37 (f38 f39 f28) f57) :named @p_3) f1) :named @p_4))
(assume a141 (forall ((?v0 S18) (?v1 S18)) (=> (= (f37 (f38 f39 ?v0) f57) f1) (=> (= (f37 (f38 f39 ?v1) f57) f1) (= (f37 (f38 f39 (f34 (f35 f36 ?v0) ?v1)) f57) f1)))))
(assume a142 (forall ((?v0 S18)) (=> (= (f37 (f38 f39 ?v0) f57) f1) (not (= (f34 (f35 f36 (f34 (f35 f36 f58) ?v0)) ?v0) f28)))))
(assume a143 (forall ((?v0 S20) (?v1 S38) (?v2 S21)) (= (f25 f26 (f31 (f59 (f60 f61 ?v0) ?v1) ?v2)) f23)))
(assume a144 (= (f37 (f38 f39 f58) f57) f1))
(assume a145 (forall ((?v0 S20) (?v1 S38) (?v2 S21) (?v3 S4) (?v4 S10)) (not (= (f31 (f59 (f60 f61 ?v0) ?v1) ?v2) (f16 (f21 f22 ?v3) ?v4)))))
(assume a146 (forall ((?v0 S4) (?v1 S10) (?v2 S20) (?v3 S38) (?v4 S21)) (not (= (f16 (f21 f22 ?v0) ?v1) (f31 (f59 (f60 f61 ?v2) ?v3) ?v4)))))
(assume a147 (forall ((?v0 S4) (?v1 S10) (?v2 S10) (?v3 S20) (?v4 S38) (?v5 S21)) (not (= (f16 (f17 (f19 f20 ?v0) ?v1) ?v2) (f31 (f59 (f60 f61 ?v3) ?v4) ?v5)))))
(assume a148 (forall ((?v0 S20) (?v1 S38) (?v2 S21) (?v3 S4) (?v4 S10) (?v5 S10)) (not (= (f31 (f59 (f60 f61 ?v0) ?v1) ?v2) (f16 (f17 (f19 f20 ?v3) ?v4) ?v5)))))
(assume a149 (forall ((?v0 S20) (?v1 S21) (?v2 S20) (?v3 S38) (?v4 S21)) (not (= (f31 (f32 f33 ?v0) ?v1) (f31 (f59 (f60 f61 ?v2) ?v3) ?v4)))))
(assume a150 (forall ((?v0 S20) (?v1 S38) (?v2 S21) (?v3 S20) (?v4 S21)) (not (= (f31 (f59 (f60 f61 ?v0) ?v1) ?v2) (f31 (f32 f33 ?v3) ?v4)))))
(assume a151 (forall ((?v0 S20) (?v1 S38) (?v2 S21) (?v3 S10) (?v4 S10)) (not (= (f31 (f59 (f60 f61 ?v0) ?v1) ?v2) (f16 (f17 f18 ?v3) ?v4)))))
(assume a152 (forall ((?v0 S10) (?v1 S10) (?v2 S20) (?v3 S38) (?v4 S21)) (not (= (f16 (f17 f18 ?v0) ?v1) (f31 (f59 (f60 f61 ?v2) ?v3) ?v4)))))
(assume a153 (forall ((?v0 S20) (?v1 S38) (?v2 S21) (?v3 S29) (?v4 S21) (?v5 S10)) (not (= (f31 (f59 (f60 f61 ?v0) ?v1) ?v2) (f16 (f41 (f42 f43 ?v3) ?v4) ?v5)))))
(assume a154 (forall ((?v0 S29) (?v1 S21) (?v2 S10) (?v3 S20) (?v4 S38) (?v5 S21)) (not (= (f16 (f41 (f42 f43 ?v0) ?v1) ?v2) (f31 (f59 (f60 f61 ?v3) ?v4) ?v5)))))
(assume a155 (forall ((?v0 S20) (?v1 S38) (?v2 S21)) (not (= (f31 (f59 (f60 f61 ?v0) ?v1) ?v2) f15))))
(assume a156 (forall ((?v0 S20) (?v1 S38) (?v2 S21)) (not (= f15 (f31 (f59 (f60 f61 ?v0) ?v1) ?v2)))))
(assume a157 (forall ((?v0 S20) (?v1 S38) (?v2 S21) (?v3 S20) (?v4 S38) (?v5 S21)) (= (= (f31 (f59 (f60 f61 ?v0) ?v1) ?v2) (f31 (f59 (f60 f61 ?v3) ?v4) ?v5)) (and (= ?v0 ?v3) (and (= ?v1 ?v4) (= ?v2 ?v5))))))
(assume a158 (forall ((?v0 S18)) (= (= f58 ?v0) (= ?v0 f58))))
(assume a159 (forall ((?v0 S7)) (= (= f62 ?v0) (= ?v0 f62))))
(assume a160 (= f62 (f11 f12 f23)))
(assume a161 (forall ((?v0 S20) (?v1 S38) (?v2 S21)) (= (f25 f27 (f31 (f59 (f60 f61 ?v0) ?v1) ?v2)) f23)))
(assume a162 (forall ((?v0 S7)) (= (f11 f12 ?v0) (f11 (f29 f30 f62) ?v0))))
(assume a163 (forall ((?v0 S7)) (= (f11 f12 ?v0) (f11 (f29 f30 ?v0) f62))))
(assume a164 (not (= f58 f28)))
(assume a165 (not (= f62 f23)))
(assume a166 (not (= f28 f58)))
(assume a167 (not (= f23 f62)))
(assume a168 (forall ((?v0 S18)) (=> (= (f37 (f38 f39 ?v0) f57) f1) (= (= (f40 (f63 f64 (f34 (f35 f36 (f34 (f35 f36 f58) ?v0)) ?v0)) f28) f1) (= (f40 (f63 f64 ?v0) f28) f1)))))
(assume a169 (forall ((?v0 S7) (?v1 S7)) (= (f11 (f29 f30 ?v0) ?v1) (ite (= ?v0 f23) ?v1 (f11 f12 (f11 (f29 f30 (f11 (f29 f65 ?v0) f62)) ?v1))))))
(assume a170 (forall ((?v0 S7)) (=> (= (f24 (f66 f67 ?v0) f23) f1) false)))
(assume a171 (forall ((?v0 S7) (?v1 S7)) (=> (= (f24 (f66 f67 ?v0) ?v1) f1) (= (f24 (f66 f67 (f11 f12 ?v0)) (f11 f12 ?v1)) f1))))
(assume a172 (forall ((?v0 S7)) (= (f24 (f66 f67 ?v0) (f11 f12 ?v0)) f1)))
(assume a173 (forall ((?v0 S7)) (= (f24 (f66 f67 f23) (f11 f12 ?v0)) f1)))
(assume a174 (forall ((?v0 S7) (?v1 S7)) (= (f11 (f29 f65 ?v0) (f11 f12 ?v1)) (f11 (f29 f65 (f11 (f29 f65 ?v0) f62)) ?v1))))
(assume a175 (not (= (f40 (f63 f64 f58) f28) f1)))
(assume a176 (not (= (f24 (f66 f67 f62) f23) f1)))
(assume a177 (= (f40 (f63 f64 f28) f58) f1))
(assume a178 (= (f24 (f66 f67 f23) f62) f1))
(assume a179 (forall ((?v0 S7)) (=> (= (f24 (f66 f67 f23) ?v0) f1) (= ?v0 (f11 f12 (f11 (f29 f65 ?v0) f62))))))
(assume a180 (forall ((?v0 S18)) (= (f40 (f63 f64 ?v0) (f34 (f35 f36 ?v0) f58)) f1)))
(assume a181 (forall ((?v0 S7)) (= (f24 (f66 f67 ?v0) (f11 (f29 f30 ?v0) f62)) f1)))
(assume a182 (forall ((?v0 S18) (?v1 S18) (?v2 S18)) (=> (= (f40 (f63 f64 f28) ?v0) f1) (=> (= (f40 (f63 f64 ?v1) ?v2) f1) (= (f40 (f63 f64 ?v1) (f34 (f35 f36 ?v0) ?v2)) f1)))))
(assume a183 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (=> (= (f24 (f66 f67 f23) ?v0) f1) (=> (= (f24 (f66 f67 ?v1) ?v2) f1) (= (f24 (f66 f67 ?v1) (f11 (f29 f30 ?v0) ?v2)) f1)))))
(assume a184 (forall ((?v0 S18) (?v1 S18)) (=> (not (= ?v0 ?v1)) (=> (=> (= (f40 (f63 f64 ?v0) ?v1) f1) false) (=> (=> (= (f40 (f63 f64 ?v1) ?v0) f1) false) false)))))
(assume a185 (forall ((?v0 S7)) (= (= (f24 (f66 f67 f23) ?v0) f1) (exists ((?v1 S7)) (= ?v0 (f11 f12 ?v1))))))
(assume a186 (forall ((?v0 S7)) (= (= (f24 (f66 f67 ?v0) (f11 f12 f23)) f1) (= ?v0 f23))))
(assume a187 (forall ((?v0 S7) (?v1 S7)) (= (= (f24 (f66 f67 ?v0) (f11 f12 ?v1)) f1) (or (= ?v0 f23) (exists ((?v2 S7)) (and (= ?v0 (f11 f12 ?v2)) (= (f24 (f66 f67 ?v2) ?v1) f1)))))))
(assume a188 (forall ((?v0 S7) (?v1 S7)) (= (= (f24 (f66 f67 f23) (f11 (f29 f30 ?v0) ?v1)) f1) (or (= (f24 (f66 f67 f23) ?v0) f1) (= (f24 (f66 f67 f23) ?v1) f1)))))
(assume a189 (forall ((?v0 S7) (?v1 S7)) (= (= (f24 (f66 f67 ?v0) ?v1) f1) (exists ((?v2 S7)) (= ?v1 (f11 f12 (f11 (f29 f30 ?v0) ?v2)))))))
(assume a190 (forall ((?v0 S7) (?v1 S7)) (= (f24 (f66 f67 ?v0) (f11 f12 (f11 (f29 f30 ?v1) ?v0))) f1)))
(assume a191 (forall ((?v0 S7) (?v1 S7)) (= (f24 (f66 f67 ?v0) (f11 f12 (f11 (f29 f30 ?v0) ?v1))) f1)))
(assume a192 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (= (= (f24 (f66 f67 ?v0) (f11 (f29 f65 ?v1) ?v2)) f1) (= (f24 (f66 f67 (f11 (f29 f30 ?v0) ?v2)) ?v1) f1))))
(assume a193 (forall ((?v0 S7) (?v1 S7)) (=> (not (= (f24 (f66 f67 ?v0) ?v1) f1)) (= (f11 (f29 f30 ?v1) (f11 (f29 f65 ?v0) ?v1)) ?v0))))
(assume a194 (forall ((?v0 S7) (?v1 S7)) (=> (= (f24 (f66 f67 f23) ?v0) f1) (= (f24 (f66 f67 (f11 (f29 f65 ?v0) (f11 f12 ?v1))) ?v0) f1))))
(assume a195 (forall ((?v0 S7)) (=> (= (f24 (f66 f67 f23) ?v0) f1) (= (f11 f12 (f11 (f29 f65 ?v0) (f11 f12 f23))) ?v0))))
(assume a196 (forall ((?v0 S16) (?v1 S7) (?v2 S7)) (= (= (f24 ?v0 (f11 (f29 f65 ?v1) ?v2)) f1) (and (=> (= (f24 (f66 f67 ?v1) ?v2) f1) (= (f24 ?v0 f23) f1)) (forall ((?v3 S7)) (=> (= ?v1 (f11 (f29 f30 ?v2) ?v3)) (= (f24 ?v0 ?v3) f1)))))))
(assume a197 (forall ((?v0 S16) (?v1 S7) (?v2 S7)) (= (= (f24 ?v0 (f11 (f29 f65 ?v1) ?v2)) f1) (not (or (and (= (f24 (f66 f67 ?v1) ?v2) f1) (not (= (f24 ?v0 f23) f1))) (exists ((?v3 S7)) (and (= ?v1 (f11 (f29 f30 ?v2) ?v3)) (not (= (f24 ?v0 ?v3) f1)))))))))
(assume a198 (forall ((?v0 S18) (?v1 S18)) (= (= (f34 (f35 f68 ?v0) ?v1) f28) (= ?v0 ?v1))))
(assume a199 (forall ((?v0 S18) (?v1 S18)) (= (= ?v0 ?v1) (= (f34 (f35 f68 ?v0) ?v1) f28))))
(assume a200 (forall ((?v0 S18)) (= (f34 (f35 f68 ?v0) ?v0) f28)))
(assume a201 (forall ((?v0 S18)) (= (f34 (f35 f68 ?v0) f28) ?v0)))
(assume a202 (forall ((?v0 S18) (?v1 S18)) (= (f34 (f35 f68 (f34 (f35 f36 ?v0) ?v1)) ?v1) ?v0)))
(assume a203 (forall ((?v0 S18) (?v1 S18)) (= (f34 (f35 f36 (f34 (f35 f68 ?v0) ?v1)) ?v1) ?v0)))
(assume a204 (forall ((?v0 S7) (?v1 S7)) (=> (= (f11 (f29 f65 ?v0) ?v1) f23) (=> (= (f11 (f29 f65 ?v1) ?v0) f23) (= ?v0 ?v1)))))
(assume a205 (forall ((?v0 S7)) (= (f11 (f29 f65 ?v0) ?v0) f23)))
(assume a206 (forall ((?v0 S7)) (= (f11 (f29 f65 ?v0) f23) ?v0)))
(assume a207 (forall ((?v0 S7)) (= (f11 (f29 f65 f23) ?v0) f23)))
(assume a208 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (= (f11 (f29 f65 (f11 (f29 f65 (f11 f12 ?v0)) ?v1)) (f11 f12 ?v2)) (f11 (f29 f65 (f11 (f29 f65 ?v0) ?v1)) ?v2))))
(assume a209 (forall ((?v0 S7) (?v1 S7)) (= (f11 (f29 f65 (f11 f12 ?v0)) (f11 f12 ?v1)) (f11 (f29 f65 ?v0) ?v1))))
(assume a210 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (= (f11 (f29 f65 (f11 (f29 f30 ?v0) ?v1)) (f11 (f29 f30 ?v2) ?v1)) (f11 (f29 f65 ?v0) ?v2))))
(assume a211 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (= (f11 (f29 f65 (f11 (f29 f30 ?v0) ?v1)) (f11 (f29 f30 ?v0) ?v2)) (f11 (f29 f65 ?v1) ?v2))))
(assume a212 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (= (f11 (f29 f65 (f11 (f29 f65 ?v0) ?v1)) ?v2) (f11 (f29 f65 ?v0) (f11 (f29 f30 ?v1) ?v2)))))
(assume a213 (forall ((?v0 S7) (?v1 S7)) (= (f11 (f29 f65 (f11 (f29 f30 ?v0) ?v1)) ?v0) ?v1)))
(assume a214 (forall ((?v0 S7) (?v1 S7)) (= (f11 (f29 f65 (f11 (f29 f30 ?v0) ?v1)) ?v1) ?v0)))
(assume a215 (forall ((?v0 S18) (?v1 S18) (?v2 S18)) (=> (= (f40 (f63 f64 (f34 (f35 f36 ?v0) ?v1)) (f34 (f35 f36 ?v0) ?v2)) f1) (= (f40 (f63 f64 ?v1) ?v2) f1))))
(assume a216 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (=> (= (f24 (f66 f67 (f11 (f29 f30 ?v0) ?v1)) (f11 (f29 f30 ?v0) ?v2)) f1) (= (f24 (f66 f67 ?v1) ?v2) f1))))
(assume a217 (forall ((?v0 S18) (?v1 S18) (?v2 S18)) (=> (= (f40 (f63 f64 (f34 (f35 f36 ?v0) ?v1)) (f34 (f35 f36 ?v2) ?v1)) f1) (= (f40 (f63 f64 ?v0) ?v2) f1))))
(assume a218 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (=> (= (f24 (f66 f67 (f11 (f29 f30 ?v0) ?v1)) (f11 (f29 f30 ?v2) ?v1)) f1) (= (f24 (f66 f67 ?v0) ?v2) f1))))
(assume a219 (forall ((?v0 S18) (?v1 S18) (?v2 S18) (?v3 S18)) (=> (= (f40 (f63 f64 ?v0) ?v1) f1) (=> (= (f40 (f63 f64 ?v2) ?v3) f1) (= (f40 (f63 f64 (f34 (f35 f36 ?v0) ?v2)) (f34 (f35 f36 ?v1) ?v3)) f1)))))
(assume a220 (forall ((?v0 S7) (?v1 S7) (?v2 S7) (?v3 S7)) (=> (= (f24 (f66 f67 ?v0) ?v1) f1) (=> (= (f24 (f66 f67 ?v2) ?v3) f1) (= (f24 (f66 f67 (f11 (f29 f30 ?v0) ?v2)) (f11 (f29 f30 ?v1) ?v3)) f1)))))
(assume a221 (forall ((?v0 S18) (?v1 S18) (?v2 S18)) (=> (= (f40 (f63 f64 ?v0) ?v1) f1) (= (f40 (f63 f64 (f34 (f35 f36 ?v2) ?v0)) (f34 (f35 f36 ?v2) ?v1)) f1))))
(assume a222 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (=> (= (f24 (f66 f67 ?v0) ?v1) f1) (= (f24 (f66 f67 (f11 (f29 f30 ?v2) ?v0)) (f11 (f29 f30 ?v2) ?v1)) f1))))
(assume a223 (forall ((?v0 S18) (?v1 S18) (?v2 S18)) (=> (= (f40 (f63 f64 ?v0) ?v1) f1) (= (f40 (f63 f64 (f34 (f35 f36 ?v0) ?v2)) (f34 (f35 f36 ?v1) ?v2)) f1))))
(assume a224 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (=> (= (f24 (f66 f67 ?v0) ?v1) f1) (= (f24 (f66 f67 (f11 (f29 f30 ?v0) ?v2)) (f11 (f29 f30 ?v1) ?v2)) f1))))
(assume a225 (forall ((?v0 S18) (?v1 S18) (?v2 S18)) (= (= (f40 (f63 f64 (f34 (f35 f36 ?v0) ?v1)) (f34 (f35 f36 ?v0) ?v2)) f1) (= (f40 (f63 f64 ?v1) ?v2) f1))))
(assume a226 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (= (= (f24 (f66 f67 (f11 (f29 f30 ?v0) ?v1)) (f11 (f29 f30 ?v0) ?v2)) f1) (= (f24 (f66 f67 ?v1) ?v2) f1))))
(assume a227 (forall ((?v0 S18) (?v1 S18) (?v2 S18)) (= (= (f40 (f63 f64 (f34 (f35 f36 ?v0) ?v1)) (f34 (f35 f36 ?v2) ?v1)) f1) (= (f40 (f63 f64 ?v0) ?v2) f1))))
(assume a228 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (= (= (f24 (f66 f67 (f11 (f29 f30 ?v0) ?v1)) (f11 (f29 f30 ?v2) ?v1)) f1) (= (f24 (f66 f67 ?v0) ?v2) f1))))
(assume a229 (forall ((?v0 S7)) (=> (=> (= ?v0 f23) false) (= (f24 (f66 f67 f23) ?v0) f1))))
(assume a230 (forall ((?v0 S7) (?v1 S7)) (=> (= (f24 (f66 f67 ?v0) ?v1) f1) (not (= ?v1 f23)))))
(assume a231 (forall ((?v0 S7)) (= (= (f24 (f66 f67 ?v0) f23) f1) false)))
(assume a232 (forall ((?v0 S7)) (= (not (= ?v0 f23)) (= (f24 (f66 f67 f23) ?v0) f1))))
(assume a233 (forall ((?v0 S7)) (not (= (f24 (f66 f67 ?v0) f23) f1))))
(assume a234 (forall ((?v0 S7) (?v1 S7)) (=> (= (f24 (f66 f67 (f11 f12 ?v0)) (f11 f12 ?v1)) f1) (= (f24 (f66 f67 ?v0) ?v1) f1))))
(assume a235 (forall ((?v0 S7) (?v1 S7)) (=> (= (f24 (f66 f67 (f11 f12 ?v0)) ?v1) f1) (= (f24 (f66 f67 ?v0) ?v1) f1))))
(assume a236 (forall ((?v0 S7) (?v1 S7)) (=> (= (f24 (f66 f67 ?v0) (f11 f12 ?v1)) f1) (=> (=> (= (f24 (f66 f67 ?v0) ?v1) f1) false) (=> (=> (= ?v0 ?v1) false) false)))))
(assume a237 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (=> (= (f24 (f66 f67 ?v0) ?v1) f1) (=> (= (f24 (f66 f67 ?v1) ?v2) f1) (= (f24 (f66 f67 (f11 f12 ?v0)) ?v2) f1)))))
(assume a238 (forall ((?v0 S7) (?v1 S7)) (=> (= (f24 (f66 f67 ?v0) ?v1) f1) (=> (not (= (f11 f12 ?v0) ?v1)) (= (f24 (f66 f67 (f11 f12 ?v0)) ?v1) f1)))))
(assume a239 (forall ((?v0 S7) (?v1 S7)) (=> (= (f24 (f66 f67 ?v0) ?v1) f1) (= (f24 (f66 f67 ?v0) (f11 f12 ?v1)) f1))))
(assume a240 (forall ((?v0 S7) (?v1 S7)) (=> (not (= (f24 (f66 f67 ?v0) ?v1) f1)) (=> (= (f24 (f66 f67 ?v0) (f11 f12 ?v1)) f1) (= ?v1 ?v0)))))
(assume a241 (forall ((?v0 S7) (?v1 S7)) (=> (not (= (f24 (f66 f67 ?v0) ?v1) f1)) (= (= (f24 (f66 f67 ?v0) (f11 f12 ?v1)) f1) (= ?v0 ?v1)))))
(assume a242 (forall ((?v0 S7) (?v1 S7)) (= (= (f24 (f66 f67 (f11 f12 ?v0)) (f11 f12 ?v1)) f1) (= (f24 (f66 f67 ?v0) ?v1) f1))))
(assume a243 (forall ((?v0 S7) (?v1 S7)) (= (= (f24 (f66 f67 ?v0) (f11 f12 ?v1)) f1) (or (= (f24 (f66 f67 ?v0) ?v1) f1) (= ?v0 ?v1)))))
(assume a244 (forall ((?v0 S7) (?v1 S7)) (= (not (= (f24 (f66 f67 ?v0) ?v1) f1)) (= (f24 (f66 f67 ?v1) (f11 f12 ?v0)) f1))))
(assume a245 (forall ((?v0 S7) (?v1 S7) (?v2 S7)) (=> (= (f24 (f66 f67 (f11 (f29 f30 ?v0) ?v1)) ?v2) f1) (= (f24 (f66 f67 ?v0) ?v2) f1))))
(assume a246 (forall ((?v0 S7) (?v1 S7) (?v2 S7) (?v3 S7)) (=> (= (f24 (f66 f67 ?v0) ?v1) f1) (=> (= (f11 (f29 f30 ?v2) ?v1) (f11 (f29 f30 ?v0) ?v3)) (= (f24 (f66 f67 ?v2) ?v3) f1)))))
(step t0 (cl (not (! (= (! (=> (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))) (! (or (! (not (! (= (! (f3 (f4 f5 @quantifiers_skolemize_1) @quantifiers_skolemize_2) :named @p_80) @p_3) :named @p_125)) :named @p_126) (! (not (! (= @p_3 (! (f3 (f6 (! (f7 (! (f8 f9 f10) :named @p_18) @quantifiers_skolemize_2) :named @p_76) (! (f11 f12 f13) :named @p_34)) @quantifiers_skolemize_3) :named @p_77)) :named @p_78)) :named @p_79) (! (= @p_3 (f3 (f4 f14 @quantifiers_skolemize_1) @quantifiers_skolemize_3)) :named @p_75)) :named @p_145)) :named @p_146) (! (=> (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))) (! (or (! (not (! (= @p_3 @p_80) :named @p_81)) :named @p_82) (! (not (! (= @p_3 (! (f3 (f6 (! (f7 (! (f8 f9 f10) :named @p_18) @quantifiers_skolemize_2) :named @p_76) (! (f11 f12 f13) :named @p_34)) @quantifiers_skolemize_3) :named @p_77)) :named @p_78)) :named @p_79) (! (= @p_3 (f3 (f4 f14 @quantifiers_skolemize_1) @quantifiers_skolemize_3)) :named @p_75)) :named @p_83)) :named @p_144)) :named @p_147)) (not @p_146) @p_144) :rule equiv_pos2)
(step t1 (cl (= (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) :rule refl)
(step t2 (cl (= @p_125 @p_81)) :rule rare_rewrite :args ("eq-symm" @p_80 @p_3))
(step t3 (cl (= @p_126 @p_82)) :rule cong :premises (t2))
(step t4 (cl (= @p_79 @p_79)) :rule refl)
(step t5 (cl (= @p_75 @p_75)) :rule refl)
(step t6 (cl (= @p_145 @p_83)) :rule cong :premises (t3 t4 t5))
(step t7 (cl @p_147) :rule cong :premises (t1 t6))
(step t8 (cl @p_146 (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))) :rule implies_neg1)
(anchor :step t9)
(assume t9.a0 (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))
(step t9.t0 (cl (or (! (not (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))) :named @p_143) @p_145)) :rule forall_inst :args (@quantifiers_skolemize_1 @quantifiers_skolemize_2 @quantifiers_skolemize_3))
(step t9.t1 (cl @p_143 @p_145) :rule or :premises (t9.t0))
(step t9.t2 (cl (not (! (= @p_2 (! (not (! (=> (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))) :named @p_5)) :named @p_6)) :named @p_7)) (not @p_2) @p_6) :rule equiv_pos2)
(anchor :step t9.t3 :args ((?v0 S2) (?v1 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1)))
(step t9.t3.t0 (cl (= (! (= (! (f3 (f4 f5 ?v0) ?v1) :named @p_24) f1) :named @p_60) (! (= f1 @p_24) :named @p_43))) :rule rare_rewrite :args ("eq-symm" @p_24 f1))
(anchor :step t9.t3.t1 :args ((?v2 S3) (:= (?v2 S3) ?v2)))
(step t9.t3.t1.t0 (cl (= (! (= (! (f3 (! (f6 (! (f7 @p_18 ?v1) :named @p_19) @p_34) :named @p_35) ?v2) :named @p_69) f1) :named @p_74) (! (= f1 @p_69) :named @p_70))) :rule rare_rewrite :args ("eq-symm" @p_69 f1))
(step t9.t3.t1.t1 (cl (= (! (= (! (f3 (! (f4 f14 ?v0) :named @p_15) ?v2) :named @p_54) f1) :named @p_63) (! (= f1 @p_54) :named @p_55))) :rule rare_rewrite :args ("eq-symm" @p_54 f1))
(step t9.t3.t1.t2 (cl (= (=> @p_74 @p_63) (! (=> @p_70 @p_55) :named @p_73))) :rule cong :premises (t9.t3.t1.t0 t9.t3.t1.t1))
(step t9.t3.t1 (cl (= (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))) (forall ((?v2 S3)) (=> (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2)) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule bind)
(anchor :step t9.t3.t2 :args ((?v2 S3) (:= (?v2 S3) ?v2)))
(step t9.t3.t2.t0 (cl (= @p_73 (! (or (not @p_70) @p_55) :named @p_71))) :rule rare_rewrite :args ("bool-impl-elim" @p_70 @p_55))
(step t9.t3.t2 (cl (= (forall ((?v2 S3)) (=> (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2)) (= f1 (f3 (f4 f14 ?v0) ?v2)))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule bind)
(step t9.t3.t3 (cl (= (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule trans :premises (t9.t3.t1 t9.t3.t2))
(step t9.t3.t4 (cl (= (=> @p_60 (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1)))) (! (=> @p_43 (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2))))) :named @p_72))) :rule cong :premises (t9.t3.t0 t9.t3.t3))
(step t9.t3 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= (f3 (f4 f5 ?v0) ?v1) f1) (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))))) (forall ((?v0 S2) (?v1 S3)) (=> (= f1 (f3 (f4 f5 ?v0) ?v1)) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))))) :rule bind)
(anchor :step t9.t4 :args ((?v0 S2) (?v1 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1)))
(step t9.t4.t0 (cl (= @p_72 (! (or (! (not @p_43) :named @p_44) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2))))) :named @p_67))) :rule rare_rewrite :args ("bool-impl-elim" @p_43 (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2))))))
(step t9.t4 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= f1 (f3 (f4 f5 ?v0) ?v1)) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))))) :rule bind)
(anchor :step t9.t5 :args ((?v0 S2) (?v1 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1)))
(step t9.t5.t0 (cl (= (forall ((BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))) (! (or @p_44 (forall ((BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))) :named @p_68))) :rule hole)
(step t9.t5.t1 (cl (= @p_44 @p_44)) :rule refl)
(anchor :step t9.t5.t2 :args ((?v2 S3) (:= (BOUND_VARIABLE_8816 S3) ?v2)))
(step t9.t5.t2.t0 (cl (= (! (or (! (not (! (= f1 (! (f3 @p_35 BOUND_VARIABLE_8816) :named @p_36)) :named @p_47)) :named @p_48) (! (= f1 (! (f3 @p_15 BOUND_VARIABLE_8816) :named @p_32)) :named @p_46)) :named @p_65) @p_71)) :rule refl)
(step t9.t5.t2 (cl (= (forall ((BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule bind)
(step t9.t5.t3 (cl (= @p_68 @p_67)) :rule cong :premises (t9.t5.t1 t9.t5.t2))
(step t9.t5.t4 (cl (= (forall ((BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))) @p_67)) :rule trans :premises (t9.t5.t0 t9.t5.t3))
(step t9.t5.t5 (cl (= @p_67 (forall ((BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))))) :rule symm :premises (t9.t5.t4))
(step t9.t5 (cl (= (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3)) (forall ((BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))))) :rule bind)
(step t9.t6 (cl (= (forall ((?v0 S2) (?v1 S3)) (forall ((BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))))) :rule qnt_join)
(step t9.t7 (cl (= (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))))) :rule trans :premises (t9.t5 t9.t6))
(anchor :step t9.t8 :args ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1) (:= (BOUND_VARIABLE_8816 S3) BOUND_VARIABLE_8816)))
(step t9.t8.t0 (cl (! (= (or @p_44 @p_65) (! (or @p_44 @p_48 @p_46) :named @p_49)) :named @p_66)) :rule hole :args ("ACI_NORM" @p_66))
(step t9.t8 (cl (= (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) :rule bind)
(step t9.t9 (cl (= (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) :rule trans :premises (t9.t7 t9.t8))
(step t9.t10 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= f1 (f3 (f4 f5 ?v0) ?v1)) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) :rule trans :premises (t9.t4 t9.t9))
(step t9.t11 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= (f3 (f4 f5 ?v0) ?v1) f1) (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) :rule trans :premises (t9.t3 t9.t10))
(anchor :step t9.t12 :args ((?v0 S2) (?v1 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1)))
(step t9.t12.t0 (cl (= @p_60 @p_43)) :rule rare_rewrite :args ("eq-symm" @p_24 f1))
(anchor :step t9.t12.t1 :args ((?v2 S3) (:= (?v2 S3) ?v2)))
(step t9.t12.t1.t0 (cl (= (! (= (! (f3 (! (f6 @p_19 f13) :named @p_20) ?v2) :named @p_56) f1) :named @p_64) (! (= f1 @p_56) :named @p_57))) :rule rare_rewrite :args ("eq-symm" @p_56 f1))
(step t9.t12.t1.t1 (cl (= @p_63 @p_55)) :rule rare_rewrite :args ("eq-symm" @p_54 f1))
(step t9.t12.t1.t2 (cl (= (=> @p_64 @p_63) (! (=> @p_57 @p_55) :named @p_62))) :rule cong :premises (t9.t12.t1.t0 t9.t12.t1.t1))
(step t9.t12.t1 (cl (= (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))) (forall ((?v2 S3)) (=> (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2)) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule bind)
(anchor :step t9.t12.t2 :args ((?v2 S3) (:= (?v2 S3) ?v2)))
(step t9.t12.t2.t0 (cl (= @p_62 (! (or (not @p_57) @p_55) :named @p_58))) :rule rare_rewrite :args ("bool-impl-elim" @p_57 @p_55))
(step t9.t12.t2 (cl (= (forall ((?v2 S3)) (=> (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2)) (= f1 (f3 (f4 f14 ?v0) ?v2)))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule bind)
(step t9.t12.t3 (cl (= (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule trans :premises (t9.t12.t1 t9.t12.t2))
(step t9.t12.t4 (cl (= (=> @p_60 (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1)))) (! (=> @p_43 (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2))))) :named @p_59))) :rule cong :premises (t9.t12.t0 t9.t12.t3))
(step t9.t12 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= (f3 (f4 f5 ?v0) ?v1) f1) (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))))) (forall ((?v0 S2) (?v1 S3)) (=> (= f1 (f3 (f4 f5 ?v0) ?v1)) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))))) :rule bind)
(anchor :step t9.t13 :args ((?v0 S2) (?v1 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1)))
(step t9.t13.t0 (cl (= @p_59 (! (or @p_44 (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2))))) :named @p_52))) :rule rare_rewrite :args ("bool-impl-elim" @p_43 (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2))))))
(step t9.t13 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= f1 (f3 (f4 f5 ?v0) ?v1)) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))))) :rule bind)
(anchor :step t9.t14 :args ((?v0 S2) (?v1 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1)))
(step t9.t14.t0 (cl (= (forall ((BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))) (! (or @p_44 (forall ((BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))) :named @p_53))) :rule hole)
(step t9.t14.t1 (cl (= @p_44 @p_44)) :rule refl)
(anchor :step t9.t14.t2 :args ((?v2 S3) (:= (BOUND_VARIABLE_8839 S3) ?v2)))
(step t9.t14.t2.t0 (cl (= (! (or (! (not (! (= f1 (! (f3 @p_20 BOUND_VARIABLE_8839) :named @p_21)) :named @p_41)) :named @p_42) (! (= f1 (! (f3 @p_15 BOUND_VARIABLE_8839) :named @p_16)) :named @p_40)) :named @p_50) @p_58)) :rule refl)
(step t9.t14.t2 (cl (= (forall ((BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule bind)
(step t9.t14.t3 (cl (= @p_53 @p_52)) :rule cong :premises (t9.t14.t1 t9.t14.t2))
(step t9.t14.t4 (cl (= (forall ((BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))) @p_52)) :rule trans :premises (t9.t14.t0 t9.t14.t3))
(step t9.t14.t5 (cl (= @p_52 (forall ((BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))))) :rule symm :premises (t9.t14.t4))
(step t9.t14 (cl (= (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3)) (forall ((BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))))) :rule bind)
(step t9.t15 (cl (= (forall ((?v0 S2) (?v1 S3)) (forall ((BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))))) :rule qnt_join)
(step t9.t16 (cl (= (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))))) :rule trans :premises (t9.t14 t9.t15))
(anchor :step t9.t17 :args ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1) (:= (BOUND_VARIABLE_8839 S3) BOUND_VARIABLE_8839)))
(step t9.t17.t0 (cl (! (= (or @p_44 @p_50) (! (or @p_44 @p_42 @p_40) :named @p_45)) :named @p_51)) :rule hole :args ("ACI_NORM" @p_51))
(step t9.t17 (cl (= (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) :rule bind)
(step t9.t18 (cl (= (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) :rule trans :premises (t9.t16 t9.t17))
(step t9.t19 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= f1 (f3 (f4 f5 ?v0) ?v1)) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) :rule trans :premises (t9.t13 t9.t18))
(step t9.t20 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= (f3 (f4 f5 ?v0) ?v1) f1) (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) :rule trans :premises (t9.t12 t9.t19))
(step t9.t21 (cl (= @p_1 (! (=> (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))) :named @p_8))) :rule cong :premises (t9.t11 t9.t20))
(step t9.t22 (cl (= @p_2 (! (not @p_8) :named @p_9))) :rule cong :premises (t9.t21))
(anchor :step t9.t23 :args ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1) (:= (BOUND_VARIABLE_8816 S3) BOUND_VARIABLE_8816)))
(step t9.t23.t0 (cl (= f1 @p_3)) :rule symm :premises (a140))
(step t9.t23.t1 (cl (= @p_24 @p_24)) :rule refl)
(step t9.t23.t2 (cl (= @p_43 (! (= @p_3 @p_24) :named @p_28))) :rule cong :premises (t9.t23.t0 t9.t23.t1))
(step t9.t23.t3 (cl (= @p_44 (! (not @p_28) :named @p_29))) :rule cong :premises (t9.t23.t2))
(step t9.t23.t4 (cl (= @p_36 @p_36)) :rule refl)
(step t9.t23.t5 (cl (= @p_47 (! (= @p_3 @p_36) :named @p_37))) :rule cong :premises (t9.t23.t0 t9.t23.t4))
(step t9.t23.t6 (cl (= @p_48 (! (not @p_37) :named @p_38))) :rule cong :premises (t9.t23.t5))
(step t9.t23.t7 (cl (= @p_32 @p_32)) :rule refl)
(step t9.t23.t8 (cl (= @p_46 (! (= @p_3 @p_32) :named @p_33))) :rule cong :premises (t9.t23.t0 t9.t23.t7))
(step t9.t23.t9 (cl (= @p_49 (! (or @p_29 @p_38 @p_33) :named @p_39))) :rule cong :premises (t9.t23.t3 t9.t23.t6 t9.t23.t8))
(step t9.t23 (cl (= (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= (f37 (f38 f39 f28) f57) (f3 (f4 f5 ?v0) ?v1))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) :rule bind)
(anchor :step t9.t24 :args ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1) (:= (BOUND_VARIABLE_8839 S3) BOUND_VARIABLE_8839)))
(step t9.t24.t0 (cl (= f1 @p_3)) :rule symm :premises (a140))
(step t9.t24.t1 (cl (= @p_24 @p_24)) :rule refl)
(step t9.t24.t2 (cl (= @p_43 @p_28)) :rule cong :premises (t9.t24.t0 t9.t24.t1))
(step t9.t24.t3 (cl (= @p_44 @p_29)) :rule cong :premises (t9.t24.t2))
(step t9.t24.t4 (cl (= @p_21 @p_21)) :rule refl)
(step t9.t24.t5 (cl (= @p_41 (! (= @p_3 @p_21) :named @p_22))) :rule cong :premises (t9.t24.t0 t9.t24.t4))
(step t9.t24.t6 (cl (= @p_42 (! (not @p_22) :named @p_23))) :rule cong :premises (t9.t24.t5))
(step t9.t24.t7 (cl (= @p_16 @p_16)) :rule refl)
(step t9.t24.t8 (cl (= @p_40 (! (= @p_3 @p_16) :named @p_17))) :rule cong :premises (t9.t24.t0 t9.t24.t7))
(step t9.t24.t9 (cl (= @p_45 (! (or @p_29 @p_23 @p_17) :named @p_30))) :rule cong :premises (t9.t24.t3 t9.t24.t6 t9.t24.t8))
(step t9.t24 (cl (= (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= (f37 (f38 f39 f28) f57) (f3 (f4 f5 ?v0) ?v1))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) :rule bind)
(step t9.t25 (cl (= @p_8 (! (=> (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= (f37 (f38 f39 f28) f57) (f3 (f4 f5 ?v0) ?v1))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= (f37 (f38 f39 f28) f57) (f3 (f4 f5 ?v0) ?v1))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))) :named @p_10))) :rule cong :premises (t9.t23 t9.t24))
(step t9.t26 (cl (= @p_9 (! (not @p_10) :named @p_11))) :rule cong :premises (t9.t25))
(anchor :step t9.t27 :args ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1) (:= (BOUND_VARIABLE_8816 S3) BOUND_VARIABLE_8816)))
(step t9.t27.t0 (cl (= @p_28 (! (= @p_24 @p_3) :named @p_25))) :rule rare_rewrite :args ("eq-symm" @p_3 @p_24))
(step t9.t27.t1 (cl (= @p_29 (! (not @p_25) :named @p_26))) :rule cong :premises (t9.t27.t0))
(step t9.t27.t2 (cl (= @p_38 @p_38)) :rule refl)
(step t9.t27.t3 (cl (= @p_33 @p_33)) :rule refl)
(step t9.t27.t4 (cl (= @p_39 (or @p_26 @p_38 @p_33))) :rule cong :premises (t9.t27.t1 t9.t27.t2 t9.t27.t3))
(step t9.t27 (cl (= (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= (f37 (f38 f39 f28) f57) (f3 (f4 f5 ?v0) ?v1))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) :rule bind)
(anchor :step t9.t28 :args ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1) (:= (BOUND_VARIABLE_8839 S3) BOUND_VARIABLE_8839)))
(step t9.t28.t0 (cl (= @p_28 @p_25)) :rule rare_rewrite :args ("eq-symm" @p_3 @p_24))
(step t9.t28.t1 (cl (= @p_29 @p_26)) :rule cong :premises (t9.t28.t0))
(step t9.t28.t2 (cl (= @p_23 @p_23)) :rule refl)
(step t9.t28.t3 (cl (= @p_17 @p_17)) :rule refl)
(step t9.t28.t4 (cl (= @p_30 (! (or @p_26 @p_23 @p_17) :named @p_27))) :rule cong :premises (t9.t28.t1 t9.t28.t2 t9.t28.t3))
(step t9.t28 (cl (= (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= (f37 (f38 f39 f28) f57) (f3 (f4 f5 ?v0) ?v1))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) :rule bind)
(step t9.t29 (cl (= @p_10 @p_5)) :rule cong :premises (t9.t27 t9.t28))
(step t9.t30 (cl (= @p_11 @p_6)) :rule cong :premises (t9.t29))
(step t9.t31 (cl (= @p_9 @p_6)) :rule trans :premises (t9.t26 t9.t30))
(step t9.t32 (cl @p_7) :rule trans :premises (t9.t22 t9.t31))
(step t9.t33 (cl @p_6) :rule resolution :premises (t9.t2 t9.t32 a1))
(step t9.t34 (cl (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))) :rule not_implies1 :premises (t9.t33))
(step t9.t35 (cl @p_145) :rule resolution :premises (t9.t1 t9.t34))
(step t9 (cl @p_143 @p_145) :rule subproof :discharge (t9.a0))
(step t10 (cl @p_146 @p_145) :rule resolution :premises (t8 t9))
(step t11 (cl @p_146 (not @p_145)) :rule implies_neg2)
(step t12 (cl @p_146 @p_146) :rule resolution :premises (t10 t11))
(step t13 (cl @p_146) :rule contraction :premises (t12))
(step t14 (cl @p_144) :rule resolution :premises (t0 t7 t13))
(step t15 (cl @p_143 @p_83) :rule implies :premises (t14))
(step t16 (cl (! (not @p_83) :named @p_84) @p_82 @p_79 @p_75) :rule or_pos)
(step t17 (cl @p_75 @p_82 @p_79 @p_84) :rule reordering :premises (t16))
(step t18 (cl (! (or @p_82 (! (not (! (= @p_3 (! (f3 (f6 @p_76 f13) @quantifiers_skolemize_3) :named @p_85)) :named @p_86)) :named @p_87) @p_75) :named @p_118) (! (not @p_75) :named @p_142)) :rule or_neg :args (2))
(step t19 (cl (not (! (= (! (or (! (not (! (not (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))) :named @p_120)) :named @p_122) (! (not @p_118) :named @p_119)) :named @p_123) (! (or (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))) @p_119) :named @p_121)) :named @p_132)) (not @p_123) @p_121) :rule equiv_pos2)
(step t20 (cl (= @p_122 (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) :rule rare_rewrite :args ("bool-double-not-elim" (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))))
(step t21 (cl (= @p_119 @p_119)) :rule refl)
(step t22 (cl @p_132) :rule cong :premises (t20 t21))
(step t23 (cl (not (! (= (! (=> @p_120 (! (not (! (or @p_126 @p_87 @p_75) :named @p_127)) :named @p_128)) :named @p_129) (! (=> @p_120 @p_119) :named @p_124)) :named @p_131)) (not @p_129) @p_124) :rule equiv_pos2)
(step t24 (cl (= @p_120 @p_120)) :rule refl)
(step t25 (cl (= @p_87 @p_87)) :rule refl)
(step t26 (cl (= @p_127 @p_118)) :rule cong :premises (t3 t25 t5))
(step t27 (cl (= @p_128 @p_119)) :rule cong :premises (t26))
(step t28 (cl @p_131) :rule cong :premises (t24 t27))
(step t29 (cl @p_129 @p_120) :rule implies_neg1)
(anchor :step t30)
(assume t30.a0 @p_120)
(step t30.t0 (cl (not (! (= @p_120 @p_128) :named @p_130)) @p_122 @p_128) :rule equiv_pos2)
(anchor :step t30.t1 :args ((:= (?v0 S2) @quantifiers_skolemize_1) (:= (?v1 S3) @quantifiers_skolemize_2) (:= (BOUND_VARIABLE_8839 S3) @quantifiers_skolemize_3)))
(step t30.t1.t0 (cl (= @p_27 @p_127)) :rule refl)
(step t30.t1 (cl (= (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))) @p_127)) :rule sko_forall)
(step t30.t2 (cl @p_130) :rule cong :premises (t30.t1))
(step t30.t3 (cl (not @p_7) (not @p_2) @p_6) :rule equiv_pos2)
(anchor :step t30.t4 :args ((?v0 S2) (?v1 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1)))
(step t30.t4.t0 (cl (= @p_60 @p_43)) :rule rare_rewrite :args ("eq-symm" @p_24 f1))
(anchor :step t30.t4.t1 :args ((?v2 S3) (:= (?v2 S3) ?v2)))
(step t30.t4.t1.t0 (cl (= @p_74 @p_70)) :rule rare_rewrite :args ("eq-symm" @p_69 f1))
(step t30.t4.t1.t1 (cl (= @p_63 @p_55)) :rule rare_rewrite :args ("eq-symm" @p_54 f1))
(step t30.t4.t1.t2 (cl (= (=> @p_74 @p_63) @p_73)) :rule cong :premises (t30.t4.t1.t0 t30.t4.t1.t1))
(step t30.t4.t1 (cl (= (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))) (forall ((?v2 S3)) (=> (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2)) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule bind)
(anchor :step t30.t4.t2 :args ((?v2 S3) (:= (?v2 S3) ?v2)))
(step t30.t4.t2.t0 (cl (= @p_73 @p_71)) :rule rare_rewrite :args ("bool-impl-elim" @p_70 @p_55))
(step t30.t4.t2 (cl (= (forall ((?v2 S3)) (=> (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2)) (= f1 (f3 (f4 f14 ?v0) ?v2)))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule bind)
(step t30.t4.t3 (cl (= (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule trans :premises (t30.t4.t1 t30.t4.t2))
(step t30.t4.t4 (cl (= (=> @p_60 (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1)))) @p_72)) :rule cong :premises (t30.t4.t0 t30.t4.t3))
(step t30.t4 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= (f3 (f4 f5 ?v0) ?v1) f1) (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))))) (forall ((?v0 S2) (?v1 S3)) (=> (= f1 (f3 (f4 f5 ?v0) ?v1)) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))))) :rule bind)
(anchor :step t30.t5 :args ((?v0 S2) (?v1 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1)))
(step t30.t5.t0 (cl (= @p_72 @p_67)) :rule rare_rewrite :args ("bool-impl-elim" @p_43 (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2))))))
(step t30.t5 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= f1 (f3 (f4 f5 ?v0) ?v1)) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))))) :rule bind)
(anchor :step t30.t6 :args ((?v0 S2) (?v1 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1)))
(step t30.t6.t0 (cl (= (forall ((BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))) @p_68)) :rule hole)
(step t30.t6.t1 (cl (= @p_44 @p_44)) :rule refl)
(anchor :step t30.t6.t2 :args ((?v2 S3) (:= (BOUND_VARIABLE_8816 S3) ?v2)))
(step t30.t6.t2.t0 (cl (= @p_65 @p_71)) :rule refl)
(step t30.t6.t2 (cl (= (forall ((BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule bind)
(step t30.t6.t3 (cl (= @p_68 @p_67)) :rule cong :premises (t30.t6.t1 t30.t6.t2))
(step t30.t6.t4 (cl (= (forall ((BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))) @p_67)) :rule trans :premises (t30.t6.t0 t30.t6.t3))
(step t30.t6.t5 (cl (= @p_67 (forall ((BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))))) :rule symm :premises (t30.t6.t4))
(step t30.t6 (cl (= (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3)) (forall ((BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))))) :rule bind)
(step t30.t7 (cl (= (forall ((?v0 S2) (?v1 S3)) (forall ((BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))))) :rule qnt_join)
(step t30.t8 (cl (= (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))))) :rule trans :premises (t30.t6 t30.t7))
(anchor :step t30.t9 :args ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1) (:= (BOUND_VARIABLE_8816 S3) BOUND_VARIABLE_8816)))
(step t30.t9.t0 (cl @p_66) :rule hole :args ("ACI_NORM" @p_66))
(step t30.t9 (cl (= (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) :rule bind)
(step t30.t10 (cl (= (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) :rule trans :premises (t30.t8 t30.t9))
(step t30.t11 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= f1 (f3 (f4 f5 ?v0) ?v1)) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) :rule trans :premises (t30.t5 t30.t10))
(step t30.t12 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= (f3 (f4 f5 ?v0) ?v1) f1) (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) :rule trans :premises (t30.t4 t30.t11))
(anchor :step t30.t13 :args ((?v0 S2) (?v1 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1)))
(step t30.t13.t0 (cl (= @p_60 @p_43)) :rule rare_rewrite :args ("eq-symm" @p_24 f1))
(anchor :step t30.t13.t1 :args ((?v2 S3) (:= (?v2 S3) ?v2)))
(step t30.t13.t1.t0 (cl (= @p_64 @p_57)) :rule rare_rewrite :args ("eq-symm" @p_56 f1))
(step t30.t13.t1.t1 (cl (= @p_63 @p_55)) :rule rare_rewrite :args ("eq-symm" @p_54 f1))
(step t30.t13.t1.t2 (cl (= (=> @p_64 @p_63) @p_62)) :rule cong :premises (t30.t13.t1.t0 t30.t13.t1.t1))
(step t30.t13.t1 (cl (= (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))) (forall ((?v2 S3)) (=> (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2)) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule bind)
(anchor :step t30.t13.t2 :args ((?v2 S3) (:= (?v2 S3) ?v2)))
(step t30.t13.t2.t0 (cl (= @p_62 @p_58)) :rule rare_rewrite :args ("bool-impl-elim" @p_57 @p_55))
(step t30.t13.t2 (cl (= (forall ((?v2 S3)) (=> (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2)) (= f1 (f3 (f4 f14 ?v0) ?v2)))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule bind)
(step t30.t13.t3 (cl (= (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule trans :premises (t30.t13.t1 t30.t13.t2))
(step t30.t13.t4 (cl (= (=> @p_60 (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1)))) @p_59)) :rule cong :premises (t30.t13.t0 t30.t13.t3))
(step t30.t13 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= (f3 (f4 f5 ?v0) ?v1) f1) (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))))) (forall ((?v0 S2) (?v1 S3)) (=> (= f1 (f3 (f4 f5 ?v0) ?v1)) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))))) :rule bind)
(anchor :step t30.t14 :args ((?v0 S2) (?v1 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1)))
(step t30.t14.t0 (cl (= @p_59 @p_52)) :rule rare_rewrite :args ("bool-impl-elim" @p_43 (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2))))))
(step t30.t14 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= f1 (f3 (f4 f5 ?v0) ?v1)) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))))) :rule bind)
(anchor :step t30.t15 :args ((?v0 S2) (?v1 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1)))
(step t30.t15.t0 (cl (= (forall ((BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))) @p_53)) :rule hole)
(step t30.t15.t1 (cl (= @p_44 @p_44)) :rule refl)
(anchor :step t30.t15.t2 :args ((?v2 S3) (:= (BOUND_VARIABLE_8839 S3) ?v2)))
(step t30.t15.t2.t0 (cl (= @p_50 @p_58)) :rule refl)
(step t30.t15.t2 (cl (= (forall ((BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule bind)
(step t30.t15.t3 (cl (= @p_53 @p_52)) :rule cong :premises (t30.t15.t1 t30.t15.t2))
(step t30.t15.t4 (cl (= (forall ((BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))) @p_52)) :rule trans :premises (t30.t15.t0 t30.t15.t3))
(step t30.t15.t5 (cl (= @p_52 (forall ((BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))))) :rule symm :premises (t30.t15.t4))
(step t30.t15 (cl (= (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3)) (forall ((BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))))) :rule bind)
(step t30.t16 (cl (= (forall ((?v0 S2) (?v1 S3)) (forall ((BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))))) :rule qnt_join)
(step t30.t17 (cl (= (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))))) :rule trans :premises (t30.t15 t30.t16))
(anchor :step t30.t18 :args ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1) (:= (BOUND_VARIABLE_8839 S3) BOUND_VARIABLE_8839)))
(step t30.t18.t0 (cl @p_51) :rule hole :args ("ACI_NORM" @p_51))
(step t30.t18 (cl (= (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) :rule bind)
(step t30.t19 (cl (= (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) :rule trans :premises (t30.t17 t30.t18))
(step t30.t20 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= f1 (f3 (f4 f5 ?v0) ?v1)) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) :rule trans :premises (t30.t14 t30.t19))
(step t30.t21 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= (f3 (f4 f5 ?v0) ?v1) f1) (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) :rule trans :premises (t30.t13 t30.t20))
(step t30.t22 (cl (= @p_1 @p_8)) :rule cong :premises (t30.t12 t30.t21))
(step t30.t23 (cl (= @p_2 @p_9)) :rule cong :premises (t30.t22))
(anchor :step t30.t24 :args ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1) (:= (BOUND_VARIABLE_8816 S3) BOUND_VARIABLE_8816)))
(step t30.t24.t0 (cl (= f1 @p_3)) :rule symm :premises (a140))
(step t30.t24.t1 (cl (= @p_24 @p_24)) :rule refl)
(step t30.t24.t2 (cl (= @p_43 @p_28)) :rule cong :premises (t30.t24.t0 t30.t24.t1))
(step t30.t24.t3 (cl (= @p_44 @p_29)) :rule cong :premises (t30.t24.t2))
(step t30.t24.t4 (cl (= @p_36 @p_36)) :rule refl)
(step t30.t24.t5 (cl (= @p_47 @p_37)) :rule cong :premises (t30.t24.t0 t30.t24.t4))
(step t30.t24.t6 (cl (= @p_48 @p_38)) :rule cong :premises (t30.t24.t5))
(step t30.t24.t7 (cl (= @p_32 @p_32)) :rule refl)
(step t30.t24.t8 (cl (= @p_46 @p_33)) :rule cong :premises (t30.t24.t0 t30.t24.t7))
(step t30.t24.t9 (cl (= @p_49 @p_39)) :rule cong :premises (t30.t24.t3 t30.t24.t6 t30.t24.t8))
(step t30.t24 (cl (= (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= (f37 (f38 f39 f28) f57) (f3 (f4 f5 ?v0) ?v1))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) :rule bind)
(anchor :step t30.t25 :args ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1) (:= (BOUND_VARIABLE_8839 S3) BOUND_VARIABLE_8839)))
(step t30.t25.t0 (cl (= f1 @p_3)) :rule symm :premises (a140))
(step t30.t25.t1 (cl (= @p_24 @p_24)) :rule refl)
(step t30.t25.t2 (cl (= @p_43 @p_28)) :rule cong :premises (t30.t25.t0 t30.t25.t1))
(step t30.t25.t3 (cl (= @p_44 @p_29)) :rule cong :premises (t30.t25.t2))
(step t30.t25.t4 (cl (= @p_21 @p_21)) :rule refl)
(step t30.t25.t5 (cl (= @p_41 @p_22)) :rule cong :premises (t30.t25.t0 t30.t25.t4))
(step t30.t25.t6 (cl (= @p_42 @p_23)) :rule cong :premises (t30.t25.t5))
(step t30.t25.t7 (cl (= @p_16 @p_16)) :rule refl)
(step t30.t25.t8 (cl (= @p_40 @p_17)) :rule cong :premises (t30.t25.t0 t30.t25.t7))
(step t30.t25.t9 (cl (= @p_45 @p_30)) :rule cong :premises (t30.t25.t3 t30.t25.t6 t30.t25.t8))
(step t30.t25 (cl (= (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= (f37 (f38 f39 f28) f57) (f3 (f4 f5 ?v0) ?v1))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) :rule bind)
(step t30.t26 (cl (= @p_8 @p_10)) :rule cong :premises (t30.t24 t30.t25))
(step t30.t27 (cl (= @p_9 @p_11)) :rule cong :premises (t30.t26))
(anchor :step t30.t28 :args ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1) (:= (BOUND_VARIABLE_8816 S3) BOUND_VARIABLE_8816)))
(step t30.t28.t0 (cl (= @p_28 @p_25)) :rule rare_rewrite :args ("eq-symm" @p_3 @p_24))
(step t30.t28.t1 (cl (= @p_29 @p_26)) :rule cong :premises (t30.t28.t0))
(step t30.t28.t2 (cl (= @p_38 @p_38)) :rule refl)
(step t30.t28.t3 (cl (= @p_33 @p_33)) :rule refl)
(step t30.t28.t4 (cl (= @p_39 (or @p_26 @p_38 @p_33))) :rule cong :premises (t30.t28.t1 t30.t28.t2 t30.t28.t3))
(step t30.t28 (cl (= (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= (f37 (f38 f39 f28) f57) (f3 (f4 f5 ?v0) ?v1))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) :rule bind)
(anchor :step t30.t29 :args ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1) (:= (BOUND_VARIABLE_8839 S3) BOUND_VARIABLE_8839)))
(step t30.t29.t0 (cl (= @p_28 @p_25)) :rule rare_rewrite :args ("eq-symm" @p_3 @p_24))
(step t30.t29.t1 (cl (= @p_29 @p_26)) :rule cong :premises (t30.t29.t0))
(step t30.t29.t2 (cl (= @p_23 @p_23)) :rule refl)
(step t30.t29.t3 (cl (= @p_17 @p_17)) :rule refl)
(step t30.t29.t4 (cl (= @p_30 @p_27)) :rule cong :premises (t30.t29.t1 t30.t29.t2 t30.t29.t3))
(step t30.t29 (cl (= (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= (f37 (f38 f39 f28) f57) (f3 (f4 f5 ?v0) ?v1))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) :rule bind)
(step t30.t30 (cl (= @p_10 @p_5)) :rule cong :premises (t30.t28 t30.t29))
(step t30.t31 (cl (= @p_11 @p_6)) :rule cong :premises (t30.t30))
(step t30.t32 (cl (= @p_9 @p_6)) :rule trans :premises (t30.t27 t30.t31))
(step t30.t33 (cl @p_7) :rule trans :premises (t30.t23 t30.t32))
(step t30.t34 (cl @p_6) :rule resolution :premises (t30.t3 t30.t33 a1))
(step t30.t35 (cl @p_120) :rule not_implies2 :premises (t30.t34))
(step t30.t36 (cl @p_128) :rule resolution :premises (t30.t0 t30.t2 t30.t35))
(step t30 (cl @p_122 @p_128) :rule subproof :discharge (t30.a0))
(step t31 (cl @p_129 @p_128) :rule resolution :premises (t29 t30))
(step t32 (cl @p_129 (not @p_128)) :rule implies_neg2)
(step t33 (cl @p_129 @p_129) :rule resolution :premises (t31 t32))
(step t34 (cl @p_129) :rule contraction :premises (t33))
(step t35 (cl @p_124) :rule resolution :premises (t23 t28 t34))
(step t36 (cl @p_122 @p_119) :rule implies :premises (t35))
(step t37 (cl @p_123 (not @p_122)) :rule or_neg :args (0))
(step t38 (cl @p_123 (not @p_119)) :rule or_neg :args (1))
(step t39 (cl @p_123 @p_123) :rule resolution :premises (t36 t37 t38))
(step t40 (cl @p_123) :rule contraction :premises (t39))
(step t41 (cl @p_121) :rule resolution :premises (t19 t22 t40))
(step t42 (cl (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))) @p_119) :rule or :premises (t41))
(step t43 (cl (not @p_7) (not @p_2) @p_6) :rule equiv_pos2)
(anchor :step t44 :args ((?v0 S2) (?v1 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1)))
(step t44.t0 (cl (= @p_60 @p_43)) :rule rare_rewrite :args ("eq-symm" @p_24 f1))
(anchor :step t44.t1 :args ((?v2 S3) (:= (?v2 S3) ?v2)))
(step t44.t1.t0 (cl (= @p_74 @p_70)) :rule rare_rewrite :args ("eq-symm" @p_69 f1))
(step t44.t1.t1 (cl (= @p_63 @p_55)) :rule rare_rewrite :args ("eq-symm" @p_54 f1))
(step t44.t1.t2 (cl (= (=> @p_74 @p_63) @p_73)) :rule cong :premises (t44.t1.t0 t44.t1.t1))
(step t44.t1 (cl (= (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))) (forall ((?v2 S3)) (=> (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2)) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule bind)
(anchor :step t44.t2 :args ((?v2 S3) (:= (?v2 S3) ?v2)))
(step t44.t2.t0 (cl (= @p_73 @p_71)) :rule rare_rewrite :args ("bool-impl-elim" @p_70 @p_55))
(step t44.t2 (cl (= (forall ((?v2 S3)) (=> (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2)) (= f1 (f3 (f4 f14 ?v0) ?v2)))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule bind)
(step t44.t3 (cl (= (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule trans :premises (t44.t1 t44.t2))
(step t44.t4 (cl (= (=> @p_60 (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1)))) @p_72)) :rule cong :premises (t44.t0 t44.t3))
(step t44 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= (f3 (f4 f5 ?v0) ?v1) f1) (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))))) (forall ((?v0 S2) (?v1 S3)) (=> (= f1 (f3 (f4 f5 ?v0) ?v1)) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))))) :rule bind)
(anchor :step t45 :args ((?v0 S2) (?v1 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1)))
(step t45.t0 (cl (= @p_72 @p_67)) :rule rare_rewrite :args ("bool-impl-elim" @p_43 (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2))))))
(step t45 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= f1 (f3 (f4 f5 ?v0) ?v1)) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))))) :rule bind)
(anchor :step t46 :args ((?v0 S2) (?v1 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1)))
(step t46.t0 (cl (= (forall ((BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))) @p_68)) :rule hole)
(step t46.t1 (cl (= @p_44 @p_44)) :rule refl)
(anchor :step t46.t2 :args ((?v2 S3) (:= (BOUND_VARIABLE_8816 S3) ?v2)))
(step t46.t2.t0 (cl (= @p_65 @p_71)) :rule refl)
(step t46.t2 (cl (= (forall ((BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule bind)
(step t46.t3 (cl (= @p_68 @p_67)) :rule cong :premises (t46.t1 t46.t2))
(step t46.t4 (cl (= (forall ((BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))) @p_67)) :rule trans :premises (t46.t0 t46.t3))
(step t46.t5 (cl (= @p_67 (forall ((BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))))) :rule symm :premises (t46.t4))
(step t46 (cl (= (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3)) (forall ((BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))))) :rule bind)
(step t47 (cl (= (forall ((?v0 S2) (?v1 S3)) (forall ((BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))))) :rule qnt_join)
(step t48 (cl (= (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))))) :rule trans :premises (t46 t47))
(anchor :step t49 :args ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1) (:= (BOUND_VARIABLE_8816 S3) BOUND_VARIABLE_8816)))
(step t49.t0 (cl @p_66) :rule hole :args ("ACI_NORM" @p_66))
(step t49 (cl (= (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) :rule bind)
(step t50 (cl (= (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) :rule trans :premises (t48 t49))
(step t51 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= f1 (f3 (f4 f5 ?v0) ?v1)) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) :rule trans :premises (t45 t50))
(step t52 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= (f3 (f4 f5 ?v0) ?v1) f1) (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) :rule trans :premises (t44 t51))
(anchor :step t53 :args ((?v0 S2) (?v1 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1)))
(step t53.t0 (cl (= @p_60 @p_43)) :rule rare_rewrite :args ("eq-symm" @p_24 f1))
(anchor :step t53.t1 :args ((?v2 S3) (:= (?v2 S3) ?v2)))
(step t53.t1.t0 (cl (= @p_64 @p_57)) :rule rare_rewrite :args ("eq-symm" @p_56 f1))
(step t53.t1.t1 (cl (= @p_63 @p_55)) :rule rare_rewrite :args ("eq-symm" @p_54 f1))
(step t53.t1.t2 (cl (= (=> @p_64 @p_63) @p_62)) :rule cong :premises (t53.t1.t0 t53.t1.t1))
(step t53.t1 (cl (= (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))) (forall ((?v2 S3)) (=> (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2)) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule bind)
(anchor :step t53.t2 :args ((?v2 S3) (:= (?v2 S3) ?v2)))
(step t53.t2.t0 (cl (= @p_62 @p_58)) :rule rare_rewrite :args ("bool-impl-elim" @p_57 @p_55))
(step t53.t2 (cl (= (forall ((?v2 S3)) (=> (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2)) (= f1 (f3 (f4 f14 ?v0) ?v2)))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule bind)
(step t53.t3 (cl (= (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule trans :premises (t53.t1 t53.t2))
(step t53.t4 (cl (= (=> @p_60 (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1)))) @p_59)) :rule cong :premises (t53.t0 t53.t3))
(step t53 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= (f3 (f4 f5 ?v0) ?v1) f1) (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))))) (forall ((?v0 S2) (?v1 S3)) (=> (= f1 (f3 (f4 f5 ?v0) ?v1)) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))))) :rule bind)
(anchor :step t54 :args ((?v0 S2) (?v1 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1)))
(step t54.t0 (cl (= @p_59 @p_52)) :rule rare_rewrite :args ("bool-impl-elim" @p_43 (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2))))))
(step t54 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= f1 (f3 (f4 f5 ?v0) ?v1)) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))))) :rule bind)
(anchor :step t55 :args ((?v0 S2) (?v1 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1)))
(step t55.t0 (cl (= (forall ((BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))) @p_53)) :rule hole)
(step t55.t1 (cl (= @p_44 @p_44)) :rule refl)
(anchor :step t55.t2 :args ((?v2 S3) (:= (BOUND_VARIABLE_8839 S3) ?v2)))
(step t55.t2.t0 (cl (= @p_50 @p_58)) :rule refl)
(step t55.t2 (cl (= (forall ((BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) :rule bind)
(step t55.t3 (cl (= @p_53 @p_52)) :rule cong :premises (t55.t1 t55.t2))
(step t55.t4 (cl (= (forall ((BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))) @p_52)) :rule trans :premises (t55.t0 t55.t3))
(step t55.t5 (cl (= @p_52 (forall ((BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))))) :rule symm :premises (t55.t4))
(step t55 (cl (= (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3)) (forall ((BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))))) :rule bind)
(step t56 (cl (= (forall ((?v0 S2) (?v1 S3)) (forall ((BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))))) :rule qnt_join)
(step t57 (cl (= (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))))) :rule trans :premises (t55 t56))
(anchor :step t58 :args ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1) (:= (BOUND_VARIABLE_8839 S3) BOUND_VARIABLE_8839)))
(step t58.t0 (cl @p_51) :rule hole :args ("ACI_NORM" @p_51))
(step t58 (cl (= (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) :rule bind)
(step t59 (cl (= (forall ((?v0 S2) (?v1 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) :rule trans :premises (t57 t58))
(step t60 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= f1 (f3 (f4 f5 ?v0) ?v1)) (forall ((?v2 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2))) (= f1 (f3 (f4 f14 ?v0) ?v2)))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) :rule trans :premises (t54 t59))
(step t61 (cl (= (forall ((?v0 S2) (?v1 S3)) (=> (= (f3 (f4 f5 ?v0) ?v1) f1) (forall ((?v2 S3)) (=> (= (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) ?v2) f1) (= (f3 (f4 f14 ?v0) ?v2) f1))))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) :rule trans :premises (t53 t60))
(step t62 (cl (= @p_1 @p_8)) :rule cong :premises (t52 t61))
(step t63 (cl (= @p_2 @p_9)) :rule cong :premises (t62))
(anchor :step t64 :args ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1) (:= (BOUND_VARIABLE_8816 S3) BOUND_VARIABLE_8816)))
(step t64.t0 (cl (= f1 @p_3)) :rule symm :premises (a140))
(step t64.t1 (cl (= @p_24 @p_24)) :rule refl)
(step t64.t2 (cl (= @p_43 @p_28)) :rule cong :premises (t64.t0 t64.t1))
(step t64.t3 (cl (= @p_44 @p_29)) :rule cong :premises (t64.t2))
(step t64.t4 (cl (= @p_36 @p_36)) :rule refl)
(step t64.t5 (cl (= @p_47 @p_37)) :rule cong :premises (t64.t0 t64.t4))
(step t64.t6 (cl (= @p_48 @p_38)) :rule cong :premises (t64.t5))
(step t64.t7 (cl (= @p_32 @p_32)) :rule refl)
(step t64.t8 (cl (= @p_46 @p_33)) :rule cong :premises (t64.t0 t64.t7))
(step t64.t9 (cl (= @p_49 @p_39)) :rule cong :premises (t64.t3 t64.t6 t64.t8))
(step t64 (cl (= (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= (f37 (f38 f39 f28) f57) (f3 (f4 f5 ?v0) ?v1))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) :rule bind)
(anchor :step t65 :args ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1) (:= (BOUND_VARIABLE_8839 S3) BOUND_VARIABLE_8839)))
(step t65.t0 (cl (= f1 @p_3)) :rule symm :premises (a140))
(step t65.t1 (cl (= @p_24 @p_24)) :rule refl)
(step t65.t2 (cl (= @p_43 @p_28)) :rule cong :premises (t65.t0 t65.t1))
(step t65.t3 (cl (= @p_44 @p_29)) :rule cong :premises (t65.t2))
(step t65.t4 (cl (= @p_21 @p_21)) :rule refl)
(step t65.t5 (cl (= @p_41 @p_22)) :rule cong :premises (t65.t0 t65.t4))
(step t65.t6 (cl (= @p_42 @p_23)) :rule cong :premises (t65.t5))
(step t65.t7 (cl (= @p_16 @p_16)) :rule refl)
(step t65.t8 (cl (= @p_40 @p_17)) :rule cong :premises (t65.t0 t65.t7))
(step t65.t9 (cl (= @p_45 @p_30)) :rule cong :premises (t65.t3 t65.t6 t65.t8))
(step t65 (cl (= (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= f1 (f3 (f4 f5 ?v0) ?v1))) (not (= f1 (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= f1 (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= (f37 (f38 f39 f28) f57) (f3 (f4 f5 ?v0) ?v1))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) :rule bind)
(step t66 (cl (= @p_8 @p_10)) :rule cong :premises (t64 t65))
(step t67 (cl (= @p_9 @p_11)) :rule cong :premises (t66))
(anchor :step t68 :args ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1) (:= (BOUND_VARIABLE_8816 S3) BOUND_VARIABLE_8816)))
(step t68.t0 (cl (= @p_28 @p_25)) :rule rare_rewrite :args ("eq-symm" @p_3 @p_24))
(step t68.t1 (cl (= @p_29 @p_26)) :rule cong :premises (t68.t0))
(step t68.t2 (cl (= @p_38 @p_38)) :rule refl)
(step t68.t3 (cl (= @p_33 @p_33)) :rule refl)
(step t68.t4 (cl (= @p_39 (or @p_26 @p_38 @p_33))) :rule cong :premises (t68.t1 t68.t2 t68.t3))
(step t68 (cl (= (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= (f37 (f38 f39 f28) f57) (f3 (f4 f5 ?v0) ?v1))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816)))))) :rule bind)
(anchor :step t69 :args ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3) (:= (?v0 S2) ?v0) (:= (?v1 S3) ?v1) (:= (BOUND_VARIABLE_8839 S3) BOUND_VARIABLE_8839)))
(step t69.t0 (cl (= @p_28 @p_25)) :rule rare_rewrite :args ("eq-symm" @p_3 @p_24))
(step t69.t1 (cl (= @p_29 @p_26)) :rule cong :premises (t69.t0))
(step t69.t2 (cl (= @p_23 @p_23)) :rule refl)
(step t69.t3 (cl (= @p_17 @p_17)) :rule refl)
(step t69.t4 (cl (= @p_30 @p_27)) :rule cong :premises (t69.t1 t69.t2 t69.t3))
(step t69 (cl (= (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= (f37 (f38 f39 f28) f57) (f3 (f4 f5 ?v0) ?v1))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))) (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8839 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) f13) BOUND_VARIABLE_8839))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8839)))))) :rule bind)
(step t70 (cl (= @p_10 @p_5)) :rule cong :premises (t68 t69))
(step t71 (cl (= @p_11 @p_6)) :rule cong :premises (t70))
(step t72 (cl (= @p_9 @p_6)) :rule trans :premises (t67 t71))
(step t73 (cl @p_7) :rule trans :premises (t63 t72))
(step t74 (cl @p_6) :rule resolution :premises (t43 t73 a1))
(step t75 (cl @p_120) :rule not_implies2 :premises (t74))
(step t76 (cl @p_119) :rule resolution :premises (t42 t75))
(step t77 (cl @p_142) :rule resolution :premises (t18 t76))
(step t78 (cl (not (! (= (! (or @p_118 (! (not @p_82) :named @p_139)) :named @p_140) (! (or @p_118 @p_81) :named @p_138)) :named @p_141)) (not @p_140) @p_138) :rule equiv_pos2)
(step t79 (cl (= @p_118 @p_118)) :rule refl)
(step t80 (cl (= @p_139 @p_81)) :rule rare_rewrite :args ("bool-double-not-elim" @p_81))
(step t81 (cl @p_141) :rule cong :premises (t79 t80))
(step t82 (cl @p_118 @p_139) :rule or_neg :args (0))
(step t83 (cl @p_140 @p_119) :rule or_neg :args (0))
(step t84 (cl @p_140 (not @p_139)) :rule or_neg :args (1))
(step t85 (cl @p_140 @p_140) :rule resolution :premises (t82 t83 t84))
(step t86 (cl @p_140) :rule contraction :premises (t85))
(step t87 (cl @p_138) :rule resolution :premises (t78 t81 t86))
(step t88 (cl @p_118 @p_81) :rule or :premises (t87))
(step t89 (cl @p_81 @p_118) :rule reordering :premises (t88))
(step t90 (cl @p_81) :rule resolution :premises (t89 t76))
(step t91 (cl (! (not (! (or @p_87 @p_78) :named @p_88)) :named @p_137) @p_87 @p_78) :rule or_pos)
(step t92 (cl @p_87 @p_78 @p_137) :rule reordering :premises (t91))
(step t93 (cl (not (! (= (! (or @p_118 (! (not @p_87) :named @p_134)) :named @p_135) (! (or @p_118 @p_86) :named @p_133)) :named @p_136)) (not @p_135) @p_133) :rule equiv_pos2)
(step t94 (cl (= @p_134 @p_86)) :rule rare_rewrite :args ("bool-double-not-elim" @p_86))
(step t95 (cl @p_136) :rule cong :premises (t79 t94))
(step t96 (cl @p_118 @p_134) :rule or_neg :args (1))
(step t97 (cl @p_135 @p_119) :rule or_neg :args (0))
(step t98 (cl @p_135 (not @p_134)) :rule or_neg :args (1))
(step t99 (cl @p_135 @p_135) :rule resolution :premises (t96 t97 t98))
(step t100 (cl @p_135) :rule contraction :premises (t99))
(step t101 (cl @p_133) :rule resolution :premises (t93 t95 t100))
(step t102 (cl @p_118 @p_86) :rule or :premises (t101))
(step t103 (cl @p_86 @p_118) :rule reordering :premises (t102))
(step t104 (cl @p_86) :rule resolution :premises (t103 t76))
(step t105 (cl (not (! (= (! (=> (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) (f37 (f38 f39 f28) f57))) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) (f37 (f38 f39 f28) f57)))) (! (or (! (not (! (= @p_85 @p_3) :named @p_113)) :named @p_114) (! (= @p_77 @p_3) :named @p_112)) :named @p_115)) :named @p_116) (! (=> (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) (f37 (f38 f39 f28) f57))) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) (f37 (f38 f39 f28) f57)))) @p_88) :named @p_111)) :named @p_117)) (not @p_116) @p_111) :rule equiv_pos2)
(step t106 (cl (= (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) (f37 (f38 f39 f28) f57))) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) (f37 (f38 f39 f28) f57)))) (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) (f37 (f38 f39 f28) f57))) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) (f37 (f38 f39 f28) f57)))))) :rule refl)
(step t107 (cl (= @p_113 @p_86)) :rule rare_rewrite :args ("eq-symm" @p_85 @p_3))
(step t108 (cl (= @p_114 @p_87)) :rule cong :premises (t107))
(step t109 (cl (= @p_112 @p_78)) :rule rare_rewrite :args ("eq-symm" @p_77 @p_3))
(step t110 (cl (= @p_115 @p_88)) :rule cong :premises (t108 t109))
(step t111 (cl @p_117) :rule cong :premises (t106 t110))
(step t112 (cl @p_116 (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) (f37 (f38 f39 f28) f57))) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) (f37 (f38 f39 f28) f57))))) :rule implies_neg1)
(anchor :step t113)
(assume t113.a0 (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) (f37 (f38 f39 f28) f57))) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) (f37 (f38 f39 f28) f57)))))
(step t113.t0 (cl (or (! (not (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) (f37 (f38 f39 f28) f57))) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) (f37 (f38 f39 f28) f57))))) :named @p_110) @p_115)) :rule forall_inst :args (f10 @quantifiers_skolemize_2 f13 @quantifiers_skolemize_3))
(step t113.t1 (cl @p_110 @p_115) :rule or :premises (t113.t0))
(step t113.t2 (cl (not (! (= (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (=> (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) f1) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) f1))) (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) (f37 (f38 f39 f28) f57))) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) (f37 (f38 f39 f28) f57))))) :named @p_89)) (not (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (=> (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) f1) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) f1)))) (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) (f37 (f38 f39 f28) f57))) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) (f37 (f38 f39 f28) f57))))) :rule equiv_pos2)
(anchor :step t113.t3 :args ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3) (:= (?v0 S10) ?v0) (:= (?v1 S3) ?v1) (:= (?v2 S7) ?v2) (:= (?v3 S3) ?v3)))
(step t113.t3.t0 (cl (= (! (= (! (f3 (f6 (! (f7 (f8 f9 ?v0) ?v1) :named @p_93) ?v2) ?v3) :named @p_96) f1) :named @p_109) (! (= f1 @p_96) :named @p_104))) :rule rare_rewrite :args ("eq-symm" @p_96 f1))
(step t113.t3.t1 (cl (= (! (= (! (f3 (f6 @p_93 (f11 f12 ?v2)) ?v3) :named @p_94) f1) :named @p_108) (! (= f1 @p_94) :named @p_103))) :rule rare_rewrite :args ("eq-symm" @p_94 f1))
(step t113.t3.t2 (cl (= (=> @p_109 @p_108) (! (=> @p_104 @p_103) :named @p_107))) :rule cong :premises (t113.t3.t0 t113.t3.t1))
(step t113.t3 (cl (= (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (=> (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) f1) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) f1))) (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (=> (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3)) (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3)))))) :rule bind)
(anchor :step t113.t4 :args ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3) (:= (?v0 S10) ?v0) (:= (?v1 S3) ?v1) (:= (?v2 S7) ?v2) (:= (?v3 S3) ?v3)))
(step t113.t4.t0 (cl (= @p_107 (! (or (! (not @p_104) :named @p_105) @p_103) :named @p_106))) :rule rare_rewrite :args ("bool-impl-elim" @p_104 @p_103))
(step t113.t4 (cl (= (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (=> (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3)) (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3)))) (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3))) (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3)))))) :rule bind)
(step t113.t5 (cl (= (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (=> (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) f1) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) f1))) (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3))) (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3)))))) :rule trans :premises (t113.t3 t113.t4))
(anchor :step t113.t6 :args ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3) (:= (?v0 S10) ?v0) (:= (?v1 S3) ?v1) (:= (?v2 S7) ?v2) (:= (?v3 S3) ?v3)))
(step t113.t6.t0 (cl (= f1 @p_3)) :rule symm :premises (a140))
(step t113.t6.t1 (cl (= @p_96 @p_96)) :rule refl)
(step t113.t6.t2 (cl (= @p_104 (! (= @p_3 @p_96) :named @p_100))) :rule cong :premises (t113.t6.t0 t113.t6.t1))
(step t113.t6.t3 (cl (= @p_105 (! (not @p_100) :named @p_101))) :rule cong :premises (t113.t6.t2))
(step t113.t6.t4 (cl (= @p_94 @p_94)) :rule refl)
(step t113.t6.t5 (cl (= @p_103 (! (= @p_3 @p_94) :named @p_99))) :rule cong :premises (t113.t6.t0 t113.t6.t4))
(step t113.t6.t6 (cl (= @p_106 (! (or @p_101 @p_99) :named @p_102))) :rule cong :premises (t113.t6.t3 t113.t6.t5))
(step t113.t6 (cl (= (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3))) (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3)))) (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3))) (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3)))))) :rule bind)
(anchor :step t113.t7 :args ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3) (:= (?v0 S10) ?v0) (:= (?v1 S3) ?v1) (:= (?v2 S7) ?v2) (:= (?v3 S3) ?v3)))
(step t113.t7.t0 (cl (= @p_100 (! (= @p_96 @p_3) :named @p_97))) :rule rare_rewrite :args ("eq-symm" @p_3 @p_96))
(step t113.t7.t1 (cl (= @p_101 (! (not @p_97) :named @p_98))) :rule cong :premises (t113.t7.t0))
(step t113.t7.t2 (cl (= @p_99 (! (= @p_94 @p_3) :named @p_95))) :rule rare_rewrite :args ("eq-symm" @p_3 @p_94))
(step t113.t7.t3 (cl (= @p_102 (or @p_98 @p_95))) :rule cong :premises (t113.t7.t1 t113.t7.t2))
(step t113.t7 (cl (= (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3))) (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3)))) (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) (f37 (f38 f39 f28) f57))) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) (f37 (f38 f39 f28) f57)))))) :rule bind)
(step t113.t8 (cl (= (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3))) (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3)))) (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) (f37 (f38 f39 f28) f57))) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) (f37 (f38 f39 f28) f57)))))) :rule trans :premises (t113.t6 t113.t7))
(step t113.t9 (cl @p_89) :rule trans :premises (t113.t5 t113.t8))
(step t113.t10 (cl (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) (f37 (f38 f39 f28) f57))) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) (f37 (f38 f39 f28) f57))))) :rule resolution :premises (t113.t2 t113.t9 a2))
(step t113.t11 (cl @p_115) :rule resolution :premises (t113.t1 t113.t10))
(step t113 (cl @p_110 @p_115) :rule subproof :discharge (t113.a0))
(step t114 (cl @p_116 @p_115) :rule resolution :premises (t112 t113))
(step t115 (cl @p_116 (not @p_115)) :rule implies_neg2)
(step t116 (cl @p_116 @p_116) :rule resolution :premises (t114 t115))
(step t117 (cl @p_116) :rule contraction :premises (t116))
(step t118 (cl @p_111) :rule resolution :premises (t105 t111 t117))
(step t119 (cl @p_110 @p_88) :rule implies :premises (t118))
(step t120 (cl (not @p_89) (not (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (=> (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) f1) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) f1)))) (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) (f37 (f38 f39 f28) f57))) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) (f37 (f38 f39 f28) f57))))) :rule equiv_pos2)
(anchor :step t121 :args ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3) (:= (?v0 S10) ?v0) (:= (?v1 S3) ?v1) (:= (?v2 S7) ?v2) (:= (?v3 S3) ?v3)))
(step t121.t0 (cl (= @p_109 @p_104)) :rule rare_rewrite :args ("eq-symm" @p_96 f1))
(step t121.t1 (cl (= @p_108 @p_103)) :rule rare_rewrite :args ("eq-symm" @p_94 f1))
(step t121.t2 (cl (= (=> @p_109 @p_108) @p_107)) :rule cong :premises (t121.t0 t121.t1))
(step t121 (cl (= (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (=> (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) f1) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) f1))) (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (=> (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3)) (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3)))))) :rule bind)
(anchor :step t122 :args ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3) (:= (?v0 S10) ?v0) (:= (?v1 S3) ?v1) (:= (?v2 S7) ?v2) (:= (?v3 S3) ?v3)))
(step t122.t0 (cl (= @p_107 @p_106)) :rule rare_rewrite :args ("bool-impl-elim" @p_104 @p_103))
(step t122 (cl (= (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (=> (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3)) (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3)))) (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3))) (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3)))))) :rule bind)
(step t123 (cl (= (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (=> (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) f1) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) f1))) (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3))) (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3)))))) :rule trans :premises (t121 t122))
(anchor :step t124 :args ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3) (:= (?v0 S10) ?v0) (:= (?v1 S3) ?v1) (:= (?v2 S7) ?v2) (:= (?v3 S3) ?v3)))
(step t124.t0 (cl (= f1 @p_3)) :rule symm :premises (a140))
(step t124.t1 (cl (= @p_96 @p_96)) :rule refl)
(step t124.t2 (cl (= @p_104 @p_100)) :rule cong :premises (t124.t0 t124.t1))
(step t124.t3 (cl (= @p_105 @p_101)) :rule cong :premises (t124.t2))
(step t124.t4 (cl (= @p_94 @p_94)) :rule refl)
(step t124.t5 (cl (= @p_103 @p_99)) :rule cong :premises (t124.t0 t124.t4))
(step t124.t6 (cl (= @p_106 @p_102)) :rule cong :premises (t124.t3 t124.t5))
(step t124 (cl (= (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3))) (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3)))) (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3))) (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3)))))) :rule bind)
(anchor :step t125 :args ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3) (:= (?v0 S10) ?v0) (:= (?v1 S3) ?v1) (:= (?v2 S7) ?v2) (:= (?v3 S3) ?v3)))
(step t125.t0 (cl (= @p_100 @p_97)) :rule rare_rewrite :args ("eq-symm" @p_3 @p_96))
(step t125.t1 (cl (= @p_101 @p_98)) :rule cong :premises (t125.t0))
(step t125.t2 (cl (= @p_99 @p_95)) :rule rare_rewrite :args ("eq-symm" @p_3 @p_94))
(step t125.t3 (cl (= @p_102 (or @p_98 @p_95))) :rule cong :premises (t125.t1 t125.t2))
(step t125 (cl (= (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3))) (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3)))) (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) (f37 (f38 f39 f28) f57))) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) (f37 (f38 f39 f28) f57)))))) :rule bind)
(step t126 (cl (= (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3))) (= f1 (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3)))) (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) (f37 (f38 f39 f28) f57))) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) (f37 (f38 f39 f28) f57)))))) :rule trans :premises (t124 t125))
(step t127 (cl @p_89) :rule trans :premises (t123 t126))
(step t128 (cl (forall ((?v0 S10) (?v1 S3) (?v2 S7) (?v3 S3)) (or (not (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) ?v2) ?v3) (f37 (f38 f39 f28) f57))) (= (f3 (f6 (f7 (f8 f9 ?v0) ?v1) (f11 f12 ?v2)) ?v3) (f37 (f38 f39 f28) f57))))) :rule resolution :premises (t120 t127 a2))
(step t129 (cl @p_88) :rule resolution :premises (t119 t128))
(step t130 (cl @p_78) :rule resolution :premises (t92 t104 t129))
(step t131 (cl @p_84) :rule resolution :premises (t17 t77 t90 t130))
(step t132 (cl (forall ((?v0 S2) (?v1 S3) (BOUND_VARIABLE_8816 S3)) (or (not (= (f3 (f4 f5 ?v0) ?v1) (f37 (f38 f39 f28) f57))) (not (= (f37 (f38 f39 f28) f57) (f3 (f6 (f7 (f8 f9 f10) ?v1) (f11 f12 f13)) BOUND_VARIABLE_8816))) (= (f37 (f38 f39 f28) f57) (f3 (f4 f14 ?v0) BOUND_VARIABLE_8816))))) :rule not_implies1 :premises (t74))
(step t133 (cl) :rule resolution :premises (t15 t131 t132))
