(assume a0 (and (and (and (and (and (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255)) (= (select_format 5) percent)) (= (select_format 5) s)) (= (select_format 5) x)) (= (select_format 5) adr_lo)) (= (select_format 5) adr_medlo)) (= (select_format 5) adr_medhi)) (= (select_format 5) adr_hi)) (= (select_format 5) 255))) (and (and (and (and (and (and (and (and (and (and (and (ite (and (= (select_format 0) percent) (= (select_format 1) s)) (= (s_count 0) 1) (= (s_count 0) 0)) (ite (and (= (select_format 1) percent) (= (select_format 2) s)) (= (s_count 1) (+ (s_count 0) 1)) (= (s_count 1) (s_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) s)) (= (s_count 2) (+ (s_count 1) 1)) (= (s_count 2) (s_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) s)) (= (s_count 3) (+ (s_count 2) 1)) (= (s_count 3) (s_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) s)) (= (s_count 4) (+ (s_count 3) 1)) (= (s_count 4) (s_count 3)))) (ite (and (= (select_format 5) percent) (= (select_format 6) s)) (= (s_count 5) (+ (s_count 4) 1)) (= (s_count 5) (s_count 4)))) (ite (and (= (select_format 0) percent) (= (select_format 1) x)) (= (x_count 0) 1) (= (x_count 0) 0))) (ite (and (= (select_format 1) percent) (= (select_format 2) x)) (= (x_count 1) (+ (x_count 0) 1)) (= (x_count 1) (x_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) x)) (= (x_count 2) (+ (x_count 1) 1)) (= (x_count 2) (x_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) x)) (= (x_count 3) (+ (x_count 2) 1)) (= (x_count 3) (x_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) x)) (= (x_count 4) (+ (x_count 3) 1)) (= (x_count 4) (x_count 3)))) (ite (and (= (select_format 5) percent) (= (select_format 6) x)) (= (x_count 5) (+ (x_count 4) 1)) (= (x_count 5) (x_count 4))))) (and (or (or (or (or (or (= fmt1 (+ fmt0 0)) (= fmt1 (+ fmt0 1))) (= fmt1 (+ fmt0 2))) (= fmt1 (+ fmt0 3))) (= fmt1 (+ fmt0 4))) (= fmt1 (+ fmt0 5))) (or (or (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1))) (= arg1 (+ fmt0 2))))) (not (and (and (and (and (and (= (select_format fmt1) percent) (= (select_format (+ fmt1 1)) s)) (= (select_format arg1) adr_lo)) (= (select_format (+ arg1 1)) adr_medlo)) (= (select_format (+ arg1 2)) adr_medhi)) (= (select_format (+ arg1 3)) adr_hi)))))
(step t0 (cl (=> (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1))) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1)))) :rule implies_neg1)
(anchor :step t1)
(assume t1.a0 (= distance 14))
(assume t1.a1 (= fmt0 0))
(assume t1.a2 (= fmt0 (+ arg0 distance)))
(assume t1.a3 (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))
(assume t1.a4 (= fmt0 (+ -1 arg1)))
(step t1.t0 (cl (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1))) (not (= distance 14)) (not (= (+ arg0 distance) 0)) (not (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (not (= 0 (+ -1 arg1)))) :rule and_neg)
(step t1.t1 (cl (and (and (and (and (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255)) (= (select_format 5) percent)) (= (select_format 5) s)) (= (select_format 5) x)) (= (select_format 5) adr_lo)) (= (select_format 5) adr_medlo)) (= (select_format 5) adr_medhi)) (= (select_format 5) adr_hi)) (= (select_format 5) 255))) (and (and (and (and (and (and (and (and (and (and (and (ite (and (= (select_format 0) percent) (= (select_format 1) s)) (= (s_count 0) 1) (= (s_count 0) 0)) (ite (and (= (select_format 1) percent) (= (select_format 2) s)) (= (s_count 1) (+ (s_count 0) 1)) (= (s_count 1) (s_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) s)) (= (s_count 2) (+ (s_count 1) 1)) (= (s_count 2) (s_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) s)) (= (s_count 3) (+ (s_count 2) 1)) (= (s_count 3) (s_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) s)) (= (s_count 4) (+ (s_count 3) 1)) (= (s_count 4) (s_count 3)))) (ite (and (= (select_format 5) percent) (= (select_format 6) s)) (= (s_count 5) (+ (s_count 4) 1)) (= (s_count 5) (s_count 4)))) (ite (and (= (select_format 0) percent) (= (select_format 1) x)) (= (x_count 0) 1) (= (x_count 0) 0))) (ite (and (= (select_format 1) percent) (= (select_format 2) x)) (= (x_count 1) (+ (x_count 0) 1)) (= (x_count 1) (x_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) x)) (= (x_count 2) (+ (x_count 1) 1)) (= (x_count 2) (x_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) x)) (= (x_count 3) (+ (x_count 2) 1)) (= (x_count 3) (x_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) x)) (= (x_count 4) (+ (x_count 3) 1)) (= (x_count 4) (x_count 3)))) (ite (and (= (select_format 5) percent) (= (select_format 6) x)) (= (x_count 5) (+ (x_count 4) 1)) (= (x_count 5) (x_count 4))))) (and (or (or (or (or (or (= fmt1 (+ fmt0 0)) (= fmt1 (+ fmt0 1))) (= fmt1 (+ fmt0 2))) (= fmt1 (+ fmt0 3))) (= fmt1 (+ fmt0 4))) (= fmt1 (+ fmt0 5))) (or (or (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1))) (= arg1 (+ fmt0 2)))))) :rule and :premises (a0) :args (0))
(step t1.t2 (cl (and (and (and (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255)) (= (select_format 5) percent)) (= (select_format 5) s)) (= (select_format 5) x)) (= (select_format 5) adr_lo)) (= (select_format 5) adr_medlo)) (= (select_format 5) adr_medhi)) (= (select_format 5) adr_hi)) (= (select_format 5) 255))) (and (and (and (and (and (and (and (and (and (and (and (ite (and (= (select_format 0) percent) (= (select_format 1) s)) (= (s_count 0) 1) (= (s_count 0) 0)) (ite (and (= (select_format 1) percent) (= (select_format 2) s)) (= (s_count 1) (+ (s_count 0) 1)) (= (s_count 1) (s_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) s)) (= (s_count 2) (+ (s_count 1) 1)) (= (s_count 2) (s_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) s)) (= (s_count 3) (+ (s_count 2) 1)) (= (s_count 3) (s_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) s)) (= (s_count 4) (+ (s_count 3) 1)) (= (s_count 4) (s_count 3)))) (ite (and (= (select_format 5) percent) (= (select_format 6) s)) (= (s_count 5) (+ (s_count 4) 1)) (= (s_count 5) (s_count 4)))) (ite (and (= (select_format 0) percent) (= (select_format 1) x)) (= (x_count 0) 1) (= (x_count 0) 0))) (ite (and (= (select_format 1) percent) (= (select_format 2) x)) (= (x_count 1) (+ (x_count 0) 1)) (= (x_count 1) (x_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) x)) (= (x_count 2) (+ (x_count 1) 1)) (= (x_count 2) (x_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) x)) (= (x_count 3) (+ (x_count 2) 1)) (= (x_count 3) (x_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) x)) (= (x_count 4) (+ (x_count 3) 1)) (= (x_count 4) (x_count 3)))) (ite (and (= (select_format 5) percent) (= (select_format 6) x)) (= (x_count 5) (+ (x_count 4) 1)) (= (x_count 5) (x_count 4)))))) :rule and :premises (t1.t1) :args (0))
(step t1.t3 (cl (and (and (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255)) (= (select_format 5) percent)) (= (select_format 5) s)) (= (select_format 5) x)) (= (select_format 5) adr_lo)) (= (select_format 5) adr_medlo)) (= (select_format 5) adr_medhi)) (= (select_format 5) adr_hi)) (= (select_format 5) 255)))) :rule and :premises (t1.t2) :args (0))
(step t1.t4 (cl (and (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))))) :rule and :premises (t1.t3) :args (0))
(step t1.t5 (cl (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120))) :rule and :premises (t1.t4) :args (0))
(step t1.t6 (cl (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115))) :rule and :premises (t1.t5) :args (0))
(step t1.t7 (cl (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37))) :rule and :premises (t1.t6) :args (0))
(step t1.t8 (cl (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2))) :rule and :premises (t1.t7) :args (0))
(step t1.t9 (cl (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5))) :rule and :premises (t1.t8) :args (0))
(step t1.t10 (cl (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1))) :rule and :premises (t1.t9) :args (0))
(step t1.t11 (cl (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3))) :rule and :premises (t1.t10) :args (0))
(step t1.t12 (cl (and (= distance 14) (= fmt_length 6))) :rule and :premises (t1.t11) :args (0))
(step t1.t13 (cl (= distance 14)) :rule and :premises (t1.t12) :args (0))
(step t1.t14 (cl (not (= (= arg0 (- fmt0 distance)) (= fmt0 (+ arg0 distance)))) (not (= arg0 (- fmt0 distance))) (= fmt0 (+ arg0 distance))) :rule equiv_pos2)
(step t1.t15 (cl (= arg0 arg0)) :rule refl)
(step t1.t16 (cl (= (- fmt0 distance) (+ fmt0 (* -1 distance)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- fmt0 distance) (+ fmt0 (* -1 distance))) 3 6))
(step t1.t17 (cl (= (= arg0 (- fmt0 distance)) (= arg0 (+ fmt0 (* -1 distance))))) :rule cong :premises (t1.t15 t1.t16))
(step t1.t18 (cl (= (= arg0 (+ fmt0 (* -1 distance))) (= fmt0 (+ arg0 distance)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= arg0 (+ fmt0 (* -1 distance))) (= fmt0 (+ arg0 distance))) 3 7))
(step t1.t19 (cl (= (= arg0 (- fmt0 distance)) (= fmt0 (+ arg0 distance)))) :rule trans :premises (t1.t17 t1.t18))
(step t1.t20 (cl (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) :rule and :premises (t1.t4) :args (1))
(step t1.t21 (cl (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)))) :rule and :premises (t1.t20) :args (0))
(step t1.t22 (cl (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance)))) :rule and :premises (t1.t21) :args (0))
(step t1.t23 (cl (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1)))) :rule and :premises (t1.t22) :args (0))
(step t1.t24 (cl (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1)))) :rule and :premises (t1.t23) :args (0))
(step t1.t25 (cl (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0))) :rule and :premises (t1.t24) :args (0))
(step t1.t26 (cl (and (= fmt0 0) (= arg0 (- fmt0 distance)))) :rule and :premises (t1.t25) :args (0))
(step t1.t27 (cl (= arg0 (- fmt0 distance))) :rule and :premises (t1.t26) :args (1))
(step t1.t28 (cl (= fmt0 (+ arg0 distance))) :rule resolution :premises (t1.t14 t1.t19 t1.t27))
(step t1.t29 (cl (= (+ arg0 distance) fmt0)) :rule symm :premises (t1.t28))
(step t1.t30 (cl (= fmt0 0)) :rule and :premises (t1.t26) :args (0))
(step t1.t31 (cl (= (+ arg0 distance) 0)) :rule trans :premises (t1.t29 t1.t30))
(step t1.t32 (cl (not (= (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) (not (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule equiv_pos2)
(step t1.t33 (cl (= arg1 arg1)) :rule refl)
(step t1.t34 (cl (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0)) 3 6))
(step t1.t35 (cl (= 4 4)) :rule refl)
(step t1.t36 (cl (= (- (- fmt1 2) fmt0) (+ (- fmt1 2) (* -1 fmt0)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- (- fmt1 2) fmt0) (+ (- fmt1 2) (* -1 fmt0))) 3 6))
(step t1.t37 (cl (= (- fmt1 2) (+ fmt1 (* -1 2)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- fmt1 2) (+ fmt1 (* -1 2))) 3 6))
(step t1.t38 (cl (= fmt1 fmt1)) :rule refl)
(step t1.t39 (cl (= (* -1 2) -2)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (* -1 2) -2) 3 7))
(step t1.t40 (cl (= (+ fmt1 (* -1 2)) (+ fmt1 -2))) :rule cong :premises (t1.t38 t1.t39))
(step t1.t41 (cl (= (+ fmt1 -2) (+ -2 fmt1))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ fmt1 -2) (+ -2 fmt1)) 3 7))
(step t1.t42 (cl (= (+ fmt1 (* -1 2)) (+ -2 fmt1))) :rule trans :premises (t1.t40 t1.t41))
(step t1.t43 (cl (= (- fmt1 2) (+ -2 fmt1))) :rule trans :premises (t1.t37 t1.t42))
(step t1.t44 (cl (= (* -1 fmt0) (* -1 fmt0))) :rule refl)
(step t1.t45 (cl (= (+ (- fmt1 2) (* -1 fmt0)) (+ (+ -2 fmt1) (* -1 fmt0)))) :rule cong :premises (t1.t43 t1.t44))
(step t1.t46 (cl (= (+ (+ -2 fmt1) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ -2 fmt1) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0))) 3 7))
(step t1.t47 (cl (= (+ (- fmt1 2) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0)))) :rule trans :premises (t1.t45 t1.t46))
(step t1.t48 (cl (= (- (- fmt1 2) fmt0) (+ -2 fmt1 (* -1 fmt0)))) :rule trans :premises (t1.t36 t1.t47))
(step t1.t49 (cl (= (x_count (- (- fmt1 2) fmt0)) (x_count (+ -2 fmt1 (* -1 fmt0))))) :rule cong :premises (t1.t48))
(step t1.t50 (cl (= (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))))) :rule cong :premises (t1.t35 t1.t49))
(step t1.t51 (cl (= (s_count (- (- fmt1 2) fmt0)) (s_count (+ -2 fmt1 (* -1 fmt0))))) :rule cong :premises (t1.t48))
(step t1.t52 (cl (= (* 4 (s_count (- (- fmt1 2) fmt0))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) :rule cong :premises (t1.t35 t1.t51))
(step t1.t53 (cl (= (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0) (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0))) :rule cong :premises (t1.t50 t1.t52 t1.t15))
(step t1.t54 (cl (= (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 3 7))
(step t1.t55 (cl (= (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule trans :premises (t1.t53 t1.t54))
(step t1.t56 (cl (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule trans :premises (t1.t34 t1.t55))
(step t1.t57 (cl (= (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule cong :premises (t1.t33 t1.t56))
(step t1.t58 (cl (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))) :rule and :premises (t1.t20) :args (1))
(step t1.t59 (cl (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule resolution :premises (t1.t32 t1.t57 t1.t58))
(step t1.t60 (cl (=> (and (= fmt0 (+ -1 arg1)) (= fmt0 0)) (= 0 (+ -1 arg1))) (and (= fmt0 (+ -1 arg1)) (= fmt0 0))) :rule implies_neg1)
(anchor :step t1.t61)
(assume t1.t61.a0 (= fmt0 (+ -1 arg1)))
(assume t1.t61.a1 (= fmt0 0))
(step t1.t61.t0 (cl (= 0 fmt0)) :rule symm :premises (t1.t30))
(step t1.t61.t1 (cl (= (+ -1 arg1) fmt0)) :rule symm :premises (t1.t61.a0))
(step t1.t61.t2 (cl (= fmt0 (+ -1 arg1))) :rule symm :premises (t1.t61.t1))
(step t1.t61.t3 (cl (= 0 (+ -1 arg1))) :rule trans :premises (t1.t61.t0 t1.t61.t2))
(step t1.t61 (cl (not (= fmt0 (+ -1 arg1))) (not (= fmt0 0)) (= 0 (+ -1 arg1))) :rule subproof :discharge (t1.t61.a0 t1.t61.a1))
(step t1.t62 (cl (not (and (= fmt0 (+ -1 arg1)) (= fmt0 0))) (= fmt0 (+ -1 arg1))) :rule and_pos :args (0))
(step t1.t63 (cl (not (and (= fmt0 (+ -1 arg1)) (= fmt0 0))) (= fmt0 0)) :rule and_pos :args (1))
(step t1.t64 (cl (= 0 (+ -1 arg1)) (not (and (= fmt0 (+ -1 arg1)) (= fmt0 0))) (not (and (= fmt0 (+ -1 arg1)) (= fmt0 0)))) :rule resolution :premises (t1.t61 t1.t62 t1.t63))
(step t1.t65 (cl (not (and (= fmt0 (+ -1 arg1)) (= fmt0 0))) (not (and (= fmt0 (+ -1 arg1)) (= fmt0 0))) (= 0 (+ -1 arg1))) :rule reordering :premises (t1.t64))
(step t1.t66 (cl (not (and (= fmt0 (+ -1 arg1)) (= fmt0 0))) (= 0 (+ -1 arg1))) :rule contraction :premises (t1.t65))
(step t1.t67 (cl (=> (and (= fmt0 (+ -1 arg1)) (= fmt0 0)) (= 0 (+ -1 arg1))) (= 0 (+ -1 arg1))) :rule resolution :premises (t1.t60 t1.t66))
(step t1.t68 (cl (=> (and (= fmt0 (+ -1 arg1)) (= fmt0 0)) (= 0 (+ -1 arg1))) (not (= 0 (+ -1 arg1)))) :rule implies_neg2)
(step t1.t69 (cl (=> (and (= fmt0 (+ -1 arg1)) (= fmt0 0)) (= 0 (+ -1 arg1))) (=> (and (= fmt0 (+ -1 arg1)) (= fmt0 0)) (= 0 (+ -1 arg1)))) :rule resolution :premises (t1.t67 t1.t68))
(step t1.t70 (cl (=> (and (= fmt0 (+ -1 arg1)) (= fmt0 0)) (= 0 (+ -1 arg1)))) :rule contraction :premises (t1.t69))
(step t1.t71 (cl (not (and (= fmt0 (+ -1 arg1)) (= fmt0 0))) (= 0 (+ -1 arg1))) :rule implies :premises (t1.t70))
(step t1.t72 (cl (and (= fmt0 (+ -1 arg1)) (= fmt0 0)) (not (= fmt0 (+ -1 arg1))) (not (= fmt0 0))) :rule and_neg)
(step t1.t73 (cl (and (= fmt0 (+ -1 arg1)) (= fmt0 0))) :rule resolution :premises (t1.t72 t1.a4 t1.t30))
(step t1.t74 (cl (= 0 (+ -1 arg1))) :rule resolution :premises (t1.t71 t1.t73))
(step t1.t75 (cl (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) :rule resolution :premises (t1.t0 t1.t13 t1.t31 t1.t59 t1.t74))
(step t1 (cl (not (= distance 14)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (not (= fmt0 (+ -1 arg1))) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) :rule subproof :discharge (t1.a0 t1.a1 t1.a2 t1.a3 t1.a4))
(step t2 (cl (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1)))) (= distance 14)) :rule and_pos :args (0))
(step t3 (cl (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1)))) (= fmt0 0)) :rule and_pos :args (1))
(step t4 (cl (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1)))) (= fmt0 (+ arg0 distance))) :rule and_pos :args (2))
(step t5 (cl (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1)))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule and_pos :args (3))
(step t6 (cl (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1)))) (= fmt0 (+ -1 arg1))) :rule and_pos :args (4))
(step t7 (cl (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1))) (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1)))) (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1)))) (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1)))) (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1)))) (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1))))) :rule resolution :premises (t1 t2 t3 t4 t5 t6))
(step t8 (cl (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1)))) (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1)))) (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1)))) (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1)))) (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1)))) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) :rule reordering :premises (t7))
(step t9 (cl (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1)))) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) :rule contraction :premises (t8))
(step t10 (cl (=> (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1))) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) :rule resolution :premises (t0 t9))
(step t11 (cl (=> (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1))) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1))))) :rule implies_neg2)
(step t12 (cl (=> (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1))) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) (=> (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1))) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1))))) :rule resolution :premises (t10 t11))
(step t13 (cl (=> (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1))) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1))))) :rule contraction :premises (t12))
(step t14 (cl (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1)))) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) :rule implies :premises (t13))
(step t15 (cl (=> (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1))) false) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) :rule implies_neg1)
(anchor :step t16)
(assume t16.a0 (= distance 14))
(assume t16.a1 (= (+ arg0 distance) 0))
(assume t16.a2 (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))
(assume t16.a3 (= 0 (+ -1 arg1)))
(step t16.t0 (cl (not (= (= 0 (+ -1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) false)) (not (= 0 (+ -1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0))))))))) false) :rule equiv_pos2)
(step t16.t1 (cl (= 0 0)) :rule refl)
(step t16.t2 (cl (= -1 -1)) :rule refl)
(step t16.t3 (cl (= (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (+ -14 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (+ -14 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 3 7))
(step t16.t4 (cl (= (+ -1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (+ -1 (+ -14 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule cong :premises (t16.t2 t16.t3))
(step t16.t5 (cl (= (+ -1 (+ -14 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (+ -15 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ -1 (+ -14 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (+ -15 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 3 7))
(step t16.t6 (cl (= (+ -1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (+ -15 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule trans :premises (t16.t4 t16.t5))
(step t16.t7 (cl (= (= 0 (+ -1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) (= 0 (+ -15 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule cong :premises (t16.t1 t16.t6))
(step t16.t8 (cl (= (= 0 (+ -15 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) false)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= 0 (+ -15 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) false) 3 7))
(step t16.t9 (cl (= (= 0 (+ -1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) false)) :rule trans :premises (t16.t7 t16.t8))
(step t16.t10 (cl (not (= (= 0 (+ -1 arg1)) (= 0 (+ -1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))))) (not (= 0 (+ -1 arg1))) (= 0 (+ -1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0))))))))) :rule equiv_pos2)
(step t16.t11 (cl (not (= (= arg1 (+ (* -1 14) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0))))))))) (not (= arg1 (+ (* -1 14) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (= arg1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule equiv_pos2)
(step t16.t12 (cl (= arg1 arg1)) :rule refl)
(step t16.t13 (cl (= (* -1 14) -14)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (* -1 14) -14) 3 7))
(step t16.t14 (cl (= (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))))) :rule refl)
(step t16.t15 (cl (= (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) :rule refl)
(step t16.t16 (cl (= (+ (* -1 14) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))) (+ -14 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule cong :premises (t16.t13 t16.t14 t16.t15))
(step t16.t17 (cl (= (= arg1 (+ (* -1 14) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 (+ -14 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule cong :premises (t16.t12 t16.t16))
(step t16.t18 (cl (= (= arg1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (= arg1 (+ -14 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule cong :premises (t16.t12 t16.t3))
(step t16.t19 (cl (= (= arg1 (+ -14 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0))))))))) :rule symm :premises (t16.t18))
(step t16.t20 (cl (= (= arg1 (+ (* -1 14) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0))))))))) :rule trans :premises (t16.t17 t16.t19))
(step t16.t21 (cl (not (= (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 (+ (* -1 14) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) (not (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (= arg1 (+ (* -1 14) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule equiv_pos2)
(step t16.t22 (cl (not (= (= (+ arg0 (* -1 -14)) 0) (= arg0 (* -1 14)))) (not (= (+ arg0 (* -1 -14)) 0)) (= arg0 (* -1 14))) :rule equiv_pos2)
(step t16.t23 (cl (= arg0 arg0)) :rule refl)
(step t16.t24 (cl (= (* -1 -14) 14)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (* -1 -14) 14) 3 7))
(step t16.t25 (cl (= (+ arg0 (* -1 -14)) (+ arg0 14))) :rule cong :premises (t16.t23 t16.t24))
(step t16.t26 (cl (= (+ arg0 14) (+ 14 arg0))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ arg0 14) (+ 14 arg0)) 3 7))
(step t16.t27 (cl (= (+ arg0 (* -1 -14)) (+ 14 arg0))) :rule trans :premises (t16.t25 t16.t26))
(step t16.t28 (cl (= (= (+ arg0 (* -1 -14)) 0) (= (+ 14 arg0) 0))) :rule cong :premises (t16.t27 t16.t1))
(step t16.t29 (cl (= (= (+ 14 arg0) 0) (= arg0 -14))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= (+ 14 arg0) 0) (= arg0 -14)) 3 7))
(step t16.t30 (cl (= (= (+ arg0 (* -1 -14)) 0) (= arg0 -14))) :rule trans :premises (t16.t28 t16.t29))
(step t16.t31 (cl (= (= arg0 (* -1 14)) (= arg0 -14))) :rule cong :premises (t16.t23 t16.t13))
(step t16.t32 (cl (= (= arg0 -14) (= arg0 (* -1 14)))) :rule symm :premises (t16.t31))
(step t16.t33 (cl (= (= (+ arg0 (* -1 -14)) 0) (= arg0 (* -1 14)))) :rule trans :premises (t16.t30 t16.t32))
(step t16.t34 (cl (not (= (= (+ arg0 distance) 0) (= (+ arg0 (* -1 -14)) 0))) (not (= (+ arg0 distance) 0)) (= (+ arg0 (* -1 -14)) 0)) :rule equiv_pos2)
(step t16.t35 (cl (not (= (= distance 14) (= distance (* -1 -14)))) (not (= distance 14)) (= distance (* -1 -14))) :rule equiv_pos2)
(step t16.t36 (cl (= distance distance)) :rule refl)
(step t16.t37 (cl (= (= distance (* -1 -14)) (= distance 14))) :rule cong :premises (t16.t36 t16.t24))
(step t16.t38 (cl (= (= distance 14) (= distance (* -1 -14)))) :rule symm :premises (t16.t37))
(step t16.t39 (cl (and (and (and (and (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255)) (= (select_format 5) percent)) (= (select_format 5) s)) (= (select_format 5) x)) (= (select_format 5) adr_lo)) (= (select_format 5) adr_medlo)) (= (select_format 5) adr_medhi)) (= (select_format 5) adr_hi)) (= (select_format 5) 255))) (and (and (and (and (and (and (and (and (and (and (and (ite (and (= (select_format 0) percent) (= (select_format 1) s)) (= (s_count 0) 1) (= (s_count 0) 0)) (ite (and (= (select_format 1) percent) (= (select_format 2) s)) (= (s_count 1) (+ (s_count 0) 1)) (= (s_count 1) (s_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) s)) (= (s_count 2) (+ (s_count 1) 1)) (= (s_count 2) (s_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) s)) (= (s_count 3) (+ (s_count 2) 1)) (= (s_count 3) (s_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) s)) (= (s_count 4) (+ (s_count 3) 1)) (= (s_count 4) (s_count 3)))) (ite (and (= (select_format 5) percent) (= (select_format 6) s)) (= (s_count 5) (+ (s_count 4) 1)) (= (s_count 5) (s_count 4)))) (ite (and (= (select_format 0) percent) (= (select_format 1) x)) (= (x_count 0) 1) (= (x_count 0) 0))) (ite (and (= (select_format 1) percent) (= (select_format 2) x)) (= (x_count 1) (+ (x_count 0) 1)) (= (x_count 1) (x_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) x)) (= (x_count 2) (+ (x_count 1) 1)) (= (x_count 2) (x_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) x)) (= (x_count 3) (+ (x_count 2) 1)) (= (x_count 3) (x_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) x)) (= (x_count 4) (+ (x_count 3) 1)) (= (x_count 4) (x_count 3)))) (ite (and (= (select_format 5) percent) (= (select_format 6) x)) (= (x_count 5) (+ (x_count 4) 1)) (= (x_count 5) (x_count 4))))) (and (or (or (or (or (or (= fmt1 (+ fmt0 0)) (= fmt1 (+ fmt0 1))) (= fmt1 (+ fmt0 2))) (= fmt1 (+ fmt0 3))) (= fmt1 (+ fmt0 4))) (= fmt1 (+ fmt0 5))) (or (or (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1))) (= arg1 (+ fmt0 2)))))) :rule and :premises (a0) :args (0))
(step t16.t40 (cl (and (and (and (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255)) (= (select_format 5) percent)) (= (select_format 5) s)) (= (select_format 5) x)) (= (select_format 5) adr_lo)) (= (select_format 5) adr_medlo)) (= (select_format 5) adr_medhi)) (= (select_format 5) adr_hi)) (= (select_format 5) 255))) (and (and (and (and (and (and (and (and (and (and (and (ite (and (= (select_format 0) percent) (= (select_format 1) s)) (= (s_count 0) 1) (= (s_count 0) 0)) (ite (and (= (select_format 1) percent) (= (select_format 2) s)) (= (s_count 1) (+ (s_count 0) 1)) (= (s_count 1) (s_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) s)) (= (s_count 2) (+ (s_count 1) 1)) (= (s_count 2) (s_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) s)) (= (s_count 3) (+ (s_count 2) 1)) (= (s_count 3) (s_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) s)) (= (s_count 4) (+ (s_count 3) 1)) (= (s_count 4) (s_count 3)))) (ite (and (= (select_format 5) percent) (= (select_format 6) s)) (= (s_count 5) (+ (s_count 4) 1)) (= (s_count 5) (s_count 4)))) (ite (and (= (select_format 0) percent) (= (select_format 1) x)) (= (x_count 0) 1) (= (x_count 0) 0))) (ite (and (= (select_format 1) percent) (= (select_format 2) x)) (= (x_count 1) (+ (x_count 0) 1)) (= (x_count 1) (x_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) x)) (= (x_count 2) (+ (x_count 1) 1)) (= (x_count 2) (x_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) x)) (= (x_count 3) (+ (x_count 2) 1)) (= (x_count 3) (x_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) x)) (= (x_count 4) (+ (x_count 3) 1)) (= (x_count 4) (x_count 3)))) (ite (and (= (select_format 5) percent) (= (select_format 6) x)) (= (x_count 5) (+ (x_count 4) 1)) (= (x_count 5) (x_count 4)))))) :rule and :premises (t16.t39) :args (0))
(step t16.t41 (cl (and (and (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255)) (= (select_format 5) percent)) (= (select_format 5) s)) (= (select_format 5) x)) (= (select_format 5) adr_lo)) (= (select_format 5) adr_medlo)) (= (select_format 5) adr_medhi)) (= (select_format 5) adr_hi)) (= (select_format 5) 255)))) :rule and :premises (t16.t40) :args (0))
(step t16.t42 (cl (and (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))))) :rule and :premises (t16.t41) :args (0))
(step t16.t43 (cl (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120))) :rule and :premises (t16.t42) :args (0))
(step t16.t44 (cl (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115))) :rule and :premises (t16.t43) :args (0))
(step t16.t45 (cl (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37))) :rule and :premises (t16.t44) :args (0))
(step t16.t46 (cl (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2))) :rule and :premises (t16.t45) :args (0))
(step t16.t47 (cl (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5))) :rule and :premises (t16.t46) :args (0))
(step t16.t48 (cl (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1))) :rule and :premises (t16.t47) :args (0))
(step t16.t49 (cl (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3))) :rule and :premises (t16.t48) :args (0))
(step t16.t50 (cl (and (= distance 14) (= fmt_length 6))) :rule and :premises (t16.t49) :args (0))
(step t16.t51 (cl (= distance 14)) :rule and :premises (t16.t50) :args (0))
(step t16.t52 (cl (= distance (* -1 -14))) :rule resolution :premises (t16.t35 t16.t38 t16.t51))
(step t16.t53 (cl (= (+ arg0 distance) (+ arg0 (* -1 -14)))) :rule cong :premises (t16.t23 t16.t52))
(step t16.t54 (cl (= (= (+ arg0 distance) 0) (= (+ arg0 (* -1 -14)) 0))) :rule cong :premises (t16.t53 t16.t1))
(step t16.t55 (cl (not (= (= arg0 (- fmt0 distance)) (= fmt0 (+ arg0 distance)))) (not (= arg0 (- fmt0 distance))) (= fmt0 (+ arg0 distance))) :rule equiv_pos2)
(step t16.t56 (cl (= (- fmt0 distance) (+ fmt0 (* -1 distance)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- fmt0 distance) (+ fmt0 (* -1 distance))) 3 6))
(step t16.t57 (cl (= (= arg0 (- fmt0 distance)) (= arg0 (+ fmt0 (* -1 distance))))) :rule cong :premises (t16.t23 t16.t56))
(step t16.t58 (cl (= (= arg0 (+ fmt0 (* -1 distance))) (= fmt0 (+ arg0 distance)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= arg0 (+ fmt0 (* -1 distance))) (= fmt0 (+ arg0 distance))) 3 7))
(step t16.t59 (cl (= (= arg0 (- fmt0 distance)) (= fmt0 (+ arg0 distance)))) :rule trans :premises (t16.t57 t16.t58))
(step t16.t60 (cl (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) :rule and :premises (t16.t42) :args (1))
(step t16.t61 (cl (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)))) :rule and :premises (t16.t60) :args (0))
(step t16.t62 (cl (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance)))) :rule and :premises (t16.t61) :args (0))
(step t16.t63 (cl (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1)))) :rule and :premises (t16.t62) :args (0))
(step t16.t64 (cl (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1)))) :rule and :premises (t16.t63) :args (0))
(step t16.t65 (cl (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0))) :rule and :premises (t16.t64) :args (0))
(step t16.t66 (cl (and (= fmt0 0) (= arg0 (- fmt0 distance)))) :rule and :premises (t16.t65) :args (0))
(step t16.t67 (cl (= arg0 (- fmt0 distance))) :rule and :premises (t16.t66) :args (1))
(step t16.t68 (cl (= fmt0 (+ arg0 distance))) :rule resolution :premises (t16.t55 t16.t59 t16.t67))
(step t16.t69 (cl (= (+ arg0 distance) fmt0)) :rule symm :premises (t16.t68))
(step t16.t70 (cl (= fmt0 0)) :rule and :premises (t16.t66) :args (0))
(step t16.t71 (cl (= (+ arg0 distance) 0)) :rule trans :premises (t16.t69 t16.t70))
(step t16.t72 (cl (= (+ arg0 (* -1 -14)) 0)) :rule resolution :premises (t16.t34 t16.t54 t16.t71))
(step t16.t73 (cl (= arg0 (* -1 14))) :rule resolution :premises (t16.t22 t16.t33 t16.t72))
(step t16.t74 (cl (= (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))) (+ (* -1 14) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule cong :premises (t16.t73 t16.t14 t16.t15))
(step t16.t75 (cl (= (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 (+ (* -1 14) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule cong :premises (t16.t12 t16.t74))
(step t16.t76 (cl (not (= (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) (not (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule equiv_pos2)
(step t16.t77 (cl (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0)) 3 6))
(step t16.t78 (cl (= 4 4)) :rule refl)
(step t16.t79 (cl (= (- (- fmt1 2) fmt0) (+ (- fmt1 2) (* -1 fmt0)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- (- fmt1 2) fmt0) (+ (- fmt1 2) (* -1 fmt0))) 3 6))
(step t16.t80 (cl (= (- fmt1 2) (+ fmt1 (* -1 2)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- fmt1 2) (+ fmt1 (* -1 2))) 3 6))
(step t16.t81 (cl (= fmt1 fmt1)) :rule refl)
(step t16.t82 (cl (= (* -1 2) -2)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (* -1 2) -2) 3 7))
(step t16.t83 (cl (= (+ fmt1 (* -1 2)) (+ fmt1 -2))) :rule cong :premises (t16.t81 t16.t82))
(step t16.t84 (cl (= (+ fmt1 -2) (+ -2 fmt1))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ fmt1 -2) (+ -2 fmt1)) 3 7))
(step t16.t85 (cl (= (+ fmt1 (* -1 2)) (+ -2 fmt1))) :rule trans :premises (t16.t83 t16.t84))
(step t16.t86 (cl (= (- fmt1 2) (+ -2 fmt1))) :rule trans :premises (t16.t80 t16.t85))
(step t16.t87 (cl (= (* -1 fmt0) (* -1 fmt0))) :rule refl)
(step t16.t88 (cl (= (+ (- fmt1 2) (* -1 fmt0)) (+ (+ -2 fmt1) (* -1 fmt0)))) :rule cong :premises (t16.t86 t16.t87))
(step t16.t89 (cl (= (+ (+ -2 fmt1) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ -2 fmt1) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0))) 3 7))
(step t16.t90 (cl (= (+ (- fmt1 2) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0)))) :rule trans :premises (t16.t88 t16.t89))
(step t16.t91 (cl (= (- (- fmt1 2) fmt0) (+ -2 fmt1 (* -1 fmt0)))) :rule trans :premises (t16.t79 t16.t90))
(step t16.t92 (cl (= (x_count (- (- fmt1 2) fmt0)) (x_count (+ -2 fmt1 (* -1 fmt0))))) :rule cong :premises (t16.t91))
(step t16.t93 (cl (= (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))))) :rule cong :premises (t16.t78 t16.t92))
(step t16.t94 (cl (= (s_count (- (- fmt1 2) fmt0)) (s_count (+ -2 fmt1 (* -1 fmt0))))) :rule cong :premises (t16.t91))
(step t16.t95 (cl (= (* 4 (s_count (- (- fmt1 2) fmt0))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) :rule cong :premises (t16.t78 t16.t94))
(step t16.t96 (cl (= (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0) (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0))) :rule cong :premises (t16.t93 t16.t95 t16.t23))
(step t16.t97 (cl (= (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 3 7))
(step t16.t98 (cl (= (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule trans :premises (t16.t96 t16.t97))
(step t16.t99 (cl (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule trans :premises (t16.t77 t16.t98))
(step t16.t100 (cl (= (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule cong :premises (t16.t12 t16.t99))
(step t16.t101 (cl (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))) :rule and :premises (t16.t60) :args (1))
(step t16.t102 (cl (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule resolution :premises (t16.t76 t16.t100 t16.t101))
(step t16.t103 (cl (= arg1 (+ (* -1 14) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule resolution :premises (t16.t21 t16.t75 t16.t102))
(step t16.t104 (cl (= arg1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule resolution :premises (t16.t11 t16.t20 t16.t103))
(step t16.t105 (cl (= (+ -1 arg1) (+ -1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0))))))))) :rule cong :premises (t16.t2 t16.t104))
(step t16.t106 (cl (= (= 0 (+ -1 arg1)) (= 0 (+ -1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))))) :rule cong :premises (t16.t1 t16.t105))
(step t16.t107 (cl (= 0 (+ -1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0))))))))) :rule resolution :premises (t16.t10 t16.t106 t16.a3))
(step t16.t108 (cl false) :rule resolution :premises (t16.t0 t16.t9 t16.t107))
(step t16 (cl (not (= distance 14)) (not (= (+ arg0 distance) 0)) (not (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (not (= 0 (+ -1 arg1))) false) :rule subproof :discharge (t16.a0 t16.a1 t16.a2 t16.a3))
(step t17 (cl (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) (= distance 14)) :rule and_pos :args (0))
(step t18 (cl (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) (= (+ arg0 distance) 0)) :rule and_pos :args (1))
(step t19 (cl (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule and_pos :args (2))
(step t20 (cl (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) (= 0 (+ -1 arg1))) :rule and_pos :args (3))
(step t21 (cl false (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1))))) :rule resolution :premises (t16 t17 t18 t19 t20))
(step t22 (cl (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) false) :rule reordering :premises (t21))
(step t23 (cl (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))) false) :rule contraction :premises (t22))
(step t24 (cl (=> (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1))) false) false) :rule resolution :premises (t15 t23))
(step t25 (cl (=> (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1))) false) (not false)) :rule implies_neg2)
(step t26 (cl (=> (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1))) false) (=> (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1))) false)) :rule resolution :premises (t24 t25))
(step t27 (cl (=> (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1))) false)) :rule contraction :premises (t26))
(step t28 (cl (= (=> (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1))) false) (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1)))))) :rule implies_simplify)
(step t29 (cl (not (=> (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1))) false)) (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1))))) :rule equiv1 :premises (t28))
(step t30 (cl (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= 0 (+ -1 arg1))))) :rule resolution :premises (t27 t29))
(step t31 (cl (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 (+ -1 arg1))))) :rule resolution :premises (t14 t30))
(step t32 (cl (not (= distance 14)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (not (= fmt0 (+ -1 arg1)))) :rule not_and :premises (t31))
(step t33 (cl (not (= (or (or (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1))) (= arg1 (+ fmt0 2))) (or (= fmt0 (+ -2 arg1)) (= fmt0 arg1) (= fmt0 (+ -1 arg1))))) (not (or (or (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1))) (= arg1 (+ fmt0 2)))) (or (= fmt0 (+ -2 arg1)) (= fmt0 arg1) (= fmt0 (+ -1 arg1)))) :rule equiv_pos2)
(step t34 (cl (= (or (or (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1))) (= arg1 (+ fmt0 2))) (or (= arg1 (+ fmt0 2)) (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (or (or (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1))) (= arg1 (+ fmt0 2))) (or (= arg1 (+ fmt0 2)) (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1)))) 1 6))
(step t35 (cl (= arg1 arg1)) :rule refl)
(step t36 (cl (= (+ fmt0 2) (+ 2 fmt0))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ fmt0 2) (+ 2 fmt0)) 3 7))
(step t37 (cl (= (= arg1 (+ fmt0 2)) (= arg1 (+ 2 fmt0)))) :rule cong :premises (t35 t36))
(step t38 (cl (= (= arg1 (+ 2 fmt0)) (= fmt0 (+ -2 arg1)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= arg1 (+ 2 fmt0)) (= fmt0 (+ -2 arg1))) 3 7))
(step t39 (cl (= (= arg1 (+ fmt0 2)) (= fmt0 (+ -2 arg1)))) :rule trans :premises (t37 t38))
(step t40 (cl (= (+ fmt0 0) fmt0)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ fmt0 0) fmt0) 3 7))
(step t41 (cl (= (= arg1 (+ fmt0 0)) (= arg1 fmt0))) :rule cong :premises (t35 t40))
(step t42 (cl (= (= arg1 fmt0) (= fmt0 arg1))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= arg1 fmt0) (= fmt0 arg1)) 3 7))
(step t43 (cl (= (= arg1 (+ fmt0 0)) (= fmt0 arg1))) :rule trans :premises (t41 t42))
(step t44 (cl (= (+ fmt0 1) (+ 1 fmt0))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ fmt0 1) (+ 1 fmt0)) 3 7))
(step t45 (cl (= (= arg1 (+ fmt0 1)) (= arg1 (+ 1 fmt0)))) :rule cong :premises (t35 t44))
(step t46 (cl (= (= arg1 (+ 1 fmt0)) (= fmt0 (+ -1 arg1)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= arg1 (+ 1 fmt0)) (= fmt0 (+ -1 arg1))) 3 7))
(step t47 (cl (= (= arg1 (+ fmt0 1)) (= fmt0 (+ -1 arg1)))) :rule trans :premises (t45 t46))
(step t48 (cl (= (or (= arg1 (+ fmt0 2)) (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1))) (or (= fmt0 (+ -2 arg1)) (= fmt0 arg1) (= fmt0 (+ -1 arg1))))) :rule cong :premises (t39 t43 t47))
(step t49 (cl (= (or (or (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1))) (= arg1 (+ fmt0 2))) (or (= fmt0 (+ -2 arg1)) (= fmt0 arg1) (= fmt0 (+ -1 arg1))))) :rule trans :premises (t34 t48))
(step t50 (cl (and (and (and (and (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255)) (= (select_format 5) percent)) (= (select_format 5) s)) (= (select_format 5) x)) (= (select_format 5) adr_lo)) (= (select_format 5) adr_medlo)) (= (select_format 5) adr_medhi)) (= (select_format 5) adr_hi)) (= (select_format 5) 255))) (and (and (and (and (and (and (and (and (and (and (and (ite (and (= (select_format 0) percent) (= (select_format 1) s)) (= (s_count 0) 1) (= (s_count 0) 0)) (ite (and (= (select_format 1) percent) (= (select_format 2) s)) (= (s_count 1) (+ (s_count 0) 1)) (= (s_count 1) (s_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) s)) (= (s_count 2) (+ (s_count 1) 1)) (= (s_count 2) (s_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) s)) (= (s_count 3) (+ (s_count 2) 1)) (= (s_count 3) (s_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) s)) (= (s_count 4) (+ (s_count 3) 1)) (= (s_count 4) (s_count 3)))) (ite (and (= (select_format 5) percent) (= (select_format 6) s)) (= (s_count 5) (+ (s_count 4) 1)) (= (s_count 5) (s_count 4)))) (ite (and (= (select_format 0) percent) (= (select_format 1) x)) (= (x_count 0) 1) (= (x_count 0) 0))) (ite (and (= (select_format 1) percent) (= (select_format 2) x)) (= (x_count 1) (+ (x_count 0) 1)) (= (x_count 1) (x_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) x)) (= (x_count 2) (+ (x_count 1) 1)) (= (x_count 2) (x_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) x)) (= (x_count 3) (+ (x_count 2) 1)) (= (x_count 3) (x_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) x)) (= (x_count 4) (+ (x_count 3) 1)) (= (x_count 4) (x_count 3)))) (ite (and (= (select_format 5) percent) (= (select_format 6) x)) (= (x_count 5) (+ (x_count 4) 1)) (= (x_count 5) (x_count 4))))) (and (or (or (or (or (or (= fmt1 (+ fmt0 0)) (= fmt1 (+ fmt0 1))) (= fmt1 (+ fmt0 2))) (= fmt1 (+ fmt0 3))) (= fmt1 (+ fmt0 4))) (= fmt1 (+ fmt0 5))) (or (or (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1))) (= arg1 (+ fmt0 2)))))) :rule and :premises (a0) :args (0))
(step t51 (cl (and (or (or (or (or (or (= fmt1 (+ fmt0 0)) (= fmt1 (+ fmt0 1))) (= fmt1 (+ fmt0 2))) (= fmt1 (+ fmt0 3))) (= fmt1 (+ fmt0 4))) (= fmt1 (+ fmt0 5))) (or (or (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1))) (= arg1 (+ fmt0 2))))) :rule and :premises (t50) :args (1))
(step t52 (cl (or (or (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1))) (= arg1 (+ fmt0 2)))) :rule and :premises (t51) :args (1))
(step t53 (cl (or (= fmt0 (+ -2 arg1)) (= fmt0 arg1) (= fmt0 (+ -1 arg1)))) :rule resolution :premises (t33 t49 t52))
(step t54 (cl (= fmt0 (+ -2 arg1)) (= fmt0 arg1) (= fmt0 (+ -1 arg1))) :rule or :premises (t53))
(step t55 (cl (not (= (or (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -2 arg1)))) (or (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) (not (= fmt0 (+ -2 arg1)))))) (not (or (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -2 arg1))))) (or (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) (not (= fmt0 (+ -2 arg1))))) :rule equiv_pos2)
(step t56 (cl (= (not (= fmt_length 6)) (not (= fmt_length 6)))) :rule refl)
(step t57 (cl (= (not (= fmt0 0)) (not (= fmt0 0)))) :rule refl)
(step t58 (cl (= (not (= fmt0 (+ arg0 distance))) (not (= fmt0 (+ arg0 distance))))) :rule refl)
(step t59 (cl (= (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) 1 6))
(step t60 (cl (= (not (= fmt0 (+ -2 arg1))) (not (= fmt0 (+ -2 arg1))))) :rule refl)
(step t61 (cl (= (or (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -2 arg1)))) (or (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) (not (= fmt0 (+ -2 arg1)))))) :rule cong :premises (t56 t57 t58 t59 t60))
(step t62 (cl (=> (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1))) (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1)))) :rule implies_neg1)
(anchor :step t63)
(assume t63.a0 (= fmt_length 6))
(assume t63.a1 (= fmt0 0))
(assume t63.a2 (= fmt0 (+ arg0 distance)))
(assume t63.a3 (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))
(assume t63.a4 (= fmt0 (+ -2 arg1)))
(step t63.t0 (cl (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt_length 6)) (not (= 0 (+ -2 arg1))) (not (= (+ arg0 distance) 0)) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule and_neg)
(step t63.t1 (cl (and (and (and (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255)) (= (select_format 5) percent)) (= (select_format 5) s)) (= (select_format 5) x)) (= (select_format 5) adr_lo)) (= (select_format 5) adr_medlo)) (= (select_format 5) adr_medhi)) (= (select_format 5) adr_hi)) (= (select_format 5) 255))) (and (and (and (and (and (and (and (and (and (and (and (ite (and (= (select_format 0) percent) (= (select_format 1) s)) (= (s_count 0) 1) (= (s_count 0) 0)) (ite (and (= (select_format 1) percent) (= (select_format 2) s)) (= (s_count 1) (+ (s_count 0) 1)) (= (s_count 1) (s_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) s)) (= (s_count 2) (+ (s_count 1) 1)) (= (s_count 2) (s_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) s)) (= (s_count 3) (+ (s_count 2) 1)) (= (s_count 3) (s_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) s)) (= (s_count 4) (+ (s_count 3) 1)) (= (s_count 4) (s_count 3)))) (ite (and (= (select_format 5) percent) (= (select_format 6) s)) (= (s_count 5) (+ (s_count 4) 1)) (= (s_count 5) (s_count 4)))) (ite (and (= (select_format 0) percent) (= (select_format 1) x)) (= (x_count 0) 1) (= (x_count 0) 0))) (ite (and (= (select_format 1) percent) (= (select_format 2) x)) (= (x_count 1) (+ (x_count 0) 1)) (= (x_count 1) (x_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) x)) (= (x_count 2) (+ (x_count 1) 1)) (= (x_count 2) (x_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) x)) (= (x_count 3) (+ (x_count 2) 1)) (= (x_count 3) (x_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) x)) (= (x_count 4) (+ (x_count 3) 1)) (= (x_count 4) (x_count 3)))) (ite (and (= (select_format 5) percent) (= (select_format 6) x)) (= (x_count 5) (+ (x_count 4) 1)) (= (x_count 5) (x_count 4)))))) :rule and :premises (t50) :args (0))
(step t63.t2 (cl (and (and (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255)) (= (select_format 5) percent)) (= (select_format 5) s)) (= (select_format 5) x)) (= (select_format 5) adr_lo)) (= (select_format 5) adr_medlo)) (= (select_format 5) adr_medhi)) (= (select_format 5) adr_hi)) (= (select_format 5) 255)))) :rule and :premises (t63.t1) :args (0))
(step t63.t3 (cl (and (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))))) :rule and :premises (t63.t2) :args (0))
(step t63.t4 (cl (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120))) :rule and :premises (t63.t3) :args (0))
(step t63.t5 (cl (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115))) :rule and :premises (t63.t4) :args (0))
(step t63.t6 (cl (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37))) :rule and :premises (t63.t5) :args (0))
(step t63.t7 (cl (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2))) :rule and :premises (t63.t6) :args (0))
(step t63.t8 (cl (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5))) :rule and :premises (t63.t7) :args (0))
(step t63.t9 (cl (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1))) :rule and :premises (t63.t8) :args (0))
(step t63.t10 (cl (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3))) :rule and :premises (t63.t9) :args (0))
(step t63.t11 (cl (and (= distance 14) (= fmt_length 6))) :rule and :premises (t63.t10) :args (0))
(step t63.t12 (cl (= fmt_length 6)) :rule and :premises (t63.t11) :args (1))
(step t63.t13 (cl (=> (and (= fmt0 (+ -2 arg1)) (= fmt0 0)) (= 0 (+ -2 arg1))) (and (= fmt0 (+ -2 arg1)) (= fmt0 0))) :rule implies_neg1)
(anchor :step t63.t14)
(assume t63.t14.a0 (= fmt0 (+ -2 arg1)))
(assume t63.t14.a1 (= fmt0 0))
(step t63.t14.t0 (cl (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) :rule and :premises (t63.t3) :args (1))
(step t63.t14.t1 (cl (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)))) :rule and :premises (t63.t14.t0) :args (0))
(step t63.t14.t2 (cl (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance)))) :rule and :premises (t63.t14.t1) :args (0))
(step t63.t14.t3 (cl (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1)))) :rule and :premises (t63.t14.t2) :args (0))
(step t63.t14.t4 (cl (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1)))) :rule and :premises (t63.t14.t3) :args (0))
(step t63.t14.t5 (cl (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0))) :rule and :premises (t63.t14.t4) :args (0))
(step t63.t14.t6 (cl (and (= fmt0 0) (= arg0 (- fmt0 distance)))) :rule and :premises (t63.t14.t5) :args (0))
(step t63.t14.t7 (cl (= fmt0 0)) :rule and :premises (t63.t14.t6) :args (0))
(step t63.t14.t8 (cl (= 0 fmt0)) :rule symm :premises (t63.t14.t7))
(step t63.t14.t9 (cl (= (+ -2 arg1) fmt0)) :rule symm :premises (t63.t14.a0))
(step t63.t14.t10 (cl (= fmt0 (+ -2 arg1))) :rule symm :premises (t63.t14.t9))
(step t63.t14.t11 (cl (= 0 (+ -2 arg1))) :rule trans :premises (t63.t14.t8 t63.t14.t10))
(step t63.t14 (cl (not (= fmt0 (+ -2 arg1))) (not (= fmt0 0)) (= 0 (+ -2 arg1))) :rule subproof :discharge (t63.t14.a0 t63.t14.a1))
(step t63.t15 (cl (not (and (= fmt0 (+ -2 arg1)) (= fmt0 0))) (= fmt0 (+ -2 arg1))) :rule and_pos :args (0))
(step t63.t16 (cl (not (and (= fmt0 (+ -2 arg1)) (= fmt0 0))) (= fmt0 0)) :rule and_pos :args (1))
(step t63.t17 (cl (= 0 (+ -2 arg1)) (not (and (= fmt0 (+ -2 arg1)) (= fmt0 0))) (not (and (= fmt0 (+ -2 arg1)) (= fmt0 0)))) :rule resolution :premises (t63.t14 t63.t15 t63.t16))
(step t63.t18 (cl (not (and (= fmt0 (+ -2 arg1)) (= fmt0 0))) (not (and (= fmt0 (+ -2 arg1)) (= fmt0 0))) (= 0 (+ -2 arg1))) :rule reordering :premises (t63.t17))
(step t63.t19 (cl (not (and (= fmt0 (+ -2 arg1)) (= fmt0 0))) (= 0 (+ -2 arg1))) :rule contraction :premises (t63.t18))
(step t63.t20 (cl (=> (and (= fmt0 (+ -2 arg1)) (= fmt0 0)) (= 0 (+ -2 arg1))) (= 0 (+ -2 arg1))) :rule resolution :premises (t63.t13 t63.t19))
(step t63.t21 (cl (=> (and (= fmt0 (+ -2 arg1)) (= fmt0 0)) (= 0 (+ -2 arg1))) (not (= 0 (+ -2 arg1)))) :rule implies_neg2)
(step t63.t22 (cl (=> (and (= fmt0 (+ -2 arg1)) (= fmt0 0)) (= 0 (+ -2 arg1))) (=> (and (= fmt0 (+ -2 arg1)) (= fmt0 0)) (= 0 (+ -2 arg1)))) :rule resolution :premises (t63.t20 t63.t21))
(step t63.t23 (cl (=> (and (= fmt0 (+ -2 arg1)) (= fmt0 0)) (= 0 (+ -2 arg1)))) :rule contraction :premises (t63.t22))
(step t63.t24 (cl (not (and (= fmt0 (+ -2 arg1)) (= fmt0 0))) (= 0 (+ -2 arg1))) :rule implies :premises (t63.t23))
(step t63.t25 (cl (and (= fmt0 (+ -2 arg1)) (= fmt0 0)) (not (= fmt0 (+ -2 arg1))) (not (= fmt0 0))) :rule and_neg)
(step t63.t26 (cl (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) :rule and :premises (t63.t3) :args (1))
(step t63.t27 (cl (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)))) :rule and :premises (t63.t26) :args (0))
(step t63.t28 (cl (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance)))) :rule and :premises (t63.t27) :args (0))
(step t63.t29 (cl (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1)))) :rule and :premises (t63.t28) :args (0))
(step t63.t30 (cl (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1)))) :rule and :premises (t63.t29) :args (0))
(step t63.t31 (cl (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0))) :rule and :premises (t63.t30) :args (0))
(step t63.t32 (cl (and (= fmt0 0) (= arg0 (- fmt0 distance)))) :rule and :premises (t63.t31) :args (0))
(step t63.t33 (cl (= fmt0 0)) :rule and :premises (t63.t32) :args (0))
(step t63.t34 (cl (and (= fmt0 (+ -2 arg1)) (= fmt0 0))) :rule resolution :premises (t63.t25 t63.a4 t63.t33))
(step t63.t35 (cl (= 0 (+ -2 arg1))) :rule resolution :premises (t63.t24 t63.t34))
(step t63.t36 (cl (not (= (= arg0 (- fmt0 distance)) (= fmt0 (+ arg0 distance)))) (not (= arg0 (- fmt0 distance))) (= fmt0 (+ arg0 distance))) :rule equiv_pos2)
(step t63.t37 (cl (= arg0 arg0)) :rule refl)
(step t63.t38 (cl (= (- fmt0 distance) (+ fmt0 (* -1 distance)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- fmt0 distance) (+ fmt0 (* -1 distance))) 3 6))
(step t63.t39 (cl (= (= arg0 (- fmt0 distance)) (= arg0 (+ fmt0 (* -1 distance))))) :rule cong :premises (t63.t37 t63.t38))
(step t63.t40 (cl (= (= arg0 (+ fmt0 (* -1 distance))) (= fmt0 (+ arg0 distance)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= arg0 (+ fmt0 (* -1 distance))) (= fmt0 (+ arg0 distance))) 3 7))
(step t63.t41 (cl (= (= arg0 (- fmt0 distance)) (= fmt0 (+ arg0 distance)))) :rule trans :premises (t63.t39 t63.t40))
(step t63.t42 (cl (= arg0 (- fmt0 distance))) :rule and :premises (t63.t32) :args (1))
(step t63.t43 (cl (= fmt0 (+ arg0 distance))) :rule resolution :premises (t63.t36 t63.t41 t63.t42))
(step t63.t44 (cl (= (+ arg0 distance) fmt0)) :rule symm :premises (t63.t43))
(step t63.t45 (cl (= (+ arg0 distance) 0)) :rule trans :premises (t63.t44 t63.t33))
(step t63.t46 (cl (not (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule equiv_pos2)
(step t63.t47 (cl (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4)))) 3 6))
(step t63.t48 (cl (= (- (+ (+ arg0 distance) fmt_length) 4) (+ (+ (+ arg0 distance) fmt_length) (* -1 4)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- (+ (+ arg0 distance) fmt_length) 4) (+ (+ (+ arg0 distance) fmt_length) (* -1 4))) 3 6))
(step t63.t49 (cl (= (+ (+ (+ arg0 distance) fmt_length) (* -1 4)) (+ distance fmt_length arg0 -4))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ (+ arg0 distance) fmt_length) (* -1 4)) (+ distance fmt_length arg0 -4)) 3 6))
(step t63.t50 (cl (= (+ distance fmt_length arg0 -4) (+ -4 arg0 fmt_length distance))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ distance fmt_length arg0 -4) (+ -4 arg0 fmt_length distance)) 3 7))
(step t63.t51 (cl (= (+ (+ (+ arg0 distance) fmt_length) (* -1 4)) (+ -4 arg0 fmt_length distance))) :rule trans :premises (t63.t49 t63.t50))
(step t63.t52 (cl (= (- (+ (+ arg0 distance) fmt_length) 4) (+ -4 arg0 fmt_length distance))) :rule trans :premises (t63.t48 t63.t51))
(step t63.t53 (cl (= (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (>= arg1 (+ -4 arg0 fmt_length distance)))) :rule cong :premises (t35 t63.t52))
(step t63.t54 (cl (= (>= arg1 (+ -4 arg0 fmt_length distance)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (>= arg1 (+ -4 arg0 fmt_length distance)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) 3 7))
(step t63.t55 (cl (= (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule trans :premises (t63.t53 t63.t54))
(step t63.t56 (cl (= (not (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule cong :premises (t63.t55))
(step t63.t57 (cl (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule trans :premises (t63.t47 t63.t56))
(step t63.t58 (cl (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) :rule and :premises (t63.t27) :args (1))
(step t63.t59 (cl (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule resolution :premises (t63.t46 t63.t57 t63.t58))
(step t63.t60 (cl (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule resolution :premises (t63.t0 t63.t12 t63.t35 t63.t45 t63.t59))
(step t63 (cl (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -2 arg1))) (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule subproof :discharge (t63.a0 t63.a1 t63.a2 t63.a3 t63.a4))
(step t64 (cl (not (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1)))) (= fmt_length 6)) :rule and_pos :args (0))
(step t65 (cl (not (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1)))) (= fmt0 0)) :rule and_pos :args (1))
(step t66 (cl (not (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1)))) (= fmt0 (+ arg0 distance))) :rule and_pos :args (2))
(step t67 (cl (not (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1)))) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule and_pos :args (3))
(step t68 (cl (not (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1)))) (= fmt0 (+ -2 arg1))) :rule and_pos :args (4))
(step t69 (cl (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1)))) (not (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1)))) (not (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1)))) (not (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1)))) (not (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1))))) :rule resolution :premises (t63 t64 t65 t66 t67 t68))
(step t70 (cl (not (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1)))) (not (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1)))) (not (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1)))) (not (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1)))) (not (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1)))) (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule reordering :premises (t69))
(step t71 (cl (not (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1)))) (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule contraction :premises (t70))
(step t72 (cl (=> (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1))) (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule resolution :premises (t62 t71))
(step t73 (cl (=> (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1))) (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))))) :rule implies_neg2)
(step t74 (cl (=> (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1))) (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (=> (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1))) (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))))) :rule resolution :premises (t72 t73))
(step t75 (cl (=> (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1))) (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))))) :rule contraction :premises (t74))
(step t76 (cl (not (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1)))) (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule implies :premises (t75))
(step t77 (cl (=> (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) false) (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule implies_neg1)
(anchor :step t78)
(assume t78.a0 (= fmt_length 6))
(assume t78.a1 (= 0 (+ -2 arg1)))
(assume t78.a2 (= (+ arg0 distance) 0))
(assume t78.a3 (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))
(step t78.t0 (cl (not (= (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) :rule equiv_pos2)
(step t78.t1 (cl (= (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) 3 6))
(step t78.t2 (cl (= (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule symm :premises (t78.t1))
(step t78.t3 (cl (not (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule equiv_pos2)
(step t78.t4 (cl (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4)))) 3 6))
(step t78.t5 (cl (= (- (+ (+ arg0 distance) fmt_length) 4) (+ (+ (+ arg0 distance) fmt_length) (* -1 4)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- (+ (+ arg0 distance) fmt_length) 4) (+ (+ (+ arg0 distance) fmt_length) (* -1 4))) 3 6))
(step t78.t6 (cl (= (+ (+ (+ arg0 distance) fmt_length) (* -1 4)) (+ distance fmt_length arg0 -4))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ (+ arg0 distance) fmt_length) (* -1 4)) (+ distance fmt_length arg0 -4)) 3 6))
(step t78.t7 (cl (= (+ distance fmt_length arg0 -4) (+ -4 arg0 fmt_length distance))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ distance fmt_length arg0 -4) (+ -4 arg0 fmt_length distance)) 3 7))
(step t78.t8 (cl (= (+ (+ (+ arg0 distance) fmt_length) (* -1 4)) (+ -4 arg0 fmt_length distance))) :rule trans :premises (t78.t6 t78.t7))
(step t78.t9 (cl (= (- (+ (+ arg0 distance) fmt_length) 4) (+ -4 arg0 fmt_length distance))) :rule trans :premises (t78.t5 t78.t8))
(step t78.t10 (cl (= (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (>= arg1 (+ -4 arg0 fmt_length distance)))) :rule cong :premises (t35 t78.t9))
(step t78.t11 (cl (= (>= arg1 (+ -4 arg0 fmt_length distance)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (>= arg1 (+ -4 arg0 fmt_length distance)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) 3 7))
(step t78.t12 (cl (= (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule trans :premises (t78.t10 t78.t11))
(step t78.t13 (cl (= (not (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule cong :premises (t78.t12))
(step t78.t14 (cl (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule trans :premises (t78.t4 t78.t13))
(step t78.t15 (cl (and (and (and (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255)) (= (select_format 5) percent)) (= (select_format 5) s)) (= (select_format 5) x)) (= (select_format 5) adr_lo)) (= (select_format 5) adr_medlo)) (= (select_format 5) adr_medhi)) (= (select_format 5) adr_hi)) (= (select_format 5) 255))) (and (and (and (and (and (and (and (and (and (and (and (ite (and (= (select_format 0) percent) (= (select_format 1) s)) (= (s_count 0) 1) (= (s_count 0) 0)) (ite (and (= (select_format 1) percent) (= (select_format 2) s)) (= (s_count 1) (+ (s_count 0) 1)) (= (s_count 1) (s_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) s)) (= (s_count 2) (+ (s_count 1) 1)) (= (s_count 2) (s_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) s)) (= (s_count 3) (+ (s_count 2) 1)) (= (s_count 3) (s_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) s)) (= (s_count 4) (+ (s_count 3) 1)) (= (s_count 4) (s_count 3)))) (ite (and (= (select_format 5) percent) (= (select_format 6) s)) (= (s_count 5) (+ (s_count 4) 1)) (= (s_count 5) (s_count 4)))) (ite (and (= (select_format 0) percent) (= (select_format 1) x)) (= (x_count 0) 1) (= (x_count 0) 0))) (ite (and (= (select_format 1) percent) (= (select_format 2) x)) (= (x_count 1) (+ (x_count 0) 1)) (= (x_count 1) (x_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) x)) (= (x_count 2) (+ (x_count 1) 1)) (= (x_count 2) (x_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) x)) (= (x_count 3) (+ (x_count 2) 1)) (= (x_count 3) (x_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) x)) (= (x_count 4) (+ (x_count 3) 1)) (= (x_count 4) (x_count 3)))) (ite (and (= (select_format 5) percent) (= (select_format 6) x)) (= (x_count 5) (+ (x_count 4) 1)) (= (x_count 5) (x_count 4)))))) :rule and :premises (t50) :args (0))
(step t78.t16 (cl (and (and (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255)) (= (select_format 5) percent)) (= (select_format 5) s)) (= (select_format 5) x)) (= (select_format 5) adr_lo)) (= (select_format 5) adr_medlo)) (= (select_format 5) adr_medhi)) (= (select_format 5) adr_hi)) (= (select_format 5) 255)))) :rule and :premises (t78.t15) :args (0))
(step t78.t17 (cl (and (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))))) :rule and :premises (t78.t16) :args (0))
(step t78.t18 (cl (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) :rule and :premises (t78.t17) :args (1))
(step t78.t19 (cl (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)))) :rule and :premises (t78.t18) :args (0))
(step t78.t20 (cl (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) :rule and :premises (t78.t19) :args (1))
(step t78.t21 (cl (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule resolution :premises (t78.t3 t78.t14 t78.t20))
(step t78.t22 (cl (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) :rule resolution :premises (t78.t0 t78.t2 t78.t21))
(step t78.t23 (cl (not (= (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule equiv_pos2)
(step t78.t24 (cl (= (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))))) :rule cong :premises (t78.t1))
(step t78.t25 (cl (= (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule trans :premises (t78.t24 t59))
(step t78.t26 (cl (= (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule symm :premises (t78.t25))
(step t78.t27 (cl (not (= (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) :rule equiv_pos2)
(step t78.t28 (cl (=> (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) false) (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) :rule implies_neg1)
(anchor :step t78.t29)
(assume t78.t29.a0 (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))
(step t78.t29.t0 (cl (not (= (< (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 2) 6)) false)) (not (< (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 2) 6))) false) :rule equiv_pos2)
(step t78.t29.t1 (cl (= (< (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 2) 6)) (not (>= (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 2) 6))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (< (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 2) 6)) (not (>= (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 2) 6)))) 3 6))
(step t78.t29.t2 (cl (= (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ (* -1 arg0) (* -1 fmt_length) (* -1 arg1) (* -1 distance) distance fmt_length arg0 arg1))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ (* -1 arg0) (* -1 fmt_length) (* -1 arg1) (* -1 distance) distance fmt_length arg0 arg1)) 3 6))
(step t78.t29.t3 (cl (= (+ (* -1 arg0) (* -1 fmt_length) (* -1 arg1) (* -1 distance) distance fmt_length arg0 arg1) 0)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (* -1 arg0) (* -1 fmt_length) (* -1 arg1) (* -1 distance) distance fmt_length arg0 arg1) 0) 3 7))
(step t78.t29.t4 (cl (= (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) 0)) :rule trans :premises (t78.t29.t2 t78.t29.t3))
(step t78.t29.t5 (cl (= -4 -4)) :rule refl)
(step t78.t29.t6 (cl (= 0 0)) :rule refl)
(step t78.t29.t7 (cl (= (* -1 2) -2)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (* -1 2) -2) 3 7))
(step t78.t29.t8 (cl (= 6 6)) :rule refl)
(step t78.t29.t9 (cl (= (+ -4 0 (* -1 2) 6) (+ -4 0 -2 6))) :rule cong :premises (t78.t29.t5 t78.t29.t6 t78.t29.t7 t78.t29.t8))
(step t78.t29.t10 (cl (= (+ -4 0 -2 6) 0)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ -4 0 -2 6) 0) 3 7))
(step t78.t29.t11 (cl (= (+ -4 0 (* -1 2) 6) 0)) :rule trans :premises (t78.t29.t9 t78.t29.t10))
(step t78.t29.t12 (cl (= (>= (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 2) 6)) (>= 0 0))) :rule cong :premises (t78.t29.t4 t78.t29.t11))
(step t78.t29.t13 (cl (= (>= 0 0) true)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (>= 0 0) true) 3 7))
(step t78.t29.t14 (cl (= (>= (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 2) 6)) true)) :rule trans :premises (t78.t29.t12 t78.t29.t13))
(step t78.t29.t15 (cl (= (not (>= (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 2) 6))) (not true))) :rule cong :premises (t78.t29.t14))
(step t78.t29.t16 (cl (= (not true) false)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (not true) false) 1 7))
(step t78.t29.t17 (cl (= (not (>= (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 2) 6))) false)) :rule trans :premises (t78.t29.t15 t78.t29.t16))
(step t78.t29.t18 (cl (= (< (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 2) 6)) false)) :rule trans :premises (t78.t29.t1 t78.t29.t17))
(step t78.t29.t19 (cl (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (not (= (+ arg0 distance) 0)) (not (= (* -1 arg1) (* -1 2))) (not (= fmt_length 6)) (< (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 2) 6))) :rule la_generic :args (1/1 -1/1 -1/1 -1/1 1/1))
(step t78.t29.t20 (cl (not (= (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) :rule equiv_pos2)
(step t78.t29.t21 (cl (not (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule equiv_pos2)
(step t78.t29.t22 (cl (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule resolution :premises (t78.t29.t21 t78.t14 t78.t20))
(step t78.t29.t23 (cl (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) :rule resolution :premises (t78.t29.t20 t78.t2 t78.t29.t22))
(step t78.t29.t24 (cl (not (= (= arg0 (- fmt0 distance)) (= fmt0 (+ arg0 distance)))) (not (= arg0 (- fmt0 distance))) (= fmt0 (+ arg0 distance))) :rule equiv_pos2)
(step t78.t29.t25 (cl (= arg0 arg0)) :rule refl)
(step t78.t29.t26 (cl (= (- fmt0 distance) (+ fmt0 (* -1 distance)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- fmt0 distance) (+ fmt0 (* -1 distance))) 3 6))
(step t78.t29.t27 (cl (= (= arg0 (- fmt0 distance)) (= arg0 (+ fmt0 (* -1 distance))))) :rule cong :premises (t78.t29.t25 t78.t29.t26))
(step t78.t29.t28 (cl (= (= arg0 (+ fmt0 (* -1 distance))) (= fmt0 (+ arg0 distance)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= arg0 (+ fmt0 (* -1 distance))) (= fmt0 (+ arg0 distance))) 3 7))
(step t78.t29.t29 (cl (= (= arg0 (- fmt0 distance)) (= fmt0 (+ arg0 distance)))) :rule trans :premises (t78.t29.t27 t78.t29.t28))
(step t78.t29.t30 (cl (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance)))) :rule and :premises (t78.t19) :args (0))
(step t78.t29.t31 (cl (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1)))) :rule and :premises (t78.t29.t30) :args (0))
(step t78.t29.t32 (cl (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1)))) :rule and :premises (t78.t29.t31) :args (0))
(step t78.t29.t33 (cl (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0))) :rule and :premises (t78.t29.t32) :args (0))
(step t78.t29.t34 (cl (and (= fmt0 0) (= arg0 (- fmt0 distance)))) :rule and :premises (t78.t29.t33) :args (0))
(step t78.t29.t35 (cl (= arg0 (- fmt0 distance))) :rule and :premises (t78.t29.t34) :args (1))
(step t78.t29.t36 (cl (= fmt0 (+ arg0 distance))) :rule resolution :premises (t78.t29.t24 t78.t29.t29 t78.t29.t35))
(step t78.t29.t37 (cl (= (+ arg0 distance) fmt0)) :rule symm :premises (t78.t29.t36))
(step t78.t29.t38 (cl (= fmt0 0)) :rule and :premises (t78.t29.t34) :args (0))
(step t78.t29.t39 (cl (= (+ arg0 distance) 0)) :rule trans :premises (t78.t29.t37 t78.t29.t38))
(step t78.t29.t40 (cl (=> (and (< -1 0) (= arg1 2)) (= (* -1 arg1) (* -1 2)))) :rule la_mult_neg)
(step t78.t29.t41 (cl (not (and (< -1 0) (= arg1 2))) (= (* -1 arg1) (* -1 2))) :rule implies :premises (t78.t29.t40))
(step t78.t29.t42 (cl (and (< -1 0) (= arg1 2)) (not (< -1 0)) (not (= arg1 2))) :rule and_neg)
(step t78.t29.t43 (cl (= (= (< -1 0) true) (< -1 0))) :rule equiv_simplify)
(step t78.t29.t44 (cl (not (= (< -1 0) true)) (< -1 0)) :rule equiv1 :premises (t78.t29.t43))
(step t78.t29.t45 (cl (= (< -1 0) true)) :rule hole :args ("evaluate"))
(step t78.t29.t46 (cl (< -1 0)) :rule resolution :premises (t78.t29.t44 t78.t29.t45))
(step t78.t29.t47 (cl (not (= (= 0 (+ -2 arg1)) (= arg1 2))) (not (= 0 (+ -2 arg1))) (= arg1 2)) :rule equiv_pos2)
(step t78.t29.t48 (cl (= (= 0 (+ -2 arg1)) (= arg1 2))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= 0 (+ -2 arg1)) (= arg1 2)) 3 7))
(step t78.t29.t49 (cl (= arg1 2)) :rule resolution :premises (t78.t29.t47 t78.t29.t48 t78.a1))
(step t78.t29.t50 (cl (and (< -1 0) (= arg1 2))) :rule resolution :premises (t78.t29.t42 t78.t29.t46 t78.t29.t49))
(step t78.t29.t51 (cl (= (* -1 arg1) (* -1 2))) :rule resolution :premises (t78.t29.t41 t78.t29.t50))
(step t78.t29.t52 (cl (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120))) :rule and :premises (t78.t17) :args (0))
(step t78.t29.t53 (cl (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115))) :rule and :premises (t78.t29.t52) :args (0))
(step t78.t29.t54 (cl (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37))) :rule and :premises (t78.t29.t53) :args (0))
(step t78.t29.t55 (cl (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2))) :rule and :premises (t78.t29.t54) :args (0))
(step t78.t29.t56 (cl (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5))) :rule and :premises (t78.t29.t55) :args (0))
(step t78.t29.t57 (cl (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1))) :rule and :premises (t78.t29.t56) :args (0))
(step t78.t29.t58 (cl (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3))) :rule and :premises (t78.t29.t57) :args (0))
(step t78.t29.t59 (cl (and (= distance 14) (= fmt_length 6))) :rule and :premises (t78.t29.t58) :args (0))
(step t78.t29.t60 (cl (= fmt_length 6)) :rule and :premises (t78.t29.t59) :args (1))
(step t78.t29.t61 (cl (< (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 2) 6))) :rule resolution :premises (t78.t29.t19 t78.t29.t23 t78.t29.t39 t78.t29.t51 t78.t29.t60))
(step t78.t29.t62 (cl false) :rule resolution :premises (t78.t29.t0 t78.t29.t18 t78.t29.t61))
(step t78.t29 (cl (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) false) :rule subproof :discharge (t78.t29.a0))
(step t78.t30 (cl (=> (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) false) false) :rule resolution :premises (t78.t28 t78.t29))
(step t78.t31 (cl (=> (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) false) (not false)) :rule implies_neg2)
(step t78.t32 (cl (=> (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) false) (=> (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) false)) :rule resolution :premises (t78.t30 t78.t31))
(step t78.t33 (cl (=> (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) false)) :rule contraction :premises (t78.t32))
(step t78.t34 (cl (= (=> (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) false) (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule implies_simplify)
(step t78.t35 (cl (not (=> (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) false)) (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule equiv1 :premises (t78.t34))
(step t78.t36 (cl (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule resolution :premises (t78.t33 t78.t35))
(step t78.t37 (cl (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) :rule resolution :premises (t78.t27 t78.t25 t78.t36))
(step t78.t38 (cl (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule resolution :premises (t78.t23 t78.t26 t78.t37))
(step t78.t39 (cl) :rule resolution :premises (t78.t22 t78.t38))
(step t78 (cl (not (= fmt_length 6)) (not (= 0 (+ -2 arg1))) (not (= (+ arg0 distance) 0)) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) false) :rule subproof :discharge (t78.a0 t78.a1 t78.a2 t78.a3))
(step t79 (cl (not (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (= fmt_length 6)) :rule and_pos :args (0))
(step t80 (cl (not (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (= 0 (+ -2 arg1))) :rule and_pos :args (1))
(step t81 (cl (not (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (= (+ arg0 distance) 0)) :rule and_pos :args (2))
(step t82 (cl (not (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule and_pos :args (3))
(step t83 (cl false (not (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))))) :rule resolution :premises (t78 t79 t80 t81 t82))
(step t84 (cl (not (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) false) :rule reordering :premises (t83))
(step t85 (cl (not (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) false) :rule contraction :premises (t84))
(step t86 (cl (=> (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) false) false) :rule resolution :premises (t77 t85))
(step t87 (cl (=> (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) false) (not false)) :rule implies_neg2)
(step t88 (cl (=> (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) false) (=> (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) false)) :rule resolution :premises (t86 t87))
(step t89 (cl (=> (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) false)) :rule contraction :premises (t88))
(step t90 (cl (= (=> (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) false) (not (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))))) :rule implies_simplify)
(step t91 (cl (not (=> (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) false)) (not (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))))) :rule equiv1 :premises (t90))
(step t92 (cl (not (and (= fmt_length 6) (= 0 (+ -2 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))))) :rule resolution :premises (t89 t91))
(step t93 (cl (not (and (= fmt_length 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -2 arg1))))) :rule resolution :premises (t76 t92))
(step t94 (cl (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -2 arg1)))) :rule not_and :premises (t93))
(step t95 (cl (or (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -2 arg1)))) (not (not (= fmt_length 6)))) :rule or_neg :args (0))
(step t96 (cl (or (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -2 arg1)))) (not (not (= fmt0 0)))) :rule or_neg :args (1))
(step t97 (cl (or (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -2 arg1)))) (not (not (= fmt0 (+ arg0 distance))))) :rule or_neg :args (2))
(step t98 (cl (or (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -2 arg1)))) (not (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))))) :rule or_neg :args (3))
(step t99 (cl (or (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -2 arg1)))) (not (not (= fmt0 (+ -2 arg1))))) :rule or_neg :args (4))
(step t100 (cl (or (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -2 arg1)))) (or (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -2 arg1)))) (or (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -2 arg1)))) (or (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -2 arg1)))) (or (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -2 arg1))))) :rule resolution :premises (t94 t95 t96 t97 t98 t99))
(step t101 (cl (or (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -2 arg1))))) :rule contraction :premises (t100))
(step t102 (cl (or (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) (not (= fmt0 (+ -2 arg1))))) :rule resolution :premises (t55 t61 t101))
(step t103 (cl (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) (not (= fmt0 (+ -2 arg1)))) :rule or :premises (t102))
(step t104 (cl (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) (not (= fmt_length 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (= fmt0 (+ -2 arg1)))) :rule reordering :premises (t103))
(step t105 (cl (not (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule equiv_pos2)
(step t106 (cl (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4)))) 3 6))
(step t107 (cl (= (- (+ (+ arg0 distance) fmt_length) 4) (+ (+ (+ arg0 distance) fmt_length) (* -1 4)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- (+ (+ arg0 distance) fmt_length) 4) (+ (+ (+ arg0 distance) fmt_length) (* -1 4))) 3 6))
(step t108 (cl (= (+ (+ (+ arg0 distance) fmt_length) (* -1 4)) (+ distance fmt_length arg0 -4))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ (+ arg0 distance) fmt_length) (* -1 4)) (+ distance fmt_length arg0 -4)) 3 6))
(step t109 (cl (= (+ distance fmt_length arg0 -4) (+ -4 arg0 fmt_length distance))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ distance fmt_length arg0 -4) (+ -4 arg0 fmt_length distance)) 3 7))
(step t110 (cl (= (+ (+ (+ arg0 distance) fmt_length) (* -1 4)) (+ -4 arg0 fmt_length distance))) :rule trans :premises (t108 t109))
(step t111 (cl (= (- (+ (+ arg0 distance) fmt_length) 4) (+ -4 arg0 fmt_length distance))) :rule trans :premises (t107 t110))
(step t112 (cl (= (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (>= arg1 (+ -4 arg0 fmt_length distance)))) :rule cong :premises (t35 t111))
(step t113 (cl (= (>= arg1 (+ -4 arg0 fmt_length distance)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (>= arg1 (+ -4 arg0 fmt_length distance)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) 3 7))
(step t114 (cl (= (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule trans :premises (t112 t113))
(step t115 (cl (= (not (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule cong :premises (t114))
(step t116 (cl (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule trans :premises (t106 t115))
(step t117 (cl (and (and (and (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255)) (= (select_format 5) percent)) (= (select_format 5) s)) (= (select_format 5) x)) (= (select_format 5) adr_lo)) (= (select_format 5) adr_medlo)) (= (select_format 5) adr_medhi)) (= (select_format 5) adr_hi)) (= (select_format 5) 255))) (and (and (and (and (and (and (and (and (and (and (and (ite (and (= (select_format 0) percent) (= (select_format 1) s)) (= (s_count 0) 1) (= (s_count 0) 0)) (ite (and (= (select_format 1) percent) (= (select_format 2) s)) (= (s_count 1) (+ (s_count 0) 1)) (= (s_count 1) (s_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) s)) (= (s_count 2) (+ (s_count 1) 1)) (= (s_count 2) (s_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) s)) (= (s_count 3) (+ (s_count 2) 1)) (= (s_count 3) (s_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) s)) (= (s_count 4) (+ (s_count 3) 1)) (= (s_count 4) (s_count 3)))) (ite (and (= (select_format 5) percent) (= (select_format 6) s)) (= (s_count 5) (+ (s_count 4) 1)) (= (s_count 5) (s_count 4)))) (ite (and (= (select_format 0) percent) (= (select_format 1) x)) (= (x_count 0) 1) (= (x_count 0) 0))) (ite (and (= (select_format 1) percent) (= (select_format 2) x)) (= (x_count 1) (+ (x_count 0) 1)) (= (x_count 1) (x_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) x)) (= (x_count 2) (+ (x_count 1) 1)) (= (x_count 2) (x_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) x)) (= (x_count 3) (+ (x_count 2) 1)) (= (x_count 3) (x_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) x)) (= (x_count 4) (+ (x_count 3) 1)) (= (x_count 4) (x_count 3)))) (ite (and (= (select_format 5) percent) (= (select_format 6) x)) (= (x_count 5) (+ (x_count 4) 1)) (= (x_count 5) (x_count 4)))))) :rule and :premises (t50) :args (0))
(step t118 (cl (and (and (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255)) (= (select_format 5) percent)) (= (select_format 5) s)) (= (select_format 5) x)) (= (select_format 5) adr_lo)) (= (select_format 5) adr_medlo)) (= (select_format 5) adr_medhi)) (= (select_format 5) adr_hi)) (= (select_format 5) 255)))) :rule and :premises (t117) :args (0))
(step t119 (cl (and (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))))) :rule and :premises (t118) :args (0))
(step t120 (cl (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) :rule and :premises (t119) :args (1))
(step t121 (cl (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)))) :rule and :premises (t120) :args (0))
(step t122 (cl (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) :rule and :premises (t121) :args (1))
(step t123 (cl (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule resolution :premises (t105 t116 t122))
(step t124 (cl (and (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115)) (= x 120))) :rule and :premises (t119) :args (0))
(step t125 (cl (and (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37)) (= s 115))) :rule and :premises (t124) :args (0))
(step t126 (cl (and (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2)) (= percent 37))) :rule and :premises (t125) :args (0))
(step t127 (cl (and (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5)) (= adr_hi 2))) :rule and :premises (t126) :args (0))
(step t128 (cl (and (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1)) (= adr_medhi 5))) :rule and :premises (t127) :args (0))
(step t129 (cl (and (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3)) (= adr_medlo 1))) :rule and :premises (t128) :args (0))
(step t130 (cl (and (and (= distance 14) (= fmt_length 6)) (= adr_lo 3))) :rule and :premises (t129) :args (0))
(step t131 (cl (and (= distance 14) (= fmt_length 6))) :rule and :premises (t130) :args (0))
(step t132 (cl (= fmt_length 6)) :rule and :premises (t131) :args (1))
(step t133 (cl (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance)))) :rule and :premises (t121) :args (0))
(step t134 (cl (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1)))) :rule and :premises (t133) :args (0))
(step t135 (cl (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1)))) :rule and :premises (t134) :args (0))
(step t136 (cl (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0))) :rule and :premises (t135) :args (0))
(step t137 (cl (and (= fmt0 0) (= arg0 (- fmt0 distance)))) :rule and :premises (t136) :args (0))
(step t138 (cl (= fmt0 0)) :rule and :premises (t137) :args (0))
(step t139 (cl (not (= (= arg0 (- fmt0 distance)) (= fmt0 (+ arg0 distance)))) (not (= arg0 (- fmt0 distance))) (= fmt0 (+ arg0 distance))) :rule equiv_pos2)
(step t140 (cl (= arg0 arg0)) :rule refl)
(step t141 (cl (= (- fmt0 distance) (+ fmt0 (* -1 distance)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- fmt0 distance) (+ fmt0 (* -1 distance))) 3 6))
(step t142 (cl (= (= arg0 (- fmt0 distance)) (= arg0 (+ fmt0 (* -1 distance))))) :rule cong :premises (t140 t141))
(step t143 (cl (= (= arg0 (+ fmt0 (* -1 distance))) (= fmt0 (+ arg0 distance)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= arg0 (+ fmt0 (* -1 distance))) (= fmt0 (+ arg0 distance))) 3 7))
(step t144 (cl (= (= arg0 (- fmt0 distance)) (= fmt0 (+ arg0 distance)))) :rule trans :premises (t142 t143))
(step t145 (cl (= arg0 (- fmt0 distance))) :rule and :premises (t137) :args (1))
(step t146 (cl (= fmt0 (+ arg0 distance))) :rule resolution :premises (t139 t144 t145))
(step t147 (cl (not (= fmt0 (+ -2 arg1)))) :rule resolution :premises (t104 t123 t132 t138 t146))
(step t148 (cl (=> (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1)) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) :rule implies_neg1)
(anchor :step t149)
(assume t149.a0 (= distance 14))
(assume t149.a1 (= fmt0 0))
(assume t149.a2 (= fmt0 (+ arg0 distance)))
(assume t149.a3 (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))
(assume t149.a4 (= fmt0 arg1))
(step t149.t0 (cl (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)) (not (= distance 14)) (not (= (+ arg0 distance) 0)) (not (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (not (= arg1 0))) :rule and_neg)
(step t149.t1 (cl (= distance 14)) :rule and :premises (t131) :args (0))
(step t149.t2 (cl (= (+ arg0 distance) fmt0)) :rule symm :premises (t146))
(step t149.t3 (cl (= (+ arg0 distance) 0)) :rule trans :premises (t149.t2 t138))
(step t149.t4 (cl (not (= (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) (not (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule equiv_pos2)
(step t149.t5 (cl (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0)) 3 6))
(step t149.t6 (cl (= 4 4)) :rule refl)
(step t149.t7 (cl (= (- (- fmt1 2) fmt0) (+ (- fmt1 2) (* -1 fmt0)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- (- fmt1 2) fmt0) (+ (- fmt1 2) (* -1 fmt0))) 3 6))
(step t149.t8 (cl (= (- fmt1 2) (+ fmt1 (* -1 2)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- fmt1 2) (+ fmt1 (* -1 2))) 3 6))
(step t149.t9 (cl (= fmt1 fmt1)) :rule refl)
(step t149.t10 (cl (= (* -1 2) -2)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (* -1 2) -2) 3 7))
(step t149.t11 (cl (= (+ fmt1 (* -1 2)) (+ fmt1 -2))) :rule cong :premises (t149.t9 t149.t10))
(step t149.t12 (cl (= (+ fmt1 -2) (+ -2 fmt1))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ fmt1 -2) (+ -2 fmt1)) 3 7))
(step t149.t13 (cl (= (+ fmt1 (* -1 2)) (+ -2 fmt1))) :rule trans :premises (t149.t11 t149.t12))
(step t149.t14 (cl (= (- fmt1 2) (+ -2 fmt1))) :rule trans :premises (t149.t8 t149.t13))
(step t149.t15 (cl (= (* -1 fmt0) (* -1 fmt0))) :rule refl)
(step t149.t16 (cl (= (+ (- fmt1 2) (* -1 fmt0)) (+ (+ -2 fmt1) (* -1 fmt0)))) :rule cong :premises (t149.t14 t149.t15))
(step t149.t17 (cl (= (+ (+ -2 fmt1) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ -2 fmt1) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0))) 3 7))
(step t149.t18 (cl (= (+ (- fmt1 2) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0)))) :rule trans :premises (t149.t16 t149.t17))
(step t149.t19 (cl (= (- (- fmt1 2) fmt0) (+ -2 fmt1 (* -1 fmt0)))) :rule trans :premises (t149.t7 t149.t18))
(step t149.t20 (cl (= (x_count (- (- fmt1 2) fmt0)) (x_count (+ -2 fmt1 (* -1 fmt0))))) :rule cong :premises (t149.t19))
(step t149.t21 (cl (= (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))))) :rule cong :premises (t149.t6 t149.t20))
(step t149.t22 (cl (= (s_count (- (- fmt1 2) fmt0)) (s_count (+ -2 fmt1 (* -1 fmt0))))) :rule cong :premises (t149.t19))
(step t149.t23 (cl (= (* 4 (s_count (- (- fmt1 2) fmt0))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) :rule cong :premises (t149.t6 t149.t22))
(step t149.t24 (cl (= (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0) (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0))) :rule cong :premises (t149.t21 t149.t23 t140))
(step t149.t25 (cl (= (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 3 7))
(step t149.t26 (cl (= (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule trans :premises (t149.t24 t149.t25))
(step t149.t27 (cl (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule trans :premises (t149.t5 t149.t26))
(step t149.t28 (cl (= (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule cong :premises (t35 t149.t27))
(step t149.t29 (cl (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))) :rule and :premises (t120) :args (1))
(step t149.t30 (cl (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule resolution :premises (t149.t4 t149.t28 t149.t29))
(step t149.t31 (cl (=> (and (= fmt0 0) (= fmt0 arg1)) (= arg1 0)) (and (= fmt0 0) (= fmt0 arg1))) :rule implies_neg1)
(anchor :step t149.t32)
(assume t149.t32.a0 (= fmt0 0))
(assume t149.t32.a1 (= fmt0 arg1))
(step t149.t32.t0 (cl (= arg1 fmt0)) :rule symm :premises (t149.t32.a1))
(step t149.t32.t1 (cl (= arg1 0)) :rule trans :premises (t149.t32.t0 t138))
(step t149.t32 (cl (not (= fmt0 0)) (not (= fmt0 arg1)) (= arg1 0)) :rule subproof :discharge (t149.t32.a0 t149.t32.a1))
(step t149.t33 (cl (not (and (= fmt0 0) (= fmt0 arg1))) (= fmt0 0)) :rule and_pos :args (0))
(step t149.t34 (cl (not (and (= fmt0 0) (= fmt0 arg1))) (= fmt0 arg1)) :rule and_pos :args (1))
(step t149.t35 (cl (= arg1 0) (not (and (= fmt0 0) (= fmt0 arg1))) (not (and (= fmt0 0) (= fmt0 arg1)))) :rule resolution :premises (t149.t32 t149.t33 t149.t34))
(step t149.t36 (cl (not (and (= fmt0 0) (= fmt0 arg1))) (not (and (= fmt0 0) (= fmt0 arg1))) (= arg1 0)) :rule reordering :premises (t149.t35))
(step t149.t37 (cl (not (and (= fmt0 0) (= fmt0 arg1))) (= arg1 0)) :rule contraction :premises (t149.t36))
(step t149.t38 (cl (=> (and (= fmt0 0) (= fmt0 arg1)) (= arg1 0)) (= arg1 0)) :rule resolution :premises (t149.t31 t149.t37))
(step t149.t39 (cl (=> (and (= fmt0 0) (= fmt0 arg1)) (= arg1 0)) (not (= arg1 0))) :rule implies_neg2)
(step t149.t40 (cl (=> (and (= fmt0 0) (= fmt0 arg1)) (= arg1 0)) (=> (and (= fmt0 0) (= fmt0 arg1)) (= arg1 0))) :rule resolution :premises (t149.t38 t149.t39))
(step t149.t41 (cl (=> (and (= fmt0 0) (= fmt0 arg1)) (= arg1 0))) :rule contraction :premises (t149.t40))
(step t149.t42 (cl (not (and (= fmt0 0) (= fmt0 arg1))) (= arg1 0)) :rule implies :premises (t149.t41))
(step t149.t43 (cl (and (= fmt0 0) (= fmt0 arg1)) (not (= fmt0 0)) (not (= fmt0 arg1))) :rule and_neg)
(step t149.t44 (cl (and (= fmt0 0) (= fmt0 arg1))) :rule resolution :premises (t149.t43 t138 t149.a4))
(step t149.t45 (cl (= arg1 0)) :rule resolution :premises (t149.t42 t149.t44))
(step t149.t46 (cl (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) :rule resolution :premises (t149.t0 t149.t1 t149.t3 t149.t30 t149.t45))
(step t149 (cl (not (= distance 14)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (not (= fmt0 arg1)) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) :rule subproof :discharge (t149.a0 t149.a1 t149.a2 t149.a3 t149.a4))
(step t150 (cl (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (= distance 14)) :rule and_pos :args (0))
(step t151 (cl (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (= fmt0 0)) :rule and_pos :args (1))
(step t152 (cl (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (= fmt0 (+ arg0 distance))) :rule and_pos :args (2))
(step t153 (cl (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule and_pos :args (3))
(step t154 (cl (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (= fmt0 arg1)) :rule and_pos :args (4))
(step t155 (cl (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)) (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1)))) :rule resolution :premises (t149 t150 t151 t152 t153 t154))
(step t156 (cl (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) :rule reordering :premises (t155))
(step t157 (cl (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) :rule contraction :premises (t156))
(step t158 (cl (=> (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1)) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) :rule resolution :premises (t148 t157))
(step t159 (cl (=> (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1)) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)))) :rule implies_neg2)
(step t160 (cl (=> (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1)) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (=> (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1)) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)))) :rule resolution :premises (t158 t159))
(step t161 (cl (=> (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1)) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)))) :rule contraction :premises (t160))
(step t162 (cl (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) :rule implies :premises (t161))
(step t163 (cl (=> (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)) false) (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) :rule implies_neg1)
(anchor :step t164)
(assume t164.a0 (= distance 14))
(assume t164.a1 (= (+ arg0 distance) 0))
(assume t164.a2 (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))
(assume t164.a3 (= arg1 0))
(step t164.t0 (cl (not (= (= (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 0) false)) (not (= (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 0)) false) :rule equiv_pos2)
(step t164.t1 (cl (= (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (+ -14 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (+ -14 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 3 7))
(step t164.t2 (cl (= 0 0)) :rule refl)
(step t164.t3 (cl (= (= (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 0) (= (+ -14 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))) 0))) :rule cong :premises (t164.t1 t164.t2))
(step t164.t4 (cl (= (= (+ -14 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))) 0) false)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= (+ -14 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))) 0) false) 3 7))
(step t164.t5 (cl (= (= (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 0) false)) :rule trans :premises (t164.t3 t164.t4))
(step t164.t6 (cl (not (= (= arg1 0) (= (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 0))) (not (= arg1 0)) (= (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 0)) :rule equiv_pos2)
(step t164.t7 (cl (not (= (= arg1 (+ (* -1 14) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0))))))))) (not (= arg1 (+ (* -1 14) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (= arg1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule equiv_pos2)
(step t164.t8 (cl (= (* -1 14) -14)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (* -1 14) -14) 3 7))
(step t164.t9 (cl (= (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))))) :rule refl)
(step t164.t10 (cl (= (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) :rule refl)
(step t164.t11 (cl (= (+ (* -1 14) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))) (+ -14 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule cong :premises (t164.t8 t164.t9 t164.t10))
(step t164.t12 (cl (= (= arg1 (+ (* -1 14) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 (+ -14 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule cong :premises (t35 t164.t11))
(step t164.t13 (cl (= (= arg1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (= arg1 (+ -14 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule cong :premises (t35 t164.t1))
(step t164.t14 (cl (= (= arg1 (+ -14 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0))))))))) :rule symm :premises (t164.t13))
(step t164.t15 (cl (= (= arg1 (+ (* -1 14) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0))))))))) :rule trans :premises (t164.t12 t164.t14))
(step t164.t16 (cl (not (= (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 (+ (* -1 14) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) (not (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (= arg1 (+ (* -1 14) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule equiv_pos2)
(step t164.t17 (cl (not (= (= (+ arg0 (* -1 -14)) 0) (= arg0 (* -1 14)))) (not (= (+ arg0 (* -1 -14)) 0)) (= arg0 (* -1 14))) :rule equiv_pos2)
(step t164.t18 (cl (= (* -1 -14) 14)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (* -1 -14) 14) 3 7))
(step t164.t19 (cl (= (+ arg0 (* -1 -14)) (+ arg0 14))) :rule cong :premises (t140 t164.t18))
(step t164.t20 (cl (= (+ arg0 14) (+ 14 arg0))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ arg0 14) (+ 14 arg0)) 3 7))
(step t164.t21 (cl (= (+ arg0 (* -1 -14)) (+ 14 arg0))) :rule trans :premises (t164.t19 t164.t20))
(step t164.t22 (cl (= (= (+ arg0 (* -1 -14)) 0) (= (+ 14 arg0) 0))) :rule cong :premises (t164.t21 t164.t2))
(step t164.t23 (cl (= (= (+ 14 arg0) 0) (= arg0 -14))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= (+ 14 arg0) 0) (= arg0 -14)) 3 7))
(step t164.t24 (cl (= (= (+ arg0 (* -1 -14)) 0) (= arg0 -14))) :rule trans :premises (t164.t22 t164.t23))
(step t164.t25 (cl (= (= arg0 (* -1 14)) (= arg0 -14))) :rule cong :premises (t140 t164.t8))
(step t164.t26 (cl (= (= arg0 -14) (= arg0 (* -1 14)))) :rule symm :premises (t164.t25))
(step t164.t27 (cl (= (= (+ arg0 (* -1 -14)) 0) (= arg0 (* -1 14)))) :rule trans :premises (t164.t24 t164.t26))
(step t164.t28 (cl (not (= (= (+ arg0 distance) 0) (= (+ arg0 (* -1 -14)) 0))) (not (= (+ arg0 distance) 0)) (= (+ arg0 (* -1 -14)) 0)) :rule equiv_pos2)
(step t164.t29 (cl (not (= (= distance 14) (= distance (* -1 -14)))) (not (= distance 14)) (= distance (* -1 -14))) :rule equiv_pos2)
(step t164.t30 (cl (= distance distance)) :rule refl)
(step t164.t31 (cl (= (= distance (* -1 -14)) (= distance 14))) :rule cong :premises (t164.t30 t164.t18))
(step t164.t32 (cl (= (= distance 14) (= distance (* -1 -14)))) :rule symm :premises (t164.t31))
(step t164.t33 (cl (= distance 14)) :rule and :premises (t131) :args (0))
(step t164.t34 (cl (= distance (* -1 -14))) :rule resolution :premises (t164.t29 t164.t32 t164.t33))
(step t164.t35 (cl (= (+ arg0 distance) (+ arg0 (* -1 -14)))) :rule cong :premises (t140 t164.t34))
(step t164.t36 (cl (= (= (+ arg0 distance) 0) (= (+ arg0 (* -1 -14)) 0))) :rule cong :premises (t164.t35 t164.t2))
(step t164.t37 (cl (= (+ arg0 distance) fmt0)) :rule symm :premises (t146))
(step t164.t38 (cl (= (+ arg0 distance) 0)) :rule trans :premises (t164.t37 t138))
(step t164.t39 (cl (= (+ arg0 (* -1 -14)) 0)) :rule resolution :premises (t164.t28 t164.t36 t164.t38))
(step t164.t40 (cl (= arg0 (* -1 14))) :rule resolution :premises (t164.t17 t164.t27 t164.t39))
(step t164.t41 (cl (= (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))) (+ (* -1 14) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule cong :premises (t164.t40 t164.t9 t164.t10))
(step t164.t42 (cl (= (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 (+ (* -1 14) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule cong :premises (t35 t164.t41))
(step t164.t43 (cl (not (= (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) (not (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule equiv_pos2)
(step t164.t44 (cl (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0)) 3 6))
(step t164.t45 (cl (= 4 4)) :rule refl)
(step t164.t46 (cl (= (- (- fmt1 2) fmt0) (+ (- fmt1 2) (* -1 fmt0)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- (- fmt1 2) fmt0) (+ (- fmt1 2) (* -1 fmt0))) 3 6))
(step t164.t47 (cl (= (- fmt1 2) (+ fmt1 (* -1 2)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- fmt1 2) (+ fmt1 (* -1 2))) 3 6))
(step t164.t48 (cl (= fmt1 fmt1)) :rule refl)
(step t164.t49 (cl (= (* -1 2) -2)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (* -1 2) -2) 3 7))
(step t164.t50 (cl (= (+ fmt1 (* -1 2)) (+ fmt1 -2))) :rule cong :premises (t164.t48 t164.t49))
(step t164.t51 (cl (= (+ fmt1 -2) (+ -2 fmt1))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ fmt1 -2) (+ -2 fmt1)) 3 7))
(step t164.t52 (cl (= (+ fmt1 (* -1 2)) (+ -2 fmt1))) :rule trans :premises (t164.t50 t164.t51))
(step t164.t53 (cl (= (- fmt1 2) (+ -2 fmt1))) :rule trans :premises (t164.t47 t164.t52))
(step t164.t54 (cl (= (* -1 fmt0) (* -1 fmt0))) :rule refl)
(step t164.t55 (cl (= (+ (- fmt1 2) (* -1 fmt0)) (+ (+ -2 fmt1) (* -1 fmt0)))) :rule cong :premises (t164.t53 t164.t54))
(step t164.t56 (cl (= (+ (+ -2 fmt1) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ -2 fmt1) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0))) 3 7))
(step t164.t57 (cl (= (+ (- fmt1 2) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0)))) :rule trans :premises (t164.t55 t164.t56))
(step t164.t58 (cl (= (- (- fmt1 2) fmt0) (+ -2 fmt1 (* -1 fmt0)))) :rule trans :premises (t164.t46 t164.t57))
(step t164.t59 (cl (= (x_count (- (- fmt1 2) fmt0)) (x_count (+ -2 fmt1 (* -1 fmt0))))) :rule cong :premises (t164.t58))
(step t164.t60 (cl (= (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))))) :rule cong :premises (t164.t45 t164.t59))
(step t164.t61 (cl (= (s_count (- (- fmt1 2) fmt0)) (s_count (+ -2 fmt1 (* -1 fmt0))))) :rule cong :premises (t164.t58))
(step t164.t62 (cl (= (* 4 (s_count (- (- fmt1 2) fmt0))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) :rule cong :premises (t164.t45 t164.t61))
(step t164.t63 (cl (= (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0) (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0))) :rule cong :premises (t164.t60 t164.t62 t140))
(step t164.t64 (cl (= (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 3 7))
(step t164.t65 (cl (= (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule trans :premises (t164.t63 t164.t64))
(step t164.t66 (cl (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule trans :premises (t164.t44 t164.t65))
(step t164.t67 (cl (= (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule cong :premises (t35 t164.t66))
(step t164.t68 (cl (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))) :rule and :premises (t120) :args (1))
(step t164.t69 (cl (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule resolution :premises (t164.t43 t164.t67 t164.t68))
(step t164.t70 (cl (= arg1 (+ (* -1 14) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule resolution :premises (t164.t16 t164.t42 t164.t69))
(step t164.t71 (cl (= arg1 (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule resolution :premises (t164.t7 t164.t15 t164.t70))
(step t164.t72 (cl (= (= arg1 0) (= (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 0))) :rule cong :premises (t164.t71 t164.t2))
(step t164.t73 (cl (= (* -1 (+ 14 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 0)) :rule resolution :premises (t164.t6 t164.t72 t164.a3))
(step t164.t74 (cl false) :rule resolution :premises (t164.t0 t164.t5 t164.t73))
(step t164 (cl (not (= distance 14)) (not (= (+ arg0 distance) 0)) (not (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (not (= arg1 0)) false) :rule subproof :discharge (t164.a0 t164.a1 t164.a2 t164.a3))
(step t165 (cl (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (= distance 14)) :rule and_pos :args (0))
(step t166 (cl (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (= (+ arg0 distance) 0)) :rule and_pos :args (1))
(step t167 (cl (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule and_pos :args (2))
(step t168 (cl (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (= arg1 0)) :rule and_pos :args (3))
(step t169 (cl false (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)))) :rule resolution :premises (t164 t165 t166 t167 t168))
(step t170 (cl (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) false) :rule reordering :premises (t169))
(step t171 (cl (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) false) :rule contraction :premises (t170))
(step t172 (cl (=> (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)) false) false) :rule resolution :premises (t163 t171))
(step t173 (cl (=> (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)) false) (not false)) :rule implies_neg2)
(step t174 (cl (=> (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)) false) (=> (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)) false)) :rule resolution :premises (t172 t173))
(step t175 (cl (=> (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)) false)) :rule contraction :premises (t174))
(step t176 (cl (= (=> (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)) false) (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))))) :rule implies_simplify)
(step t177 (cl (not (=> (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)) false)) (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)))) :rule equiv1 :premises (t176))
(step t178 (cl (not (and (= distance 14) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)))) :rule resolution :premises (t175 t177))
(step t179 (cl (not (and (= distance 14) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1)))) :rule resolution :premises (t162 t178))
(step t180 (cl (not (= distance 14)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (not (= fmt0 arg1))) :rule not_and :premises (t179))
(step t181 (cl (= distance 14)) :rule and :premises (t131) :args (0))
(step t182 (cl (not (= (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) (not (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule equiv_pos2)
(step t183 (cl (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0)) 3 6))
(step t184 (cl (= 4 4)) :rule refl)
(step t185 (cl (= (- (- fmt1 2) fmt0) (+ (- fmt1 2) (* -1 fmt0)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- (- fmt1 2) fmt0) (+ (- fmt1 2) (* -1 fmt0))) 3 6))
(step t186 (cl (= (- fmt1 2) (+ fmt1 (* -1 2)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- fmt1 2) (+ fmt1 (* -1 2))) 3 6))
(step t187 (cl (= fmt1 fmt1)) :rule refl)
(step t188 (cl (= (* -1 2) -2)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (* -1 2) -2) 3 7))
(step t189 (cl (= (+ fmt1 (* -1 2)) (+ fmt1 -2))) :rule cong :premises (t187 t188))
(step t190 (cl (= (+ fmt1 -2) (+ -2 fmt1))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ fmt1 -2) (+ -2 fmt1)) 3 7))
(step t191 (cl (= (+ fmt1 (* -1 2)) (+ -2 fmt1))) :rule trans :premises (t189 t190))
(step t192 (cl (= (- fmt1 2) (+ -2 fmt1))) :rule trans :premises (t186 t191))
(step t193 (cl (= (* -1 fmt0) (* -1 fmt0))) :rule refl)
(step t194 (cl (= (+ (- fmt1 2) (* -1 fmt0)) (+ (+ -2 fmt1) (* -1 fmt0)))) :rule cong :premises (t192 t193))
(step t195 (cl (= (+ (+ -2 fmt1) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ -2 fmt1) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0))) 3 7))
(step t196 (cl (= (+ (- fmt1 2) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0)))) :rule trans :premises (t194 t195))
(step t197 (cl (= (- (- fmt1 2) fmt0) (+ -2 fmt1 (* -1 fmt0)))) :rule trans :premises (t185 t196))
(step t198 (cl (= (x_count (- (- fmt1 2) fmt0)) (x_count (+ -2 fmt1 (* -1 fmt0))))) :rule cong :premises (t197))
(step t199 (cl (= (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))))) :rule cong :premises (t184 t198))
(step t200 (cl (= (s_count (- (- fmt1 2) fmt0)) (s_count (+ -2 fmt1 (* -1 fmt0))))) :rule cong :premises (t197))
(step t201 (cl (= (* 4 (s_count (- (- fmt1 2) fmt0))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) :rule cong :premises (t184 t200))
(step t202 (cl (= (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0) (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0))) :rule cong :premises (t199 t201 t140))
(step t203 (cl (= (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 3 7))
(step t204 (cl (= (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule trans :premises (t202 t203))
(step t205 (cl (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule trans :premises (t183 t204))
(step t206 (cl (= (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule cong :premises (t35 t205))
(step t207 (cl (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))) :rule and :premises (t120) :args (1))
(step t208 (cl (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule resolution :premises (t182 t206 t207))
(step t209 (cl (not (= fmt0 arg1))) :rule resolution :premises (t180 t181 t138 t146 t208))
(step t210 (cl (= fmt0 (+ -1 arg1))) :rule resolution :premises (t54 t147 t209))
(step t211 (cl) :rule resolution :premises (t32 t210 t208 t146 t138 t181))
