(assume a0 (and (and (and (and (and (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255))) (and (and (and (and (and (and (and (and (and (ite (and (= (select_format 0) percent) (= (select_format 1) s)) (= (s_count 0) 1) (= (s_count 0) 0)) (ite (and (= (select_format 1) percent) (= (select_format 2) s)) (= (s_count 1) (+ (s_count 0) 1)) (= (s_count 1) (s_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) s)) (= (s_count 2) (+ (s_count 1) 1)) (= (s_count 2) (s_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) s)) (= (s_count 3) (+ (s_count 2) 1)) (= (s_count 3) (s_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) s)) (= (s_count 4) (+ (s_count 3) 1)) (= (s_count 4) (s_count 3)))) (ite (and (= (select_format 0) percent) (= (select_format 1) x)) (= (x_count 0) 1) (= (x_count 0) 0))) (ite (and (= (select_format 1) percent) (= (select_format 2) x)) (= (x_count 1) (+ (x_count 0) 1)) (= (x_count 1) (x_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) x)) (= (x_count 2) (+ (x_count 1) 1)) (= (x_count 2) (x_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) x)) (= (x_count 3) (+ (x_count 2) 1)) (= (x_count 3) (x_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) x)) (= (x_count 4) (+ (x_count 3) 1)) (= (x_count 4) (x_count 3))))) (and (or (or (or (or (= fmt1 (+ fmt0 0)) (= fmt1 (+ fmt0 1))) (= fmt1 (+ fmt0 2))) (= fmt1 (+ fmt0 3))) (= fmt1 (+ fmt0 4))) (or (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1))))) (not (and (and (and (and (and (= (select_format fmt1) percent) (= (select_format (+ fmt1 1)) s)) (= (select_format arg1) adr_lo)) (= (select_format (+ arg1 1)) adr_medlo)) (= (select_format (+ arg1 2)) adr_medhi)) (= (select_format (+ arg1 3)) adr_hi)))))
(step t0 (cl (not (= (or (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -1 arg1)))) (or (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) (not (= fmt0 (+ -1 arg1)))))) (not (or (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -1 arg1))))) (or (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) (not (= fmt0 (+ -1 arg1))))) :rule equiv_pos2)
(step t1 (cl (= (not (= fmt_length 5)) (not (= fmt_length 5)))) :rule refl)
(step t2 (cl (= (not (= fmt0 0)) (not (= fmt0 0)))) :rule refl)
(step t3 (cl (= (not (= fmt0 (+ arg0 distance))) (not (= fmt0 (+ arg0 distance))))) :rule refl)
(step t4 (cl (= (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) 1 6))
(step t5 (cl (= (not (= fmt0 (+ -1 arg1))) (not (= fmt0 (+ -1 arg1))))) :rule refl)
(step t6 (cl (= (or (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -1 arg1)))) (or (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) (not (= fmt0 (+ -1 arg1)))))) :rule cong :premises (t1 t2 t3 t4 t5))
(step t7 (cl (=> (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1))) (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1)))) :rule implies_neg1)
(anchor :step t8)
(assume t8.a0 (= fmt_length 5))
(assume t8.a1 (= fmt0 0))
(assume t8.a2 (= fmt0 (+ arg0 distance)))
(assume t8.a3 (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))
(assume t8.a4 (= fmt0 (+ -1 arg1)))
(step t8.t0 (cl (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt_length 5)) (not (= 0 (+ -1 arg1))) (not (= (+ arg0 distance) 0)) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule and_neg)
(step t8.t1 (cl (and (and (and (and (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255))) (and (and (and (and (and (and (and (and (and (ite (and (= (select_format 0) percent) (= (select_format 1) s)) (= (s_count 0) 1) (= (s_count 0) 0)) (ite (and (= (select_format 1) percent) (= (select_format 2) s)) (= (s_count 1) (+ (s_count 0) 1)) (= (s_count 1) (s_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) s)) (= (s_count 2) (+ (s_count 1) 1)) (= (s_count 2) (s_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) s)) (= (s_count 3) (+ (s_count 2) 1)) (= (s_count 3) (s_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) s)) (= (s_count 4) (+ (s_count 3) 1)) (= (s_count 4) (s_count 3)))) (ite (and (= (select_format 0) percent) (= (select_format 1) x)) (= (x_count 0) 1) (= (x_count 0) 0))) (ite (and (= (select_format 1) percent) (= (select_format 2) x)) (= (x_count 1) (+ (x_count 0) 1)) (= (x_count 1) (x_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) x)) (= (x_count 2) (+ (x_count 1) 1)) (= (x_count 2) (x_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) x)) (= (x_count 3) (+ (x_count 2) 1)) (= (x_count 3) (x_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) x)) (= (x_count 4) (+ (x_count 3) 1)) (= (x_count 4) (x_count 3))))) (and (or (or (or (or (= fmt1 (+ fmt0 0)) (= fmt1 (+ fmt0 1))) (= fmt1 (+ fmt0 2))) (= fmt1 (+ fmt0 3))) (= fmt1 (+ fmt0 4))) (or (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1)))))) :rule and :premises (a0) :args (0))
(step t8.t2 (cl (and (and (and (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255))) (and (and (and (and (and (and (and (and (and (ite (and (= (select_format 0) percent) (= (select_format 1) s)) (= (s_count 0) 1) (= (s_count 0) 0)) (ite (and (= (select_format 1) percent) (= (select_format 2) s)) (= (s_count 1) (+ (s_count 0) 1)) (= (s_count 1) (s_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) s)) (= (s_count 2) (+ (s_count 1) 1)) (= (s_count 2) (s_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) s)) (= (s_count 3) (+ (s_count 2) 1)) (= (s_count 3) (s_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) s)) (= (s_count 4) (+ (s_count 3) 1)) (= (s_count 4) (s_count 3)))) (ite (and (= (select_format 0) percent) (= (select_format 1) x)) (= (x_count 0) 1) (= (x_count 0) 0))) (ite (and (= (select_format 1) percent) (= (select_format 2) x)) (= (x_count 1) (+ (x_count 0) 1)) (= (x_count 1) (x_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) x)) (= (x_count 2) (+ (x_count 1) 1)) (= (x_count 2) (x_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) x)) (= (x_count 3) (+ (x_count 2) 1)) (= (x_count 3) (x_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) x)) (= (x_count 4) (+ (x_count 3) 1)) (= (x_count 4) (x_count 3)))))) :rule and :premises (t8.t1) :args (0))
(step t8.t3 (cl (and (and (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255)))) :rule and :premises (t8.t2) :args (0))
(step t8.t4 (cl (and (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))))) :rule and :premises (t8.t3) :args (0))
(step t8.t5 (cl (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120))) :rule and :premises (t8.t4) :args (0))
(step t8.t6 (cl (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115))) :rule and :premises (t8.t5) :args (0))
(step t8.t7 (cl (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37))) :rule and :premises (t8.t6) :args (0))
(step t8.t8 (cl (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1))) :rule and :premises (t8.t7) :args (0))
(step t8.t9 (cl (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5))) :rule and :premises (t8.t8) :args (0))
(step t8.t10 (cl (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2))) :rule and :premises (t8.t9) :args (0))
(step t8.t11 (cl (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2))) :rule and :premises (t8.t10) :args (0))
(step t8.t12 (cl (and (= distance 6) (= fmt_length 5))) :rule and :premises (t8.t11) :args (0))
(step t8.t13 (cl (= fmt_length 5)) :rule and :premises (t8.t12) :args (1))
(step t8.t14 (cl (=> (and (= fmt0 (+ -1 arg1)) (= fmt0 0)) (= 0 (+ -1 arg1))) (and (= fmt0 (+ -1 arg1)) (= fmt0 0))) :rule implies_neg1)
(anchor :step t8.t15)
(assume t8.t15.a0 (= fmt0 (+ -1 arg1)))
(assume t8.t15.a1 (= fmt0 0))
(step t8.t15.t0 (cl (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) :rule and :premises (t8.t4) :args (1))
(step t8.t15.t1 (cl (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)))) :rule and :premises (t8.t15.t0) :args (0))
(step t8.t15.t2 (cl (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance)))) :rule and :premises (t8.t15.t1) :args (0))
(step t8.t15.t3 (cl (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1)))) :rule and :premises (t8.t15.t2) :args (0))
(step t8.t15.t4 (cl (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1)))) :rule and :premises (t8.t15.t3) :args (0))
(step t8.t15.t5 (cl (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0))) :rule and :premises (t8.t15.t4) :args (0))
(step t8.t15.t6 (cl (and (= fmt0 0) (= arg0 (- fmt0 distance)))) :rule and :premises (t8.t15.t5) :args (0))
(step t8.t15.t7 (cl (= fmt0 0)) :rule and :premises (t8.t15.t6) :args (0))
(step t8.t15.t8 (cl (= 0 fmt0)) :rule symm :premises (t8.t15.t7))
(step t8.t15.t9 (cl (= (+ -1 arg1) fmt0)) :rule symm :premises (t8.t15.a0))
(step t8.t15.t10 (cl (= fmt0 (+ -1 arg1))) :rule symm :premises (t8.t15.t9))
(step t8.t15.t11 (cl (= 0 (+ -1 arg1))) :rule trans :premises (t8.t15.t8 t8.t15.t10))
(step t8.t15 (cl (not (= fmt0 (+ -1 arg1))) (not (= fmt0 0)) (= 0 (+ -1 arg1))) :rule subproof :discharge (t8.t15.a0 t8.t15.a1))
(step t8.t16 (cl (not (and (= fmt0 (+ -1 arg1)) (= fmt0 0))) (= fmt0 (+ -1 arg1))) :rule and_pos :args (0))
(step t8.t17 (cl (not (and (= fmt0 (+ -1 arg1)) (= fmt0 0))) (= fmt0 0)) :rule and_pos :args (1))
(step t8.t18 (cl (= 0 (+ -1 arg1)) (not (and (= fmt0 (+ -1 arg1)) (= fmt0 0))) (not (and (= fmt0 (+ -1 arg1)) (= fmt0 0)))) :rule resolution :premises (t8.t15 t8.t16 t8.t17))
(step t8.t19 (cl (not (and (= fmt0 (+ -1 arg1)) (= fmt0 0))) (not (and (= fmt0 (+ -1 arg1)) (= fmt0 0))) (= 0 (+ -1 arg1))) :rule reordering :premises (t8.t18))
(step t8.t20 (cl (not (and (= fmt0 (+ -1 arg1)) (= fmt0 0))) (= 0 (+ -1 arg1))) :rule contraction :premises (t8.t19))
(step t8.t21 (cl (=> (and (= fmt0 (+ -1 arg1)) (= fmt0 0)) (= 0 (+ -1 arg1))) (= 0 (+ -1 arg1))) :rule resolution :premises (t8.t14 t8.t20))
(step t8.t22 (cl (=> (and (= fmt0 (+ -1 arg1)) (= fmt0 0)) (= 0 (+ -1 arg1))) (not (= 0 (+ -1 arg1)))) :rule implies_neg2)
(step t8.t23 (cl (=> (and (= fmt0 (+ -1 arg1)) (= fmt0 0)) (= 0 (+ -1 arg1))) (=> (and (= fmt0 (+ -1 arg1)) (= fmt0 0)) (= 0 (+ -1 arg1)))) :rule resolution :premises (t8.t21 t8.t22))
(step t8.t24 (cl (=> (and (= fmt0 (+ -1 arg1)) (= fmt0 0)) (= 0 (+ -1 arg1)))) :rule contraction :premises (t8.t23))
(step t8.t25 (cl (not (and (= fmt0 (+ -1 arg1)) (= fmt0 0))) (= 0 (+ -1 arg1))) :rule implies :premises (t8.t24))
(step t8.t26 (cl (and (= fmt0 (+ -1 arg1)) (= fmt0 0)) (not (= fmt0 (+ -1 arg1))) (not (= fmt0 0))) :rule and_neg)
(step t8.t27 (cl (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) :rule and :premises (t8.t4) :args (1))
(step t8.t28 (cl (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)))) :rule and :premises (t8.t27) :args (0))
(step t8.t29 (cl (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance)))) :rule and :premises (t8.t28) :args (0))
(step t8.t30 (cl (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1)))) :rule and :premises (t8.t29) :args (0))
(step t8.t31 (cl (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1)))) :rule and :premises (t8.t30) :args (0))
(step t8.t32 (cl (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0))) :rule and :premises (t8.t31) :args (0))
(step t8.t33 (cl (and (= fmt0 0) (= arg0 (- fmt0 distance)))) :rule and :premises (t8.t32) :args (0))
(step t8.t34 (cl (= fmt0 0)) :rule and :premises (t8.t33) :args (0))
(step t8.t35 (cl (and (= fmt0 (+ -1 arg1)) (= fmt0 0))) :rule resolution :premises (t8.t26 t8.a4 t8.t34))
(step t8.t36 (cl (= 0 (+ -1 arg1))) :rule resolution :premises (t8.t25 t8.t35))
(step t8.t37 (cl (not (= (= arg0 (- fmt0 distance)) (= fmt0 (+ arg0 distance)))) (not (= arg0 (- fmt0 distance))) (= fmt0 (+ arg0 distance))) :rule equiv_pos2)
(step t8.t38 (cl (= arg0 arg0)) :rule refl)
(step t8.t39 (cl (= (- fmt0 distance) (+ fmt0 (* -1 distance)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- fmt0 distance) (+ fmt0 (* -1 distance))) 3 6))
(step t8.t40 (cl (= (= arg0 (- fmt0 distance)) (= arg0 (+ fmt0 (* -1 distance))))) :rule cong :premises (t8.t38 t8.t39))
(step t8.t41 (cl (= (= arg0 (+ fmt0 (* -1 distance))) (= fmt0 (+ arg0 distance)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= arg0 (+ fmt0 (* -1 distance))) (= fmt0 (+ arg0 distance))) 3 7))
(step t8.t42 (cl (= (= arg0 (- fmt0 distance)) (= fmt0 (+ arg0 distance)))) :rule trans :premises (t8.t40 t8.t41))
(step t8.t43 (cl (= arg0 (- fmt0 distance))) :rule and :premises (t8.t33) :args (1))
(step t8.t44 (cl (= fmt0 (+ arg0 distance))) :rule resolution :premises (t8.t37 t8.t42 t8.t43))
(step t8.t45 (cl (= (+ arg0 distance) fmt0)) :rule symm :premises (t8.t44))
(step t8.t46 (cl (= (+ arg0 distance) 0)) :rule trans :premises (t8.t45 t8.t34))
(step t8.t47 (cl (not (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule equiv_pos2)
(step t8.t48 (cl (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4)))) 3 6))
(step t8.t49 (cl (= arg1 arg1)) :rule refl)
(step t8.t50 (cl (= (- (+ (+ arg0 distance) fmt_length) 4) (+ (+ (+ arg0 distance) fmt_length) (* -1 4)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- (+ (+ arg0 distance) fmt_length) 4) (+ (+ (+ arg0 distance) fmt_length) (* -1 4))) 3 6))
(step t8.t51 (cl (= (+ (+ (+ arg0 distance) fmt_length) (* -1 4)) (+ distance fmt_length arg0 -4))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ (+ arg0 distance) fmt_length) (* -1 4)) (+ distance fmt_length arg0 -4)) 3 6))
(step t8.t52 (cl (= (+ distance fmt_length arg0 -4) (+ -4 arg0 fmt_length distance))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ distance fmt_length arg0 -4) (+ -4 arg0 fmt_length distance)) 3 7))
(step t8.t53 (cl (= (+ (+ (+ arg0 distance) fmt_length) (* -1 4)) (+ -4 arg0 fmt_length distance))) :rule trans :premises (t8.t51 t8.t52))
(step t8.t54 (cl (= (- (+ (+ arg0 distance) fmt_length) 4) (+ -4 arg0 fmt_length distance))) :rule trans :premises (t8.t50 t8.t53))
(step t8.t55 (cl (= (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (>= arg1 (+ -4 arg0 fmt_length distance)))) :rule cong :premises (t8.t49 t8.t54))
(step t8.t56 (cl (= (>= arg1 (+ -4 arg0 fmt_length distance)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (>= arg1 (+ -4 arg0 fmt_length distance)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) 3 7))
(step t8.t57 (cl (= (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule trans :premises (t8.t55 t8.t56))
(step t8.t58 (cl (= (not (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule cong :premises (t8.t57))
(step t8.t59 (cl (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule trans :premises (t8.t48 t8.t58))
(step t8.t60 (cl (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) :rule and :premises (t8.t28) :args (1))
(step t8.t61 (cl (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule resolution :premises (t8.t47 t8.t59 t8.t60))
(step t8.t62 (cl (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule resolution :premises (t8.t0 t8.t13 t8.t36 t8.t46 t8.t61))
(step t8 (cl (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -1 arg1))) (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule subproof :discharge (t8.a0 t8.a1 t8.a2 t8.a3 t8.a4))
(step t9 (cl (not (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1)))) (= fmt_length 5)) :rule and_pos :args (0))
(step t10 (cl (not (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1)))) (= fmt0 0)) :rule and_pos :args (1))
(step t11 (cl (not (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1)))) (= fmt0 (+ arg0 distance))) :rule and_pos :args (2))
(step t12 (cl (not (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1)))) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule and_pos :args (3))
(step t13 (cl (not (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1)))) (= fmt0 (+ -1 arg1))) :rule and_pos :args (4))
(step t14 (cl (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1)))) (not (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1)))) (not (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1)))) (not (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1)))) (not (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1))))) :rule resolution :premises (t8 t9 t10 t11 t12 t13))
(step t15 (cl (not (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1)))) (not (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1)))) (not (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1)))) (not (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1)))) (not (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1)))) (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule reordering :premises (t14))
(step t16 (cl (not (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1)))) (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule contraction :premises (t15))
(step t17 (cl (=> (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1))) (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule resolution :premises (t7 t16))
(step t18 (cl (=> (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1))) (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))))) :rule implies_neg2)
(step t19 (cl (=> (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1))) (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (=> (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1))) (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))))) :rule resolution :premises (t17 t18))
(step t20 (cl (=> (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1))) (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))))) :rule contraction :premises (t19))
(step t21 (cl (not (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1)))) (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule implies :premises (t20))
(step t22 (cl (=> (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) false) (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule implies_neg1)
(anchor :step t23)
(assume t23.a0 (= fmt_length 5))
(assume t23.a1 (= 0 (+ -1 arg1)))
(assume t23.a2 (= (+ arg0 distance) 0))
(assume t23.a3 (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))
(step t23.t0 (cl (not (= (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) :rule equiv_pos2)
(step t23.t1 (cl (= (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) 3 6))
(step t23.t2 (cl (= (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule symm :premises (t23.t1))
(step t23.t3 (cl (not (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule equiv_pos2)
(step t23.t4 (cl (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4)))) 3 6))
(step t23.t5 (cl (= arg1 arg1)) :rule refl)
(step t23.t6 (cl (= (- (+ (+ arg0 distance) fmt_length) 4) (+ (+ (+ arg0 distance) fmt_length) (* -1 4)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- (+ (+ arg0 distance) fmt_length) 4) (+ (+ (+ arg0 distance) fmt_length) (* -1 4))) 3 6))
(step t23.t7 (cl (= (+ (+ (+ arg0 distance) fmt_length) (* -1 4)) (+ distance fmt_length arg0 -4))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ (+ arg0 distance) fmt_length) (* -1 4)) (+ distance fmt_length arg0 -4)) 3 6))
(step t23.t8 (cl (= (+ distance fmt_length arg0 -4) (+ -4 arg0 fmt_length distance))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ distance fmt_length arg0 -4) (+ -4 arg0 fmt_length distance)) 3 7))
(step t23.t9 (cl (= (+ (+ (+ arg0 distance) fmt_length) (* -1 4)) (+ -4 arg0 fmt_length distance))) :rule trans :premises (t23.t7 t23.t8))
(step t23.t10 (cl (= (- (+ (+ arg0 distance) fmt_length) 4) (+ -4 arg0 fmt_length distance))) :rule trans :premises (t23.t6 t23.t9))
(step t23.t11 (cl (= (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (>= arg1 (+ -4 arg0 fmt_length distance)))) :rule cong :premises (t23.t5 t23.t10))
(step t23.t12 (cl (= (>= arg1 (+ -4 arg0 fmt_length distance)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (>= arg1 (+ -4 arg0 fmt_length distance)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) 3 7))
(step t23.t13 (cl (= (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule trans :premises (t23.t11 t23.t12))
(step t23.t14 (cl (= (not (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule cong :premises (t23.t13))
(step t23.t15 (cl (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule trans :premises (t23.t4 t23.t14))
(step t23.t16 (cl (and (and (and (and (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255))) (and (and (and (and (and (and (and (and (and (ite (and (= (select_format 0) percent) (= (select_format 1) s)) (= (s_count 0) 1) (= (s_count 0) 0)) (ite (and (= (select_format 1) percent) (= (select_format 2) s)) (= (s_count 1) (+ (s_count 0) 1)) (= (s_count 1) (s_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) s)) (= (s_count 2) (+ (s_count 1) 1)) (= (s_count 2) (s_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) s)) (= (s_count 3) (+ (s_count 2) 1)) (= (s_count 3) (s_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) s)) (= (s_count 4) (+ (s_count 3) 1)) (= (s_count 4) (s_count 3)))) (ite (and (= (select_format 0) percent) (= (select_format 1) x)) (= (x_count 0) 1) (= (x_count 0) 0))) (ite (and (= (select_format 1) percent) (= (select_format 2) x)) (= (x_count 1) (+ (x_count 0) 1)) (= (x_count 1) (x_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) x)) (= (x_count 2) (+ (x_count 1) 1)) (= (x_count 2) (x_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) x)) (= (x_count 3) (+ (x_count 2) 1)) (= (x_count 3) (x_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) x)) (= (x_count 4) (+ (x_count 3) 1)) (= (x_count 4) (x_count 3))))) (and (or (or (or (or (= fmt1 (+ fmt0 0)) (= fmt1 (+ fmt0 1))) (= fmt1 (+ fmt0 2))) (= fmt1 (+ fmt0 3))) (= fmt1 (+ fmt0 4))) (or (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1)))))) :rule and :premises (a0) :args (0))
(step t23.t17 (cl (and (and (and (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255))) (and (and (and (and (and (and (and (and (and (ite (and (= (select_format 0) percent) (= (select_format 1) s)) (= (s_count 0) 1) (= (s_count 0) 0)) (ite (and (= (select_format 1) percent) (= (select_format 2) s)) (= (s_count 1) (+ (s_count 0) 1)) (= (s_count 1) (s_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) s)) (= (s_count 2) (+ (s_count 1) 1)) (= (s_count 2) (s_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) s)) (= (s_count 3) (+ (s_count 2) 1)) (= (s_count 3) (s_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) s)) (= (s_count 4) (+ (s_count 3) 1)) (= (s_count 4) (s_count 3)))) (ite (and (= (select_format 0) percent) (= (select_format 1) x)) (= (x_count 0) 1) (= (x_count 0) 0))) (ite (and (= (select_format 1) percent) (= (select_format 2) x)) (= (x_count 1) (+ (x_count 0) 1)) (= (x_count 1) (x_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) x)) (= (x_count 2) (+ (x_count 1) 1)) (= (x_count 2) (x_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) x)) (= (x_count 3) (+ (x_count 2) 1)) (= (x_count 3) (x_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) x)) (= (x_count 4) (+ (x_count 3) 1)) (= (x_count 4) (x_count 3)))))) :rule and :premises (t23.t16) :args (0))
(step t23.t18 (cl (and (and (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255)))) :rule and :premises (t23.t17) :args (0))
(step t23.t19 (cl (and (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))))) :rule and :premises (t23.t18) :args (0))
(step t23.t20 (cl (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) :rule and :premises (t23.t19) :args (1))
(step t23.t21 (cl (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)))) :rule and :premises (t23.t20) :args (0))
(step t23.t22 (cl (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) :rule and :premises (t23.t21) :args (1))
(step t23.t23 (cl (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule resolution :premises (t23.t3 t23.t15 t23.t22))
(step t23.t24 (cl (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) :rule resolution :premises (t23.t0 t23.t2 t23.t23))
(step t23.t25 (cl (not (= (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule equiv_pos2)
(step t23.t26 (cl (= (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))))) :rule cong :premises (t23.t1))
(step t23.t27 (cl (= (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule trans :premises (t23.t26 t4))
(step t23.t28 (cl (= (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule symm :premises (t23.t27))
(step t23.t29 (cl (not (= (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) :rule equiv_pos2)
(step t23.t30 (cl (=> (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) false) (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) :rule implies_neg1)
(anchor :step t23.t31)
(assume t23.t31.a0 (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))
(step t23.t31.t0 (cl (not (= (< (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 1) 5)) false)) (not (< (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 1) 5))) false) :rule equiv_pos2)
(step t23.t31.t1 (cl (= (< (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 1) 5)) (not (>= (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 1) 5))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (< (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 1) 5)) (not (>= (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 1) 5)))) 3 6))
(step t23.t31.t2 (cl (= (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ (* -1 arg0) (* -1 fmt_length) (* -1 arg1) (* -1 distance) distance fmt_length arg0 arg1))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ (* -1 arg0) (* -1 fmt_length) (* -1 arg1) (* -1 distance) distance fmt_length arg0 arg1)) 3 6))
(step t23.t31.t3 (cl (= (+ (* -1 arg0) (* -1 fmt_length) (* -1 arg1) (* -1 distance) distance fmt_length arg0 arg1) 0)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (* -1 arg0) (* -1 fmt_length) (* -1 arg1) (* -1 distance) distance fmt_length arg0 arg1) 0) 3 7))
(step t23.t31.t4 (cl (= (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) 0)) :rule trans :premises (t23.t31.t2 t23.t31.t3))
(step t23.t31.t5 (cl (= -4 -4)) :rule refl)
(step t23.t31.t6 (cl (= 0 0)) :rule refl)
(step t23.t31.t7 (cl (= (* -1 1) -1)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (* -1 1) -1) 3 7))
(step t23.t31.t8 (cl (= 5 5)) :rule refl)
(step t23.t31.t9 (cl (= (+ -4 0 (* -1 1) 5) (+ -4 0 -1 5))) :rule cong :premises (t23.t31.t5 t23.t31.t6 t23.t31.t7 t23.t31.t8))
(step t23.t31.t10 (cl (= (+ -4 0 -1 5) 0)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ -4 0 -1 5) 0) 3 7))
(step t23.t31.t11 (cl (= (+ -4 0 (* -1 1) 5) 0)) :rule trans :premises (t23.t31.t9 t23.t31.t10))
(step t23.t31.t12 (cl (= (>= (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 1) 5)) (>= 0 0))) :rule cong :premises (t23.t31.t4 t23.t31.t11))
(step t23.t31.t13 (cl (= (>= 0 0) true)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (>= 0 0) true) 3 7))
(step t23.t31.t14 (cl (= (>= (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 1) 5)) true)) :rule trans :premises (t23.t31.t12 t23.t31.t13))
(step t23.t31.t15 (cl (= (not (>= (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 1) 5))) (not true))) :rule cong :premises (t23.t31.t14))
(step t23.t31.t16 (cl (= (not true) false)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (not true) false) 1 7))
(step t23.t31.t17 (cl (= (not (>= (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 1) 5))) false)) :rule trans :premises (t23.t31.t15 t23.t31.t16))
(step t23.t31.t18 (cl (= (< (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 1) 5)) false)) :rule trans :premises (t23.t31.t1 t23.t31.t17))
(step t23.t31.t19 (cl (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (not (= (+ arg0 distance) 0)) (not (= (* -1 arg1) (* -1 1))) (not (= fmt_length 5)) (< (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 1) 5))) :rule la_generic :args (1/1 -1/1 -1/1 -1/1 1/1))
(step t23.t31.t20 (cl (not (= (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) :rule equiv_pos2)
(step t23.t31.t21 (cl (not (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule equiv_pos2)
(step t23.t31.t22 (cl (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule resolution :premises (t23.t31.t21 t23.t15 t23.t22))
(step t23.t31.t23 (cl (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) :rule resolution :premises (t23.t31.t20 t23.t2 t23.t31.t22))
(step t23.t31.t24 (cl (not (= (= arg0 (- fmt0 distance)) (= fmt0 (+ arg0 distance)))) (not (= arg0 (- fmt0 distance))) (= fmt0 (+ arg0 distance))) :rule equiv_pos2)
(step t23.t31.t25 (cl (= arg0 arg0)) :rule refl)
(step t23.t31.t26 (cl (= (- fmt0 distance) (+ fmt0 (* -1 distance)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- fmt0 distance) (+ fmt0 (* -1 distance))) 3 6))
(step t23.t31.t27 (cl (= (= arg0 (- fmt0 distance)) (= arg0 (+ fmt0 (* -1 distance))))) :rule cong :premises (t23.t31.t25 t23.t31.t26))
(step t23.t31.t28 (cl (= (= arg0 (+ fmt0 (* -1 distance))) (= fmt0 (+ arg0 distance)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= arg0 (+ fmt0 (* -1 distance))) (= fmt0 (+ arg0 distance))) 3 7))
(step t23.t31.t29 (cl (= (= arg0 (- fmt0 distance)) (= fmt0 (+ arg0 distance)))) :rule trans :premises (t23.t31.t27 t23.t31.t28))
(step t23.t31.t30 (cl (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance)))) :rule and :premises (t23.t21) :args (0))
(step t23.t31.t31 (cl (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1)))) :rule and :premises (t23.t31.t30) :args (0))
(step t23.t31.t32 (cl (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1)))) :rule and :premises (t23.t31.t31) :args (0))
(step t23.t31.t33 (cl (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0))) :rule and :premises (t23.t31.t32) :args (0))
(step t23.t31.t34 (cl (and (= fmt0 0) (= arg0 (- fmt0 distance)))) :rule and :premises (t23.t31.t33) :args (0))
(step t23.t31.t35 (cl (= arg0 (- fmt0 distance))) :rule and :premises (t23.t31.t34) :args (1))
(step t23.t31.t36 (cl (= fmt0 (+ arg0 distance))) :rule resolution :premises (t23.t31.t24 t23.t31.t29 t23.t31.t35))
(step t23.t31.t37 (cl (= (+ arg0 distance) fmt0)) :rule symm :premises (t23.t31.t36))
(step t23.t31.t38 (cl (= fmt0 0)) :rule and :premises (t23.t31.t34) :args (0))
(step t23.t31.t39 (cl (= (+ arg0 distance) 0)) :rule trans :premises (t23.t31.t37 t23.t31.t38))
(step t23.t31.t40 (cl (=> (and (< -1 0) (= arg1 1)) (= (* -1 arg1) (* -1 1)))) :rule la_mult_neg)
(step t23.t31.t41 (cl (not (and (< -1 0) (= arg1 1))) (= (* -1 arg1) (* -1 1))) :rule implies :premises (t23.t31.t40))
(step t23.t31.t42 (cl (and (< -1 0) (= arg1 1)) (not (< -1 0)) (not (= arg1 1))) :rule and_neg)
(step t23.t31.t43 (cl (= (= (< -1 0) true) (< -1 0))) :rule equiv_simplify)
(step t23.t31.t44 (cl (not (= (< -1 0) true)) (< -1 0)) :rule equiv1 :premises (t23.t31.t43))
(step t23.t31.t45 (cl (= (< -1 0) true)) :rule hole :args ("evaluate"))
(step t23.t31.t46 (cl (< -1 0)) :rule resolution :premises (t23.t31.t44 t23.t31.t45))
(step t23.t31.t47 (cl (not (= (= 0 (+ -1 arg1)) (= arg1 1))) (not (= 0 (+ -1 arg1))) (= arg1 1)) :rule equiv_pos2)
(step t23.t31.t48 (cl (= (= 0 (+ -1 arg1)) (= arg1 1))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= 0 (+ -1 arg1)) (= arg1 1)) 3 7))
(step t23.t31.t49 (cl (= arg1 1)) :rule resolution :premises (t23.t31.t47 t23.t31.t48 t23.a1))
(step t23.t31.t50 (cl (and (< -1 0) (= arg1 1))) :rule resolution :premises (t23.t31.t42 t23.t31.t46 t23.t31.t49))
(step t23.t31.t51 (cl (= (* -1 arg1) (* -1 1))) :rule resolution :premises (t23.t31.t41 t23.t31.t50))
(step t23.t31.t52 (cl (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120))) :rule and :premises (t23.t19) :args (0))
(step t23.t31.t53 (cl (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115))) :rule and :premises (t23.t31.t52) :args (0))
(step t23.t31.t54 (cl (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37))) :rule and :premises (t23.t31.t53) :args (0))
(step t23.t31.t55 (cl (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1))) :rule and :premises (t23.t31.t54) :args (0))
(step t23.t31.t56 (cl (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5))) :rule and :premises (t23.t31.t55) :args (0))
(step t23.t31.t57 (cl (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2))) :rule and :premises (t23.t31.t56) :args (0))
(step t23.t31.t58 (cl (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2))) :rule and :premises (t23.t31.t57) :args (0))
(step t23.t31.t59 (cl (and (= distance 6) (= fmt_length 5))) :rule and :premises (t23.t31.t58) :args (0))
(step t23.t31.t60 (cl (= fmt_length 5)) :rule and :premises (t23.t31.t59) :args (1))
(step t23.t31.t61 (cl (< (+ (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) (+ arg0 distance) (* -1 arg1) fmt_length) (+ -4 0 (* -1 1) 5))) :rule resolution :premises (t23.t31.t19 t23.t31.t23 t23.t31.t39 t23.t31.t51 t23.t31.t60))
(step t23.t31.t62 (cl false) :rule resolution :premises (t23.t31.t0 t23.t31.t18 t23.t31.t61))
(step t23.t31 (cl (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) false) :rule subproof :discharge (t23.t31.a0))
(step t23.t32 (cl (=> (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) false) false) :rule resolution :premises (t23.t30 t23.t31))
(step t23.t33 (cl (=> (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) false) (not false)) :rule implies_neg2)
(step t23.t34 (cl (=> (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) false) (=> (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) false)) :rule resolution :premises (t23.t32 t23.t33))
(step t23.t35 (cl (=> (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) false)) :rule contraction :premises (t23.t34))
(step t23.t36 (cl (= (=> (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) false) (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule implies_simplify)
(step t23.t37 (cl (not (=> (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) false)) (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule equiv1 :premises (t23.t36))
(step t23.t38 (cl (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule resolution :premises (t23.t35 t23.t37))
(step t23.t39 (cl (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) :rule resolution :premises (t23.t29 t23.t27 t23.t38))
(step t23.t40 (cl (not (< (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule resolution :premises (t23.t25 t23.t28 t23.t39))
(step t23.t41 (cl) :rule resolution :premises (t23.t24 t23.t40))
(step t23 (cl (not (= fmt_length 5)) (not (= 0 (+ -1 arg1))) (not (= (+ arg0 distance) 0)) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) false) :rule subproof :discharge (t23.a0 t23.a1 t23.a2 t23.a3))
(step t24 (cl (not (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (= fmt_length 5)) :rule and_pos :args (0))
(step t25 (cl (not (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (= 0 (+ -1 arg1))) :rule and_pos :args (1))
(step t26 (cl (not (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (= (+ arg0 distance) 0)) :rule and_pos :args (2))
(step t27 (cl (not (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule and_pos :args (3))
(step t28 (cl false (not (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))))) :rule resolution :premises (t23 t24 t25 t26 t27))
(step t29 (cl (not (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) false) :rule reordering :premises (t28))
(step t30 (cl (not (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) false) :rule contraction :premises (t29))
(step t31 (cl (=> (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) false) false) :rule resolution :premises (t22 t30))
(step t32 (cl (=> (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) false) (not false)) :rule implies_neg2)
(step t33 (cl (=> (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) false) (=> (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) false)) :rule resolution :premises (t31 t32))
(step t34 (cl (=> (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) false)) :rule contraction :premises (t33))
(step t35 (cl (= (=> (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) false) (not (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))))) :rule implies_simplify)
(step t36 (cl (not (=> (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) false)) (not (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))))) :rule equiv1 :premises (t35))
(step t37 (cl (not (and (= fmt_length 5) (= 0 (+ -1 arg1)) (= (+ arg0 distance) 0) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))))) :rule resolution :premises (t34 t36))
(step t38 (cl (not (and (= fmt_length 5) (= fmt0 0) (= fmt0 (+ arg0 distance)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) (= fmt0 (+ -1 arg1))))) :rule resolution :premises (t21 t37))
(step t39 (cl (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -1 arg1)))) :rule not_and :premises (t38))
(step t40 (cl (or (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -1 arg1)))) (not (not (= fmt_length 5)))) :rule or_neg :args (0))
(step t41 (cl (or (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -1 arg1)))) (not (not (= fmt0 0)))) :rule or_neg :args (1))
(step t42 (cl (or (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -1 arg1)))) (not (not (= fmt0 (+ arg0 distance))))) :rule or_neg :args (2))
(step t43 (cl (or (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -1 arg1)))) (not (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))))) :rule or_neg :args (3))
(step t44 (cl (or (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -1 arg1)))) (not (not (= fmt0 (+ -1 arg1))))) :rule or_neg :args (4))
(step t45 (cl (or (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -1 arg1)))) (or (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -1 arg1)))) (or (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -1 arg1)))) (or (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -1 arg1)))) (or (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -1 arg1))))) :rule resolution :premises (t39 t40 t41 t42 t43 t44))
(step t46 (cl (or (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) (not (= fmt0 (+ -1 arg1))))) :rule contraction :premises (t45))
(step t47 (cl (or (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) (not (= fmt0 (+ -1 arg1))))) :rule resolution :premises (t0 t6 t46))
(step t48 (cl (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) (not (= fmt0 (+ -1 arg1)))) :rule or :premises (t47))
(step t49 (cl (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4) (not (= fmt_length 5)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (= fmt0 (+ -1 arg1)))) :rule reordering :premises (t48))
(step t50 (cl (not (= (or (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1))) (or (= fmt0 arg1) (= fmt0 (+ -1 arg1))))) (not (or (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1)))) (or (= fmt0 arg1) (= fmt0 (+ -1 arg1)))) :rule equiv_pos2)
(step t51 (cl (= arg1 arg1)) :rule refl)
(step t52 (cl (= (+ fmt0 0) fmt0)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ fmt0 0) fmt0) 3 7))
(step t53 (cl (= (= arg1 (+ fmt0 0)) (= arg1 fmt0))) :rule cong :premises (t51 t52))
(step t54 (cl (= (= arg1 fmt0) (= fmt0 arg1))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= arg1 fmt0) (= fmt0 arg1)) 3 7))
(step t55 (cl (= (= arg1 (+ fmt0 0)) (= fmt0 arg1))) :rule trans :premises (t53 t54))
(step t56 (cl (= (+ fmt0 1) (+ 1 fmt0))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ fmt0 1) (+ 1 fmt0)) 3 7))
(step t57 (cl (= (= arg1 (+ fmt0 1)) (= arg1 (+ 1 fmt0)))) :rule cong :premises (t51 t56))
(step t58 (cl (= (= arg1 (+ 1 fmt0)) (= fmt0 (+ -1 arg1)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= arg1 (+ 1 fmt0)) (= fmt0 (+ -1 arg1))) 3 7))
(step t59 (cl (= (= arg1 (+ fmt0 1)) (= fmt0 (+ -1 arg1)))) :rule trans :premises (t57 t58))
(step t60 (cl (= (or (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1))) (or (= fmt0 arg1) (= fmt0 (+ -1 arg1))))) :rule cong :premises (t55 t59))
(step t61 (cl (and (and (and (and (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255))) (and (and (and (and (and (and (and (and (and (ite (and (= (select_format 0) percent) (= (select_format 1) s)) (= (s_count 0) 1) (= (s_count 0) 0)) (ite (and (= (select_format 1) percent) (= (select_format 2) s)) (= (s_count 1) (+ (s_count 0) 1)) (= (s_count 1) (s_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) s)) (= (s_count 2) (+ (s_count 1) 1)) (= (s_count 2) (s_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) s)) (= (s_count 3) (+ (s_count 2) 1)) (= (s_count 3) (s_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) s)) (= (s_count 4) (+ (s_count 3) 1)) (= (s_count 4) (s_count 3)))) (ite (and (= (select_format 0) percent) (= (select_format 1) x)) (= (x_count 0) 1) (= (x_count 0) 0))) (ite (and (= (select_format 1) percent) (= (select_format 2) x)) (= (x_count 1) (+ (x_count 0) 1)) (= (x_count 1) (x_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) x)) (= (x_count 2) (+ (x_count 1) 1)) (= (x_count 2) (x_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) x)) (= (x_count 3) (+ (x_count 2) 1)) (= (x_count 3) (x_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) x)) (= (x_count 4) (+ (x_count 3) 1)) (= (x_count 4) (x_count 3))))) (and (or (or (or (or (= fmt1 (+ fmt0 0)) (= fmt1 (+ fmt0 1))) (= fmt1 (+ fmt0 2))) (= fmt1 (+ fmt0 3))) (= fmt1 (+ fmt0 4))) (or (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1)))))) :rule and :premises (a0) :args (0))
(step t62 (cl (and (or (or (or (or (= fmt1 (+ fmt0 0)) (= fmt1 (+ fmt0 1))) (= fmt1 (+ fmt0 2))) (= fmt1 (+ fmt0 3))) (= fmt1 (+ fmt0 4))) (or (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1))))) :rule and :premises (t61) :args (1))
(step t63 (cl (or (= arg1 (+ fmt0 0)) (= arg1 (+ fmt0 1)))) :rule and :premises (t62) :args (1))
(step t64 (cl (or (= fmt0 arg1) (= fmt0 (+ -1 arg1)))) :rule resolution :premises (t50 t60 t63))
(step t65 (cl (= fmt0 arg1) (= fmt0 (+ -1 arg1))) :rule or :premises (t64))
(step t66 (cl (=> (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1)) (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) :rule implies_neg1)
(anchor :step t67)
(assume t67.a0 (= distance 6))
(assume t67.a1 (= fmt0 0))
(assume t67.a2 (= fmt0 (+ arg0 distance)))
(assume t67.a3 (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))
(assume t67.a4 (= fmt0 arg1))
(step t67.t0 (cl (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)) (not (= distance 6)) (not (= (+ arg0 distance) 0)) (not (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (not (= arg1 0))) :rule and_neg)
(step t67.t1 (cl (and (and (and (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255))) (and (and (and (and (and (and (and (and (and (ite (and (= (select_format 0) percent) (= (select_format 1) s)) (= (s_count 0) 1) (= (s_count 0) 0)) (ite (and (= (select_format 1) percent) (= (select_format 2) s)) (= (s_count 1) (+ (s_count 0) 1)) (= (s_count 1) (s_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) s)) (= (s_count 2) (+ (s_count 1) 1)) (= (s_count 2) (s_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) s)) (= (s_count 3) (+ (s_count 2) 1)) (= (s_count 3) (s_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) s)) (= (s_count 4) (+ (s_count 3) 1)) (= (s_count 4) (s_count 3)))) (ite (and (= (select_format 0) percent) (= (select_format 1) x)) (= (x_count 0) 1) (= (x_count 0) 0))) (ite (and (= (select_format 1) percent) (= (select_format 2) x)) (= (x_count 1) (+ (x_count 0) 1)) (= (x_count 1) (x_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) x)) (= (x_count 2) (+ (x_count 1) 1)) (= (x_count 2) (x_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) x)) (= (x_count 3) (+ (x_count 2) 1)) (= (x_count 3) (x_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) x)) (= (x_count 4) (+ (x_count 3) 1)) (= (x_count 4) (x_count 3)))))) :rule and :premises (t61) :args (0))
(step t67.t2 (cl (and (and (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255)))) :rule and :premises (t67.t1) :args (0))
(step t67.t3 (cl (and (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))))) :rule and :premises (t67.t2) :args (0))
(step t67.t4 (cl (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120))) :rule and :premises (t67.t3) :args (0))
(step t67.t5 (cl (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115))) :rule and :premises (t67.t4) :args (0))
(step t67.t6 (cl (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37))) :rule and :premises (t67.t5) :args (0))
(step t67.t7 (cl (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1))) :rule and :premises (t67.t6) :args (0))
(step t67.t8 (cl (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5))) :rule and :premises (t67.t7) :args (0))
(step t67.t9 (cl (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2))) :rule and :premises (t67.t8) :args (0))
(step t67.t10 (cl (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2))) :rule and :premises (t67.t9) :args (0))
(step t67.t11 (cl (and (= distance 6) (= fmt_length 5))) :rule and :premises (t67.t10) :args (0))
(step t67.t12 (cl (= distance 6)) :rule and :premises (t67.t11) :args (0))
(step t67.t13 (cl (not (= (= arg0 (- fmt0 distance)) (= fmt0 (+ arg0 distance)))) (not (= arg0 (- fmt0 distance))) (= fmt0 (+ arg0 distance))) :rule equiv_pos2)
(step t67.t14 (cl (= arg0 arg0)) :rule refl)
(step t67.t15 (cl (= (- fmt0 distance) (+ fmt0 (* -1 distance)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- fmt0 distance) (+ fmt0 (* -1 distance))) 3 6))
(step t67.t16 (cl (= (= arg0 (- fmt0 distance)) (= arg0 (+ fmt0 (* -1 distance))))) :rule cong :premises (t67.t14 t67.t15))
(step t67.t17 (cl (= (= arg0 (+ fmt0 (* -1 distance))) (= fmt0 (+ arg0 distance)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= arg0 (+ fmt0 (* -1 distance))) (= fmt0 (+ arg0 distance))) 3 7))
(step t67.t18 (cl (= (= arg0 (- fmt0 distance)) (= fmt0 (+ arg0 distance)))) :rule trans :premises (t67.t16 t67.t17))
(step t67.t19 (cl (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) :rule and :premises (t67.t3) :args (1))
(step t67.t20 (cl (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)))) :rule and :premises (t67.t19) :args (0))
(step t67.t21 (cl (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance)))) :rule and :premises (t67.t20) :args (0))
(step t67.t22 (cl (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1)))) :rule and :premises (t67.t21) :args (0))
(step t67.t23 (cl (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1)))) :rule and :premises (t67.t22) :args (0))
(step t67.t24 (cl (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0))) :rule and :premises (t67.t23) :args (0))
(step t67.t25 (cl (and (= fmt0 0) (= arg0 (- fmt0 distance)))) :rule and :premises (t67.t24) :args (0))
(step t67.t26 (cl (= arg0 (- fmt0 distance))) :rule and :premises (t67.t25) :args (1))
(step t67.t27 (cl (= fmt0 (+ arg0 distance))) :rule resolution :premises (t67.t13 t67.t18 t67.t26))
(step t67.t28 (cl (= (+ arg0 distance) fmt0)) :rule symm :premises (t67.t27))
(step t67.t29 (cl (= fmt0 0)) :rule and :premises (t67.t25) :args (0))
(step t67.t30 (cl (= (+ arg0 distance) 0)) :rule trans :premises (t67.t28 t67.t29))
(step t67.t31 (cl (not (= (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) (not (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule equiv_pos2)
(step t67.t32 (cl (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0)) 3 6))
(step t67.t33 (cl (= 4 4)) :rule refl)
(step t67.t34 (cl (= (- (- fmt1 2) fmt0) (+ (- fmt1 2) (* -1 fmt0)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- (- fmt1 2) fmt0) (+ (- fmt1 2) (* -1 fmt0))) 3 6))
(step t67.t35 (cl (= (- fmt1 2) (+ fmt1 (* -1 2)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- fmt1 2) (+ fmt1 (* -1 2))) 3 6))
(step t67.t36 (cl (= fmt1 fmt1)) :rule refl)
(step t67.t37 (cl (= (* -1 2) -2)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (* -1 2) -2) 3 7))
(step t67.t38 (cl (= (+ fmt1 (* -1 2)) (+ fmt1 -2))) :rule cong :premises (t67.t36 t67.t37))
(step t67.t39 (cl (= (+ fmt1 -2) (+ -2 fmt1))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ fmt1 -2) (+ -2 fmt1)) 3 7))
(step t67.t40 (cl (= (+ fmt1 (* -1 2)) (+ -2 fmt1))) :rule trans :premises (t67.t38 t67.t39))
(step t67.t41 (cl (= (- fmt1 2) (+ -2 fmt1))) :rule trans :premises (t67.t35 t67.t40))
(step t67.t42 (cl (= (* -1 fmt0) (* -1 fmt0))) :rule refl)
(step t67.t43 (cl (= (+ (- fmt1 2) (* -1 fmt0)) (+ (+ -2 fmt1) (* -1 fmt0)))) :rule cong :premises (t67.t41 t67.t42))
(step t67.t44 (cl (= (+ (+ -2 fmt1) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ -2 fmt1) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0))) 3 7))
(step t67.t45 (cl (= (+ (- fmt1 2) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0)))) :rule trans :premises (t67.t43 t67.t44))
(step t67.t46 (cl (= (- (- fmt1 2) fmt0) (+ -2 fmt1 (* -1 fmt0)))) :rule trans :premises (t67.t34 t67.t45))
(step t67.t47 (cl (= (x_count (- (- fmt1 2) fmt0)) (x_count (+ -2 fmt1 (* -1 fmt0))))) :rule cong :premises (t67.t46))
(step t67.t48 (cl (= (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))))) :rule cong :premises (t67.t33 t67.t47))
(step t67.t49 (cl (= (s_count (- (- fmt1 2) fmt0)) (s_count (+ -2 fmt1 (* -1 fmt0))))) :rule cong :premises (t67.t46))
(step t67.t50 (cl (= (* 4 (s_count (- (- fmt1 2) fmt0))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) :rule cong :premises (t67.t33 t67.t49))
(step t67.t51 (cl (= (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0) (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0))) :rule cong :premises (t67.t48 t67.t50 t67.t14))
(step t67.t52 (cl (= (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 3 7))
(step t67.t53 (cl (= (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule trans :premises (t67.t51 t67.t52))
(step t67.t54 (cl (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule trans :premises (t67.t32 t67.t53))
(step t67.t55 (cl (= (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule cong :premises (t51 t67.t54))
(step t67.t56 (cl (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))) :rule and :premises (t67.t19) :args (1))
(step t67.t57 (cl (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule resolution :premises (t67.t31 t67.t55 t67.t56))
(step t67.t58 (cl (=> (and (= fmt0 0) (= fmt0 arg1)) (= arg1 0)) (and (= fmt0 0) (= fmt0 arg1))) :rule implies_neg1)
(anchor :step t67.t59)
(assume t67.t59.a0 (= fmt0 0))
(assume t67.t59.a1 (= fmt0 arg1))
(step t67.t59.t0 (cl (= arg1 fmt0)) :rule symm :premises (t67.t59.a1))
(step t67.t59.t1 (cl (= arg1 0)) :rule trans :premises (t67.t59.t0 t67.t29))
(step t67.t59 (cl (not (= fmt0 0)) (not (= fmt0 arg1)) (= arg1 0)) :rule subproof :discharge (t67.t59.a0 t67.t59.a1))
(step t67.t60 (cl (not (and (= fmt0 0) (= fmt0 arg1))) (= fmt0 0)) :rule and_pos :args (0))
(step t67.t61 (cl (not (and (= fmt0 0) (= fmt0 arg1))) (= fmt0 arg1)) :rule and_pos :args (1))
(step t67.t62 (cl (= arg1 0) (not (and (= fmt0 0) (= fmt0 arg1))) (not (and (= fmt0 0) (= fmt0 arg1)))) :rule resolution :premises (t67.t59 t67.t60 t67.t61))
(step t67.t63 (cl (not (and (= fmt0 0) (= fmt0 arg1))) (not (and (= fmt0 0) (= fmt0 arg1))) (= arg1 0)) :rule reordering :premises (t67.t62))
(step t67.t64 (cl (not (and (= fmt0 0) (= fmt0 arg1))) (= arg1 0)) :rule contraction :premises (t67.t63))
(step t67.t65 (cl (=> (and (= fmt0 0) (= fmt0 arg1)) (= arg1 0)) (= arg1 0)) :rule resolution :premises (t67.t58 t67.t64))
(step t67.t66 (cl (=> (and (= fmt0 0) (= fmt0 arg1)) (= arg1 0)) (not (= arg1 0))) :rule implies_neg2)
(step t67.t67 (cl (=> (and (= fmt0 0) (= fmt0 arg1)) (= arg1 0)) (=> (and (= fmt0 0) (= fmt0 arg1)) (= arg1 0))) :rule resolution :premises (t67.t65 t67.t66))
(step t67.t68 (cl (=> (and (= fmt0 0) (= fmt0 arg1)) (= arg1 0))) :rule contraction :premises (t67.t67))
(step t67.t69 (cl (not (and (= fmt0 0) (= fmt0 arg1))) (= arg1 0)) :rule implies :premises (t67.t68))
(step t67.t70 (cl (and (= fmt0 0) (= fmt0 arg1)) (not (= fmt0 0)) (not (= fmt0 arg1))) :rule and_neg)
(step t67.t71 (cl (and (= fmt0 0) (= fmt0 arg1))) :rule resolution :premises (t67.t70 t67.t29 t67.a4))
(step t67.t72 (cl (= arg1 0)) :rule resolution :premises (t67.t69 t67.t71))
(step t67.t73 (cl (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) :rule resolution :premises (t67.t0 t67.t12 t67.t30 t67.t57 t67.t72))
(step t67 (cl (not (= distance 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (not (= fmt0 arg1)) (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) :rule subproof :discharge (t67.a0 t67.a1 t67.a2 t67.a3 t67.a4))
(step t68 (cl (not (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (= distance 6)) :rule and_pos :args (0))
(step t69 (cl (not (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (= fmt0 0)) :rule and_pos :args (1))
(step t70 (cl (not (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (= fmt0 (+ arg0 distance))) :rule and_pos :args (2))
(step t71 (cl (not (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule and_pos :args (3))
(step t72 (cl (not (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (= fmt0 arg1)) :rule and_pos :args (4))
(step t73 (cl (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)) (not (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (not (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (not (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (not (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (not (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1)))) :rule resolution :premises (t67 t68 t69 t70 t71 t72))
(step t74 (cl (not (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (not (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (not (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (not (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (not (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) :rule reordering :premises (t73))
(step t75 (cl (not (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) :rule contraction :premises (t74))
(step t76 (cl (=> (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1)) (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) :rule resolution :premises (t66 t75))
(step t77 (cl (=> (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1)) (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (not (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)))) :rule implies_neg2)
(step t78 (cl (=> (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1)) (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (=> (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1)) (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)))) :rule resolution :premises (t76 t77))
(step t79 (cl (=> (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1)) (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)))) :rule contraction :premises (t78))
(step t80 (cl (not (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1))) (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) :rule implies :premises (t79))
(step t81 (cl (=> (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)) false) (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) :rule implies_neg1)
(anchor :step t82)
(assume t82.a0 (= distance 6))
(assume t82.a1 (= (+ arg0 distance) 0))
(assume t82.a2 (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))
(assume t82.a3 (= arg1 0))
(step t82.t0 (cl (not (= (= (* -1 (+ 6 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 0) false)) (not (= (* -1 (+ 6 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 0)) false) :rule equiv_pos2)
(step t82.t1 (cl (= (* -1 (+ 6 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (+ -6 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (* -1 (+ 6 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (+ -6 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 3 7))
(step t82.t2 (cl (= 0 0)) :rule refl)
(step t82.t3 (cl (= (= (* -1 (+ 6 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 0) (= (+ -6 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))) 0))) :rule cong :premises (t82.t1 t82.t2))
(step t82.t4 (cl (= (= (+ -6 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))) 0) false)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= (+ -6 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))) 0) false) 3 7))
(step t82.t5 (cl (= (= (* -1 (+ 6 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 0) false)) :rule trans :premises (t82.t3 t82.t4))
(step t82.t6 (cl (not (= (= arg1 0) (= (* -1 (+ 6 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 0))) (not (= arg1 0)) (= (* -1 (+ 6 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 0)) :rule equiv_pos2)
(step t82.t7 (cl (not (= (= arg1 (+ (* -1 6) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 (* -1 (+ 6 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0))))))))) (not (= arg1 (+ (* -1 6) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (= arg1 (* -1 (+ 6 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule equiv_pos2)
(step t82.t8 (cl (= (* -1 6) -6)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (* -1 6) -6) 3 7))
(step t82.t9 (cl (= (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))))) :rule refl)
(step t82.t10 (cl (= (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) :rule refl)
(step t82.t11 (cl (= (+ (* -1 6) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))) (+ -6 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule cong :premises (t82.t8 t82.t9 t82.t10))
(step t82.t12 (cl (= (= arg1 (+ (* -1 6) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 (+ -6 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule cong :premises (t51 t82.t11))
(step t82.t13 (cl (= (= arg1 (* -1 (+ 6 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (= arg1 (+ -6 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule cong :premises (t51 t82.t1))
(step t82.t14 (cl (= (= arg1 (+ -6 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 (* -1 (+ 6 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0))))))))) :rule symm :premises (t82.t13))
(step t82.t15 (cl (= (= arg1 (+ (* -1 6) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 (* -1 (+ 6 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0))))))))) :rule trans :premises (t82.t12 t82.t14))
(step t82.t16 (cl (not (= (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 (+ (* -1 6) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) (not (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (= arg1 (+ (* -1 6) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule equiv_pos2)
(step t82.t17 (cl (not (= (= (+ arg0 (* -1 -6)) 0) (= arg0 (* -1 6)))) (not (= (+ arg0 (* -1 -6)) 0)) (= arg0 (* -1 6))) :rule equiv_pos2)
(step t82.t18 (cl (= arg0 arg0)) :rule refl)
(step t82.t19 (cl (= (* -1 -6) 6)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (* -1 -6) 6) 3 7))
(step t82.t20 (cl (= (+ arg0 (* -1 -6)) (+ arg0 6))) :rule cong :premises (t82.t18 t82.t19))
(step t82.t21 (cl (= (+ arg0 6) (+ 6 arg0))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ arg0 6) (+ 6 arg0)) 3 7))
(step t82.t22 (cl (= (+ arg0 (* -1 -6)) (+ 6 arg0))) :rule trans :premises (t82.t20 t82.t21))
(step t82.t23 (cl (= (= (+ arg0 (* -1 -6)) 0) (= (+ 6 arg0) 0))) :rule cong :premises (t82.t22 t82.t2))
(step t82.t24 (cl (= (= (+ 6 arg0) 0) (= arg0 -6))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= (+ 6 arg0) 0) (= arg0 -6)) 3 7))
(step t82.t25 (cl (= (= (+ arg0 (* -1 -6)) 0) (= arg0 -6))) :rule trans :premises (t82.t23 t82.t24))
(step t82.t26 (cl (= (= arg0 (* -1 6)) (= arg0 -6))) :rule cong :premises (t82.t18 t82.t8))
(step t82.t27 (cl (= (= arg0 -6) (= arg0 (* -1 6)))) :rule symm :premises (t82.t26))
(step t82.t28 (cl (= (= (+ arg0 (* -1 -6)) 0) (= arg0 (* -1 6)))) :rule trans :premises (t82.t25 t82.t27))
(step t82.t29 (cl (not (= (= (+ arg0 distance) 0) (= (+ arg0 (* -1 -6)) 0))) (not (= (+ arg0 distance) 0)) (= (+ arg0 (* -1 -6)) 0)) :rule equiv_pos2)
(step t82.t30 (cl (not (= (= distance 6) (= distance (* -1 -6)))) (not (= distance 6)) (= distance (* -1 -6))) :rule equiv_pos2)
(step t82.t31 (cl (= distance distance)) :rule refl)
(step t82.t32 (cl (= (= distance (* -1 -6)) (= distance 6))) :rule cong :premises (t82.t31 t82.t19))
(step t82.t33 (cl (= (= distance 6) (= distance (* -1 -6)))) :rule symm :premises (t82.t32))
(step t82.t34 (cl (and (and (and (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255))) (and (and (and (and (and (and (and (and (and (ite (and (= (select_format 0) percent) (= (select_format 1) s)) (= (s_count 0) 1) (= (s_count 0) 0)) (ite (and (= (select_format 1) percent) (= (select_format 2) s)) (= (s_count 1) (+ (s_count 0) 1)) (= (s_count 1) (s_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) s)) (= (s_count 2) (+ (s_count 1) 1)) (= (s_count 2) (s_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) s)) (= (s_count 3) (+ (s_count 2) 1)) (= (s_count 3) (s_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) s)) (= (s_count 4) (+ (s_count 3) 1)) (= (s_count 4) (s_count 3)))) (ite (and (= (select_format 0) percent) (= (select_format 1) x)) (= (x_count 0) 1) (= (x_count 0) 0))) (ite (and (= (select_format 1) percent) (= (select_format 2) x)) (= (x_count 1) (+ (x_count 0) 1)) (= (x_count 1) (x_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) x)) (= (x_count 2) (+ (x_count 1) 1)) (= (x_count 2) (x_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) x)) (= (x_count 3) (+ (x_count 2) 1)) (= (x_count 3) (x_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) x)) (= (x_count 4) (+ (x_count 3) 1)) (= (x_count 4) (x_count 3)))))) :rule and :premises (t61) :args (0))
(step t82.t35 (cl (and (and (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255)))) :rule and :premises (t82.t34) :args (0))
(step t82.t36 (cl (and (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))))) :rule and :premises (t82.t35) :args (0))
(step t82.t37 (cl (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120))) :rule and :premises (t82.t36) :args (0))
(step t82.t38 (cl (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115))) :rule and :premises (t82.t37) :args (0))
(step t82.t39 (cl (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37))) :rule and :premises (t82.t38) :args (0))
(step t82.t40 (cl (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1))) :rule and :premises (t82.t39) :args (0))
(step t82.t41 (cl (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5))) :rule and :premises (t82.t40) :args (0))
(step t82.t42 (cl (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2))) :rule and :premises (t82.t41) :args (0))
(step t82.t43 (cl (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2))) :rule and :premises (t82.t42) :args (0))
(step t82.t44 (cl (and (= distance 6) (= fmt_length 5))) :rule and :premises (t82.t43) :args (0))
(step t82.t45 (cl (= distance 6)) :rule and :premises (t82.t44) :args (0))
(step t82.t46 (cl (= distance (* -1 -6))) :rule resolution :premises (t82.t30 t82.t33 t82.t45))
(step t82.t47 (cl (= (+ arg0 distance) (+ arg0 (* -1 -6)))) :rule cong :premises (t82.t18 t82.t46))
(step t82.t48 (cl (= (= (+ arg0 distance) 0) (= (+ arg0 (* -1 -6)) 0))) :rule cong :premises (t82.t47 t82.t2))
(step t82.t49 (cl (not (= (= arg0 (- fmt0 distance)) (= fmt0 (+ arg0 distance)))) (not (= arg0 (- fmt0 distance))) (= fmt0 (+ arg0 distance))) :rule equiv_pos2)
(step t82.t50 (cl (= (- fmt0 distance) (+ fmt0 (* -1 distance)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- fmt0 distance) (+ fmt0 (* -1 distance))) 3 6))
(step t82.t51 (cl (= (= arg0 (- fmt0 distance)) (= arg0 (+ fmt0 (* -1 distance))))) :rule cong :premises (t82.t18 t82.t50))
(step t82.t52 (cl (= (= arg0 (+ fmt0 (* -1 distance))) (= fmt0 (+ arg0 distance)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= arg0 (+ fmt0 (* -1 distance))) (= fmt0 (+ arg0 distance))) 3 7))
(step t82.t53 (cl (= (= arg0 (- fmt0 distance)) (= fmt0 (+ arg0 distance)))) :rule trans :premises (t82.t51 t82.t52))
(step t82.t54 (cl (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) :rule and :premises (t82.t36) :args (1))
(step t82.t55 (cl (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)))) :rule and :premises (t82.t54) :args (0))
(step t82.t56 (cl (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance)))) :rule and :premises (t82.t55) :args (0))
(step t82.t57 (cl (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1)))) :rule and :premises (t82.t56) :args (0))
(step t82.t58 (cl (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1)))) :rule and :premises (t82.t57) :args (0))
(step t82.t59 (cl (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0))) :rule and :premises (t82.t58) :args (0))
(step t82.t60 (cl (and (= fmt0 0) (= arg0 (- fmt0 distance)))) :rule and :premises (t82.t59) :args (0))
(step t82.t61 (cl (= arg0 (- fmt0 distance))) :rule and :premises (t82.t60) :args (1))
(step t82.t62 (cl (= fmt0 (+ arg0 distance))) :rule resolution :premises (t82.t49 t82.t53 t82.t61))
(step t82.t63 (cl (= (+ arg0 distance) fmt0)) :rule symm :premises (t82.t62))
(step t82.t64 (cl (= fmt0 0)) :rule and :premises (t82.t60) :args (0))
(step t82.t65 (cl (= (+ arg0 distance) 0)) :rule trans :premises (t82.t63 t82.t64))
(step t82.t66 (cl (= (+ arg0 (* -1 -6)) 0)) :rule resolution :premises (t82.t29 t82.t48 t82.t65))
(step t82.t67 (cl (= arg0 (* -1 6))) :rule resolution :premises (t82.t17 t82.t28 t82.t66))
(step t82.t68 (cl (= (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))) (+ (* -1 6) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule cong :premises (t82.t67 t82.t9 t82.t10))
(step t82.t69 (cl (= (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 (+ (* -1 6) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule cong :premises (t51 t82.t68))
(step t82.t70 (cl (not (= (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) (not (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule equiv_pos2)
(step t82.t71 (cl (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0)) 3 6))
(step t82.t72 (cl (= 4 4)) :rule refl)
(step t82.t73 (cl (= (- (- fmt1 2) fmt0) (+ (- fmt1 2) (* -1 fmt0)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- (- fmt1 2) fmt0) (+ (- fmt1 2) (* -1 fmt0))) 3 6))
(step t82.t74 (cl (= (- fmt1 2) (+ fmt1 (* -1 2)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- fmt1 2) (+ fmt1 (* -1 2))) 3 6))
(step t82.t75 (cl (= fmt1 fmt1)) :rule refl)
(step t82.t76 (cl (= (* -1 2) -2)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (* -1 2) -2) 3 7))
(step t82.t77 (cl (= (+ fmt1 (* -1 2)) (+ fmt1 -2))) :rule cong :premises (t82.t75 t82.t76))
(step t82.t78 (cl (= (+ fmt1 -2) (+ -2 fmt1))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ fmt1 -2) (+ -2 fmt1)) 3 7))
(step t82.t79 (cl (= (+ fmt1 (* -1 2)) (+ -2 fmt1))) :rule trans :premises (t82.t77 t82.t78))
(step t82.t80 (cl (= (- fmt1 2) (+ -2 fmt1))) :rule trans :premises (t82.t74 t82.t79))
(step t82.t81 (cl (= (* -1 fmt0) (* -1 fmt0))) :rule refl)
(step t82.t82 (cl (= (+ (- fmt1 2) (* -1 fmt0)) (+ (+ -2 fmt1) (* -1 fmt0)))) :rule cong :premises (t82.t80 t82.t81))
(step t82.t83 (cl (= (+ (+ -2 fmt1) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ -2 fmt1) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0))) 3 7))
(step t82.t84 (cl (= (+ (- fmt1 2) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0)))) :rule trans :premises (t82.t82 t82.t83))
(step t82.t85 (cl (= (- (- fmt1 2) fmt0) (+ -2 fmt1 (* -1 fmt0)))) :rule trans :premises (t82.t73 t82.t84))
(step t82.t86 (cl (= (x_count (- (- fmt1 2) fmt0)) (x_count (+ -2 fmt1 (* -1 fmt0))))) :rule cong :premises (t82.t85))
(step t82.t87 (cl (= (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))))) :rule cong :premises (t82.t72 t82.t86))
(step t82.t88 (cl (= (s_count (- (- fmt1 2) fmt0)) (s_count (+ -2 fmt1 (* -1 fmt0))))) :rule cong :premises (t82.t85))
(step t82.t89 (cl (= (* 4 (s_count (- (- fmt1 2) fmt0))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) :rule cong :premises (t82.t72 t82.t88))
(step t82.t90 (cl (= (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0) (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0))) :rule cong :premises (t82.t87 t82.t89 t82.t18))
(step t82.t91 (cl (= (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 3 7))
(step t82.t92 (cl (= (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule trans :premises (t82.t90 t82.t91))
(step t82.t93 (cl (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule trans :premises (t82.t71 t82.t92))
(step t82.t94 (cl (= (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule cong :premises (t51 t82.t93))
(step t82.t95 (cl (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))) :rule and :premises (t82.t54) :args (1))
(step t82.t96 (cl (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule resolution :premises (t82.t70 t82.t94 t82.t95))
(step t82.t97 (cl (= arg1 (+ (* -1 6) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule resolution :premises (t82.t16 t82.t69 t82.t96))
(step t82.t98 (cl (= arg1 (* -1 (+ 6 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule resolution :premises (t82.t7 t82.t15 t82.t97))
(step t82.t99 (cl (= (= arg1 0) (= (* -1 (+ 6 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 0))) :rule cong :premises (t82.t98 t82.t2))
(step t82.t100 (cl (= (* -1 (+ 6 (* -4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* -4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 0)) :rule resolution :premises (t82.t6 t82.t99 t82.a3))
(step t82.t101 (cl false) :rule resolution :premises (t82.t0 t82.t5 t82.t100))
(step t82 (cl (not (= distance 6)) (not (= (+ arg0 distance) 0)) (not (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (not (= arg1 0)) false) :rule subproof :discharge (t82.a0 t82.a1 t82.a2 t82.a3))
(step t83 (cl (not (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (= distance 6)) :rule and_pos :args (0))
(step t84 (cl (not (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (= (+ arg0 distance) 0)) :rule and_pos :args (1))
(step t85 (cl (not (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule and_pos :args (2))
(step t86 (cl (not (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (= arg1 0)) :rule and_pos :args (3))
(step t87 (cl false (not (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (not (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (not (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (not (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)))) :rule resolution :premises (t82 t83 t84 t85 t86))
(step t88 (cl (not (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (not (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (not (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) (not (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) false) :rule reordering :premises (t87))
(step t89 (cl (not (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))) false) :rule contraction :premises (t88))
(step t90 (cl (=> (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)) false) false) :rule resolution :premises (t81 t89))
(step t91 (cl (=> (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)) false) (not false)) :rule implies_neg2)
(step t92 (cl (=> (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)) false) (=> (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)) false)) :rule resolution :premises (t90 t91))
(step t93 (cl (=> (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)) false)) :rule contraction :premises (t92))
(step t94 (cl (= (=> (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)) false) (not (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0))))) :rule implies_simplify)
(step t95 (cl (not (=> (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)) false)) (not (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)))) :rule equiv1 :premises (t94))
(step t96 (cl (not (and (= distance 6) (= (+ arg0 distance) 0) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= arg1 0)))) :rule resolution :premises (t93 t95))
(step t97 (cl (not (and (= distance 6) (= fmt0 0) (= fmt0 (+ arg0 distance)) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) (= fmt0 arg1)))) :rule resolution :premises (t80 t96))
(step t98 (cl (not (= distance 6)) (not (= fmt0 0)) (not (= fmt0 (+ arg0 distance))) (not (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) (not (= fmt0 arg1))) :rule not_and :premises (t97))
(step t99 (cl (and (and (and (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255))) (and (and (and (and (and (and (and (and (and (ite (and (= (select_format 0) percent) (= (select_format 1) s)) (= (s_count 0) 1) (= (s_count 0) 0)) (ite (and (= (select_format 1) percent) (= (select_format 2) s)) (= (s_count 1) (+ (s_count 0) 1)) (= (s_count 1) (s_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) s)) (= (s_count 2) (+ (s_count 1) 1)) (= (s_count 2) (s_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) s)) (= (s_count 3) (+ (s_count 2) 1)) (= (s_count 3) (s_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) s)) (= (s_count 4) (+ (s_count 3) 1)) (= (s_count 4) (s_count 3)))) (ite (and (= (select_format 0) percent) (= (select_format 1) x)) (= (x_count 0) 1) (= (x_count 0) 0))) (ite (and (= (select_format 1) percent) (= (select_format 2) x)) (= (x_count 1) (+ (x_count 0) 1)) (= (x_count 1) (x_count 0)))) (ite (and (= (select_format 2) percent) (= (select_format 3) x)) (= (x_count 2) (+ (x_count 1) 1)) (= (x_count 2) (x_count 1)))) (ite (and (= (select_format 3) percent) (= (select_format 4) x)) (= (x_count 3) (+ (x_count 2) 1)) (= (x_count 3) (x_count 2)))) (ite (and (= (select_format 4) percent) (= (select_format 5) x)) (= (x_count 4) (+ (x_count 3) 1)) (= (x_count 4) (x_count 3)))))) :rule and :premises (t61) :args (0))
(step t100 (cl (and (and (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (or (= (select_format 0) percent) (= (select_format 0) s)) (= (select_format 0) x)) (= (select_format 0) adr_lo)) (= (select_format 0) adr_medlo)) (= (select_format 0) adr_medhi)) (= (select_format 0) adr_hi)) (= (select_format 0) 255)) (= (select_format 1) percent)) (= (select_format 1) s)) (= (select_format 1) x)) (= (select_format 1) adr_lo)) (= (select_format 1) adr_medlo)) (= (select_format 1) adr_medhi)) (= (select_format 1) adr_hi)) (= (select_format 1) 255)) (= (select_format 2) percent)) (= (select_format 2) s)) (= (select_format 2) x)) (= (select_format 2) adr_lo)) (= (select_format 2) adr_medlo)) (= (select_format 2) adr_medhi)) (= (select_format 2) adr_hi)) (= (select_format 2) 255)) (= (select_format 3) percent)) (= (select_format 3) s)) (= (select_format 3) x)) (= (select_format 3) adr_lo)) (= (select_format 3) adr_medlo)) (= (select_format 3) adr_medhi)) (= (select_format 3) adr_hi)) (= (select_format 3) 255)) (= (select_format 4) percent)) (= (select_format 4) s)) (= (select_format 4) x)) (= (select_format 4) adr_lo)) (= (select_format 4) adr_medlo)) (= (select_format 4) adr_medhi)) (= (select_format 4) adr_hi)) (= (select_format 4) 255)))) :rule and :premises (t99) :args (0))
(step t101 (cl (and (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120)) (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))))) :rule and :premises (t100) :args (0))
(step t102 (cl (and (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115)) (= x 120))) :rule and :premises (t101) :args (0))
(step t103 (cl (and (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37)) (= s 115))) :rule and :premises (t102) :args (0))
(step t104 (cl (and (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1)) (= percent 37))) :rule and :premises (t103) :args (0))
(step t105 (cl (and (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5)) (= adr_hi 1))) :rule and :premises (t104) :args (0))
(step t106 (cl (and (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2)) (= adr_medhi 5))) :rule and :premises (t105) :args (0))
(step t107 (cl (and (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2)) (= adr_medlo 2))) :rule and :premises (t106) :args (0))
(step t108 (cl (and (and (= distance 6) (= fmt_length 5)) (= adr_lo 2))) :rule and :premises (t107) :args (0))
(step t109 (cl (and (= distance 6) (= fmt_length 5))) :rule and :premises (t108) :args (0))
(step t110 (cl (= distance 6)) :rule and :premises (t109) :args (0))
(step t111 (cl (and (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))))) :rule and :premises (t101) :args (1))
(step t112 (cl (and (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance))) (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)))) :rule and :premises (t111) :args (0))
(step t113 (cl (and (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1))) (>= arg1 (+ arg0 distance)))) :rule and :premises (t112) :args (0))
(step t114 (cl (and (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1))) (> fmt1 (+ fmt0 1)))) :rule and :premises (t113) :args (0))
(step t115 (cl (and (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0)) (< fmt1 (- (+ fmt0 fmt_length) 1)))) :rule and :premises (t114) :args (0))
(step t116 (cl (and (and (= fmt0 0) (= arg0 (- fmt0 distance))) (>= arg1 fmt0))) :rule and :premises (t115) :args (0))
(step t117 (cl (and (= fmt0 0) (= arg0 (- fmt0 distance)))) :rule and :premises (t116) :args (0))
(step t118 (cl (= fmt0 0)) :rule and :premises (t117) :args (0))
(step t119 (cl (not (= (= arg0 (- fmt0 distance)) (= fmt0 (+ arg0 distance)))) (not (= arg0 (- fmt0 distance))) (= fmt0 (+ arg0 distance))) :rule equiv_pos2)
(step t120 (cl (= arg0 arg0)) :rule refl)
(step t121 (cl (= (- fmt0 distance) (+ fmt0 (* -1 distance)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- fmt0 distance) (+ fmt0 (* -1 distance))) 3 6))
(step t122 (cl (= (= arg0 (- fmt0 distance)) (= arg0 (+ fmt0 (* -1 distance))))) :rule cong :premises (t120 t121))
(step t123 (cl (= (= arg0 (+ fmt0 (* -1 distance))) (= fmt0 (+ arg0 distance)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (= arg0 (+ fmt0 (* -1 distance))) (= fmt0 (+ arg0 distance))) 3 7))
(step t124 (cl (= (= arg0 (- fmt0 distance)) (= fmt0 (+ arg0 distance)))) :rule trans :premises (t122 t123))
(step t125 (cl (= arg0 (- fmt0 distance))) :rule and :premises (t117) :args (1))
(step t126 (cl (= fmt0 (+ arg0 distance))) :rule resolution :premises (t119 t124 t125))
(step t127 (cl (not (= (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) (not (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule equiv_pos2)
(step t128 (cl (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0)) 3 6))
(step t129 (cl (= 4 4)) :rule refl)
(step t130 (cl (= (- (- fmt1 2) fmt0) (+ (- fmt1 2) (* -1 fmt0)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- (- fmt1 2) fmt0) (+ (- fmt1 2) (* -1 fmt0))) 3 6))
(step t131 (cl (= (- fmt1 2) (+ fmt1 (* -1 2)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- fmt1 2) (+ fmt1 (* -1 2))) 3 6))
(step t132 (cl (= fmt1 fmt1)) :rule refl)
(step t133 (cl (= (* -1 2) -2)) :rule hole :args ("TRUST_THEORY_REWRITE" (= (* -1 2) -2) 3 7))
(step t134 (cl (= (+ fmt1 (* -1 2)) (+ fmt1 -2))) :rule cong :premises (t132 t133))
(step t135 (cl (= (+ fmt1 -2) (+ -2 fmt1))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ fmt1 -2) (+ -2 fmt1)) 3 7))
(step t136 (cl (= (+ fmt1 (* -1 2)) (+ -2 fmt1))) :rule trans :premises (t134 t135))
(step t137 (cl (= (- fmt1 2) (+ -2 fmt1))) :rule trans :premises (t131 t136))
(step t138 (cl (= (* -1 fmt0) (* -1 fmt0))) :rule refl)
(step t139 (cl (= (+ (- fmt1 2) (* -1 fmt0)) (+ (+ -2 fmt1) (* -1 fmt0)))) :rule cong :premises (t137 t138))
(step t140 (cl (= (+ (+ -2 fmt1) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ -2 fmt1) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0))) 3 7))
(step t141 (cl (= (+ (- fmt1 2) (* -1 fmt0)) (+ -2 fmt1 (* -1 fmt0)))) :rule trans :premises (t139 t140))
(step t142 (cl (= (- (- fmt1 2) fmt0) (+ -2 fmt1 (* -1 fmt0)))) :rule trans :premises (t130 t141))
(step t143 (cl (= (x_count (- (- fmt1 2) fmt0)) (x_count (+ -2 fmt1 (* -1 fmt0))))) :rule cong :premises (t142))
(step t144 (cl (= (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))))) :rule cong :premises (t129 t143))
(step t145 (cl (= (s_count (- (- fmt1 2) fmt0)) (s_count (+ -2 fmt1 (* -1 fmt0))))) :rule cong :premises (t142))
(step t146 (cl (= (* 4 (s_count (- (- fmt1 2) fmt0))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) :rule cong :premises (t129 t145))
(step t147 (cl (= (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0) (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0))) :rule cong :premises (t144 t146 t120))
(step t148 (cl (= (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))) 3 7))
(step t149 (cl (= (+ (* 4 (x_count (- (- fmt1 2) fmt0))) (* 4 (s_count (- (- fmt1 2) fmt0))) arg0) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule trans :premises (t147 t148))
(step t150 (cl (= (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))) (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule trans :premises (t128 t149))
(step t151 (cl (= (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0))))) (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0)))))))) :rule cong :premises (t51 t150))
(step t152 (cl (= arg1 (+ (+ arg0 (* 4 (s_count (- (- fmt1 2) fmt0)))) (* 4 (x_count (- (- fmt1 2) fmt0)))))) :rule and :premises (t111) :args (1))
(step t153 (cl (= arg1 (+ arg0 (* 4 (x_count (+ -2 fmt1 (* -1 fmt0)))) (* 4 (s_count (+ -2 fmt1 (* -1 fmt0))))))) :rule resolution :premises (t127 t151 t152))
(step t154 (cl (not (= fmt0 arg1))) :rule resolution :premises (t98 t110 t118 t126 t153))
(step t155 (cl (= fmt0 (+ -1 arg1))) :rule resolution :premises (t65 t154))
(step t156 (cl (not (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) (not (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule equiv_pos2)
(step t157 (cl (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4))))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4)))) 3 6))
(step t158 (cl (= (- (+ (+ arg0 distance) fmt_length) 4) (+ (+ (+ arg0 distance) fmt_length) (* -1 4)))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (- (+ (+ arg0 distance) fmt_length) 4) (+ (+ (+ arg0 distance) fmt_length) (* -1 4))) 3 6))
(step t159 (cl (= (+ (+ (+ arg0 distance) fmt_length) (* -1 4)) (+ distance fmt_length arg0 -4))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ (+ (+ arg0 distance) fmt_length) (* -1 4)) (+ distance fmt_length arg0 -4)) 3 6))
(step t160 (cl (= (+ distance fmt_length arg0 -4) (+ -4 arg0 fmt_length distance))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (+ distance fmt_length arg0 -4) (+ -4 arg0 fmt_length distance)) 3 7))
(step t161 (cl (= (+ (+ (+ arg0 distance) fmt_length) (* -1 4)) (+ -4 arg0 fmt_length distance))) :rule trans :premises (t159 t160))
(step t162 (cl (= (- (+ (+ arg0 distance) fmt_length) 4) (+ -4 arg0 fmt_length distance))) :rule trans :premises (t158 t161))
(step t163 (cl (= (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (>= arg1 (+ -4 arg0 fmt_length distance)))) :rule cong :premises (t51 t162))
(step t164 (cl (= (>= arg1 (+ -4 arg0 fmt_length distance)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule hole :args ("TRUST_THEORY_REWRITE" (= (>= arg1 (+ -4 arg0 fmt_length distance)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)) 3 7))
(step t165 (cl (= (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule trans :premises (t163 t164))
(step t166 (cl (= (not (>= arg1 (- (+ (+ arg0 distance) fmt_length) 4))) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule cong :premises (t165))
(step t167 (cl (= (< arg1 (- (+ (+ arg0 distance) fmt_length) 4)) (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4)))) :rule trans :premises (t157 t166))
(step t168 (cl (< arg1 (- (+ (+ arg0 distance) fmt_length) 4))) :rule and :premises (t112) :args (1))
(step t169 (cl (not (>= (+ arg1 (* -1 arg0) (* -1 fmt_length) (* -1 distance)) -4))) :rule resolution :premises (t156 t167 t168))
(step t170 (cl (= fmt_length 5)) :rule and :premises (t109) :args (1))
(step t171 (cl) :rule resolution :premises (t49 t155 t169 t126 t118 t170))
